 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/U7j4q7I1Ly/SOT_MQsEBL
% (1722632)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_MQsEBL for (1ds/27Mi)
% (1722636)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_MQsEBL for (1ds/18Mi)
% (1722631)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_MQsEBL for (1ds/4Mi)
Permission denied
% (1722630)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_MQsEBL for (1ds/183Mi)
% (1722635)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_MQsEBL for (1ds/275Mi)
% (1722633)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_MQsEBL for (1ds/2Mi)
% (1722634)Refutation not found, incomplete strategy
% (1722634)------------------------------
% (1722634)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (1722634)Termination reason: Refutation not found, incomplete strategy


% (1722634)Memory used [KB]: 5500
% (1722634)Time elapsed: 0.004 s
% (1722634)------------------------------
% (1722634)------------------------------
% (1722637)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_MQsEBL for (1ds/3Mi)
% (1722635)First to succeed.
% (1722630)Also succeeded, but the first one will report.
% (1722635)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_MQsEBL
% SZS output start Proof for SOT_MQsEBL
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Prohibition: m).
thf(func_def_2, type, accreln1: m > w > w > $o).
thf(func_def_3, type, accreln2: m > $i > w > w > $o).
thf(func_def_4, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_6, type, s__instance: $i > $i > $o).
thf(func_def_7, type, s__agent: $i > $i > w > $o).
thf(func_def_8, type, s__patient: $i > $i > w > $o).
thf(func_def_9, type, s__origin: $i > $i > w > $o).
thf(func_def_10, type, s__destination: $i > $i > w > $o).
thf(func_def_11, type, s__located: $i > $i > w > $o).
thf(func_def_20, type, sK0: w > $i > $i).
thf(func_def_21, type, sK1: w > w).
thf(func_def_23, type, sK3: w).
thf(func_def_24, type, sK4: w > $i).
thf(f155,plain,(
  $false),
  inference(subsumption_resolution,[status(thm)],[f154,f63])).
thf(f63,plain,(
  ($true = (accreln1 @ s__Prohibition @ cw @ sK3))),
  inference(cnf_transformation,[status(thm)],[f49])).
thf(f49,plain,(
  ($true = (accreln1 @ s__Prohibition @ cw @ sK3)) & ! [X1 : $i] : (($true != (s__instance @ X1 @ s__BusinessTrip)) | ($true != (s__agent @ X1 @ s__Bill @ sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f37,f48])).
thf(f48,plain,(
  ? [X0 : w] : (($true = (accreln1 @ s__Prohibition @ cw @ X0)) & ! [X1 : $i] : (($true != (s__instance @ X1 @ s__BusinessTrip)) | ($true != (s__agent @ X1 @ s__Bill @ X0)))) => (($true = (accreln1 @ s__Prohibition @ cw @ sK3)) & ! [X1 : $i] : (($true != (s__instance @ X1 @ s__BusinessTrip)) | ($true != (s__agent @ X1 @ s__Bill @ sK3))))),
  introduced(definition,[],[choice_axiom])).
thf(f37,plain,(
  ? [X0 : w] : (($true = (accreln1 @ s__Prohibition @ cw @ X0)) & ! [X1 : $i] : (($true != (s__instance @ X1 @ s__BusinessTrip)) | ($true != (s__agent @ X1 @ s__Bill @ X0))))),
  inference(ennf_transformation,[status(thm)],[f26])).
thf(f26,plain,(
  ~! [X0 : w] : (($true = (accreln1 @ s__Prohibition @ cw @ X0)) => ? [X1 : $i] : (($true = (s__instance @ X1 @ s__BusinessTrip)) & ($true = (s__agent @ X1 @ s__Bill @ X0))))),
  inference(fool_elimination,[status(thm)],[f25])).
thf(f25,plain,(
  ~! [X0 : w] : ((accreln1 @ s__Prohibition @ cw @ X0) => ? [X1 : $i] : ((s__agent @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__BusinessTrip)))),
  inference(rectify,[status(thm)],[f15])).
thf(f15,negated_conjecture,(
  ~! [X6 : w] : ((accreln1 @ s__Prohibition @ cw @ X6) => ? [X5 : $i] : ((s__agent @ X5 @ s__Bill @ X6) & (s__instance @ X5 @ s__BusinessTrip)))),
  inference(negated_conjecture,[status(cth)],[f14])).
thf(f14,conjecture,(
  ! [X6 : w] : ((accreln1 @ s__Prohibition @ cw @ X6) => ? [X5 : $i] : ((s__agent @ X5 @ s__Bill @ X6) & (s__instance @ X5 @ s__BusinessTrip)))),
  file('/tmp/U7j4q7I1Ly/SOT_MQsEBL',conjecture_not_go_on_trip)).
thf(f154,plain,(
  ($true != (accreln1 @ s__Prohibition @ cw @ sK3))),
  inference(subsumption_resolution,[status(thm)],[f153,f59])).
thf(f59,plain,(
  ($true = (s__instance @ sK2 @ s__OfficeBuilding))),
  inference(cnf_transformation,[status(thm)],[f47])).
thf(f47,plain,(
  ($true = (s__located @ s__Bill @ sK2 @ cw)) & ($true = (s__instance @ sK2 @ s__OfficeBuilding))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f20,f46])).
thf(f46,plain,(
  ? [X0 : $i] : (($true = (s__located @ s__Bill @ X0 @ cw)) & ($true = (s__instance @ X0 @ s__OfficeBuilding))) => (($true = (s__located @ s__Bill @ sK2 @ cw)) & ($true = (s__instance @ sK2 @ s__OfficeBuilding)))),
  introduced(definition,[],[choice_axiom])).
thf(f20,plain,(
  ? [X0 : $i] : (($true = (s__located @ s__Bill @ X0 @ cw)) & ($true = (s__instance @ X0 @ s__OfficeBuilding)))),
  inference(fool_elimination,[status(thm)],[f19])).
thf(f19,plain,(
  ? [X0 : $i] : ((s__instance @ X0 @ s__OfficeBuilding) & (s__located @ s__Bill @ X0 @ cw))),
  inference(rectify,[status(thm)],[f13])).
thf(f13,axiom,(
  ? [X3 : $i] : ((s__instance @ X3 @ s__OfficeBuilding) & (s__located @ s__Bill @ X3 @ cw))),
  file('/tmp/U7j4q7I1Ly/SOT_MQsEBL',fact_bill_at_office)).
thf(f153,plain,(
  ($true != (s__instance @ sK2 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Prohibition @ cw @ sK3))),
  inference(trivial_inequality_removal,[status(thm)],[f152])).
thf(f152,plain,(
  ($true != (s__instance @ sK2 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Prohibition @ cw @ sK3)) | ($true != $true)),
  inference(superposition,[status(thm)],[f142,f60])).
thf(f60,plain,(
  ($true = (s__located @ s__Bill @ sK2 @ cw))),
  inference(cnf_transformation,[status(thm)],[f47])).
thf(f142,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__located @ s__Bill @ X1 @ X0)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ((accreln1 @ s__Prohibition @ X0 @ sK3) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f138])).
thf(f138,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__located @ s__Bill @ X1 @ X0)) | ((accreln1 @ s__Prohibition @ X0 @ sK3) != $true) | ($true != $true) | ($true != (s__instance @ X1 @ s__OfficeBuilding))) )),
  inference(superposition,[status(thm)],[f119,f63])).
thf(f119,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Prohibition @ cw @ X0)) | ((accreln1 @ s__Prohibition @ X2 @ X0) != $true) | ($true != (s__located @ s__Bill @ X1 @ X2)) | ($true != (s__instance @ X1 @ s__OfficeBuilding))) )),
  inference(subsumption_resolution,[status(thm)],[f118,f57])).
thf(f57,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (($true = (s__instance @ (sK0 @ X1 @ X0) @ s__Smoking)) | ($true != (accreln1 @ s__Prohibition @ X2 @ X1)) | ($true != (s__instance @ X0 @ s__OfficeBuilding)) | ($true != (s__located @ s__Bill @ X0 @ X2))) )),
  inference(cnf_transformation,[status(thm)],[f43])).
thf(f43,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (($true != (accreln1 @ s__Prohibition @ X2 @ X1)) | ($true != (s__located @ s__Bill @ X0 @ X2)) | ($true != (s__instance @ X0 @ s__OfficeBuilding)) | (($true = (s__instance @ (sK0 @ X1 @ X0) @ s__Smoking)) & ($true = (s__located @ (sK0 @ X1 @ X0) @ X0 @ X1)) & ($true = (s__agent @ (sK0 @ X1 @ X0) @ s__Bill @ X1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f34,f42])).
thf(f42,plain,(
  ! [X0 : $i,X1 : w] : (? [X3 : $i] : (($true = (s__instance @ X3 @ s__Smoking)) & ($true = (s__located @ X3 @ X0 @ X1)) & ($true = (s__agent @ X3 @ s__Bill @ X1))) => (($true = (s__instance @ (sK0 @ X1 @ X0) @ s__Smoking)) & ($true = (s__located @ (sK0 @ X1 @ X0) @ X0 @ X1)) & ($true = (s__agent @ (sK0 @ X1 @ X0) @ s__Bill @ X1))))),
  introduced(definition,[],[choice_axiom])).
thf(f34,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (($true != (accreln1 @ s__Prohibition @ X2 @ X1)) | ($true != (s__located @ s__Bill @ X0 @ X2)) | ($true != (s__instance @ X0 @ s__OfficeBuilding)) | ? [X3 : $i] : (($true = (s__instance @ X3 @ s__Smoking)) & ($true = (s__located @ X3 @ X0 @ X1)) & ($true = (s__agent @ X3 @ s__Bill @ X1))))),
  inference(flattening,[status(thm)],[f33])).
thf(f33,plain,(
  ! [X0 : $i,X2 : w,X1 : w] : ((? [X3 : $i] : (($true = (s__instance @ X3 @ s__Smoking)) & ($true = (s__located @ X3 @ X0 @ X1)) & ($true = (s__agent @ X3 @ s__Bill @ X1))) | ($true != (accreln1 @ s__Prohibition @ X2 @ X1))) | (($true != (s__instance @ X0 @ s__OfficeBuilding)) | ($true != (s__located @ s__Bill @ X0 @ X2))))),
  inference(ennf_transformation,[status(thm)],[f30])).
thf(f30,plain,(
  ! [X0 : $i,X2 : w,X1 : w] : ((($true = (s__instance @ X0 @ s__OfficeBuilding)) & ($true = (s__located @ s__Bill @ X0 @ X2))) => (($true = (accreln1 @ s__Prohibition @ X2 @ X1)) => ? [X3 : $i] : (($true = (s__instance @ X3 @ s__Smoking)) & ($true = (s__located @ X3 @ X0 @ X1)) & ($true = (s__agent @ X3 @ s__Bill @ X1)))))),
  inference(fool_elimination,[status(thm)],[f29])).
thf(f29,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (((s__instance @ X0 @ s__OfficeBuilding) & (s__located @ s__Bill @ X0 @ X2)) => ((accreln1 @ s__Prohibition @ X2 @ X1) => ? [X3 : $i] : ((s__instance @ X3 @ s__Smoking) & (s__agent @ X3 @ s__Bill @ X1) & (s__located @ X3 @ X0 @ X1))))),
  inference(rectify,[status(thm)],[f11])).
thf(f11,axiom,(
  ! [X3 : $i,X1 : w,X0 : w] : (((s__instance @ X3 @ s__OfficeBuilding) & (s__located @ s__Bill @ X3 @ X0)) => ((accreln1 @ s__Prohibition @ X0 @ X1) => ? [X4 : $i] : ((s__instance @ X4 @ s__Smoking) & (s__agent @ X4 @ s__Bill @ X1) & (s__located @ X4 @ X3 @ X1))))),
  file('/tmp/U7j4q7I1Ly/SOT_MQsEBL',prohibition1_smoking_in_office)).
thf(f118,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Prohibition @ cw @ X0)) | ((accreln1 @ s__Prohibition @ X2 @ X0) != $true) | ($true != (s__located @ s__Bill @ X1 @ X2)) | ($true != (s__instance @ (sK0 @ X0 @ X1) @ s__Smoking)) | ($true != (s__instance @ X1 @ s__OfficeBuilding))) )),
  inference(subsumption_resolution,[status(thm)],[f117,f55])).
thf(f55,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (($true = (s__agent @ (sK0 @ X1 @ X0) @ s__Bill @ X1)) | ($true != (s__instance @ X0 @ s__OfficeBuilding)) | ($true != (s__located @ s__Bill @ X0 @ X2)) | ($true != (accreln1 @ s__Prohibition @ X2 @ X1))) )),
  inference(cnf_transformation,[status(thm)],[f43])).
thf(f117,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (s__located @ s__Bill @ X1 @ X2)) | ($true != (accreln1 @ s__Prohibition @ cw @ X0)) | ($true != (s__agent @ (sK0 @ X0 @ X1) @ s__Bill @ X0)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__instance @ (sK0 @ X0 @ X1) @ s__Smoking)) | ((accreln1 @ s__Prohibition @ X2 @ X0) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f116])).
thf(f116,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (s__agent @ (sK0 @ X0 @ X1) @ s__Bill @ X0)) | ($true != (s__instance @ (sK0 @ X0 @ X1) @ s__Smoking)) | ($true != (accreln1 @ s__Prohibition @ cw @ X0)) | ($true != (s__located @ s__Bill @ X1 @ X2)) | ($true != $true) | ((accreln1 @ s__Prohibition @ X2 @ X0) != $true) | ($true != (s__instance @ X1 @ s__OfficeBuilding))) )),
  inference(duplicate_literal_removal,[status(thm)],[f114])).
thf(f114,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__located @ s__Bill @ X1 @ X2)) | ($true != (s__agent @ (sK0 @ X0 @ X1) @ s__Bill @ X0)) | ($true != $true) | ($true != (accreln1 @ s__Prohibition @ cw @ X0)) | ($true != (s__instance @ (sK0 @ X0 @ X1) @ s__Smoking)) | ((accreln1 @ s__Prohibition @ X2 @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f102,f56])).
thf(f56,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (($true = (s__located @ (sK0 @ X1 @ X0) @ X0 @ X1)) | ($true != (s__located @ s__Bill @ X0 @ X2)) | ($true != (s__instance @ X0 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Prohibition @ X2 @ X1))) )),
  inference(cnf_transformation,[status(thm)],[f43])).
thf(f102,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != (s__located @ X0 @ X1 @ X2)) | ($true != (s__instance @ X0 @ s__Smoking)) | ((accreln1 @ s__Prohibition @ cw @ X2) != $true) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__agent @ X0 @ s__Bill @ X2))) )),
  inference(trivial_inequality_removal,[status(thm)],[f99])).
thf(f99,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != $true) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__instance @ X0 @ s__Smoking)) | ($true != (s__located @ X0 @ X1 @ X2)) | ($true != (s__agent @ X0 @ s__Bill @ X2)) | ((accreln1 @ s__Prohibition @ cw @ X2) != $true)) )),
  inference(superposition,[status(thm)],[f77,f63])).
thf(f77,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Prohibition @ X3 @ sK3)) | ($true != (s__instance @ X2 @ s__Smoking)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ($true != (accreln1 @ s__Prohibition @ X3 @ X0)) | ((s__located @ X2 @ X1 @ X0) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f75])).
thf(f75,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Prohibition @ X3 @ X0)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Prohibition @ X3 @ sK3)) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ($true != (s__instance @ X2 @ s__Smoking)) | ($true != $true) | ((s__located @ X2 @ X1 @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f73,f68])).
thf(f68,plain,(
  ( ! [X2 : w,X0 : w,X1 : w] : (((accreln1 @ s__Prohibition @ X1 @ X2) = $true) | ((accreln1 @ s__Prohibition @ X0 @ X2) != $true) | ((accreln1 @ s__Prohibition @ X0 @ X1) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f41])).
thf(f41,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Prohibition @ X1 @ X2) = $true) | ((accreln1 @ s__Prohibition @ X0 @ X2) != $true) | ((accreln1 @ s__Prohibition @ X0 @ X1) != $true))),
  inference(flattening,[status(thm)],[f40])).
thf(f40,plain,(
  ! [X1 : w,X2 : w,X0 : w] : (((accreln1 @ s__Prohibition @ X1 @ X2) = $true) | (((accreln1 @ s__Prohibition @ X0 @ X1) != $true) | ((accreln1 @ s__Prohibition @ X0 @ X2) != $true)))),
  inference(ennf_transformation,[status(thm)],[f18])).
thf(f18,plain,(
  ! [X1 : w,X2 : w,X0 : w] : ((((accreln1 @ s__Prohibition @ X0 @ X1) = $true) & ((accreln1 @ s__Prohibition @ X0 @ X2) = $true)) => ((accreln1 @ s__Prohibition @ X1 @ X2) = $true))),
  inference(fool_elimination,[status(thm)],[f17])).
thf(f17,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Prohibition @ X0 @ X2) & (accreln1 @ s__Prohibition @ X0 @ X1)) => (accreln1 @ s__Prohibition @ X1 @ X2))),
  inference(rectify,[status(thm)],[f3])).
thf(f3,axiom,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Prohibition @ X0 @ X2) & (accreln1 @ s__Prohibition @ X0 @ X1)) => (accreln1 @ s__Prohibition @ X1 @ X2))),
  file('/tmp/U7j4q7I1Ly/SOT_MQsEBL',euclid_prohibition)).
thf(f73,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != (accreln1 @ s__Prohibition @ X2 @ sK3)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__located @ X0 @ X1 @ X2)) | ($true != (s__instance @ X0 @ s__Smoking)) | ($true != (s__agent @ X0 @ s__Bill @ X2))) )),
  inference(subsumption_resolution,[status(thm)],[f72,f67])).
thf(f67,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__located @ X3 @ X2 @ X0)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__instance @ X3 @ s__Smoking)) | ($true = (s__instance @ (sK4 @ X1) @ s__BusinessTrip)) | ((accreln1 @ s__Prohibition @ X0 @ X1) != $true) | ($true != (s__agent @ X3 @ s__Bill @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f53])).
thf(f53,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (($true != (s__instance @ X3 @ s__Smoking)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__agent @ X3 @ s__Bill @ X0)) | ((accreln1 @ s__Prohibition @ X0 @ X1) != $true) | (($true = (s__instance @ (sK4 @ X1) @ s__BusinessTrip)) & ($true = (s__agent @ (sK4 @ X1) @ s__Bill @ X1))) | ($true != (s__located @ X3 @ X2 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f51,f52])).
thf(f52,plain,(
  ! [X1 : w] : (? [X4 : $i] : (($true = (s__instance @ X4 @ s__BusinessTrip)) & ((s__agent @ X4 @ s__Bill @ X1) = $true)) => (($true = (s__instance @ (sK4 @ X1) @ s__BusinessTrip)) & ($true = (s__agent @ (sK4 @ X1) @ s__Bill @ X1))))),
  introduced(definition,[],[choice_axiom])).
thf(f51,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (($true != (s__instance @ X3 @ s__Smoking)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__agent @ X3 @ s__Bill @ X0)) | ((accreln1 @ s__Prohibition @ X0 @ X1) != $true) | ? [X4 : $i] : (($true = (s__instance @ X4 @ s__BusinessTrip)) & ((s__agent @ X4 @ s__Bill @ X1) = $true)) | ($true != (s__located @ X3 @ X2 @ X0)))),
  inference(rectify,[status(thm)],[f36])).
thf(f36,plain,(
  ! [X1 : w,X3 : w,X2 : $i,X0 : $i] : (($true != (s__instance @ X0 @ s__Smoking)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__agent @ X0 @ s__Bill @ X1)) | ((accreln1 @ s__Prohibition @ X1 @ X3) != $true) | ? [X4 : $i] : (($true = (s__instance @ X4 @ s__BusinessTrip)) & ($true = (s__agent @ X4 @ s__Bill @ X3))) | ($true != (s__located @ X0 @ X2 @ X1)))),
  inference(flattening,[status(thm)],[f35])).
thf(f35,plain,(
  ! [X1 : w,X0 : $i,X3 : w,X2 : $i] : ((? [X4 : $i] : (($true = (s__instance @ X4 @ s__BusinessTrip)) & ($true = (s__agent @ X4 @ s__Bill @ X3))) | ((accreln1 @ s__Prohibition @ X1 @ X3) != $true)) | (($true != (s__located @ X0 @ X2 @ X1)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__agent @ X0 @ s__Bill @ X1)) | ($true != (s__instance @ X0 @ s__Smoking))))),
  inference(ennf_transformation,[status(thm)],[f24])).
thf(f24,plain,(
  ! [X1 : w,X0 : $i,X3 : w,X2 : $i] : ((($true = (s__located @ X0 @ X2 @ X1)) & ($true = (s__instance @ X2 @ s__OfficeBuilding)) & ($true = (s__agent @ X0 @ s__Bill @ X1)) & ($true = (s__instance @ X0 @ s__Smoking))) => (((accreln1 @ s__Prohibition @ X1 @ X3) = $true) => ? [X4 : $i] : (($true = (s__instance @ X4 @ s__BusinessTrip)) & ($true = (s__agent @ X4 @ s__Bill @ X3)))))),
  inference(fool_elimination,[status(thm)],[f23])).
thf(f23,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__instance @ X0 @ s__Smoking) & (s__located @ X0 @ X2 @ X1) & (s__agent @ X0 @ s__Bill @ X1) & (s__instance @ X2 @ s__OfficeBuilding)) => ((accreln1 @ s__Prohibition @ X1 @ X3) => ? [X4 : $i] : ((s__instance @ X4 @ s__BusinessTrip) & (s__agent @ X4 @ s__Bill @ X3))))),
  inference(rectify,[status(thm)],[f12])).
thf(f12,axiom,(
  ! [X4 : $i,X0 : w,X3 : $i,X1 : w] : (((s__instance @ X4 @ s__Smoking) & (s__located @ X4 @ X3 @ X0) & (s__agent @ X4 @ s__Bill @ X0) & (s__instance @ X3 @ s__OfficeBuilding)) => ((accreln1 @ s__Prohibition @ X0 @ X1) => ? [X5 : $i] : ((s__instance @ X5 @ s__BusinessTrip) & (s__agent @ X5 @ s__Bill @ X1))))),
  file('/tmp/U7j4q7I1Ly/SOT_MQsEBL',prohibition2_business_trip)).
thf(f72,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != (s__located @ X0 @ X1 @ X2)) | ($true != (s__agent @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Prohibition @ X2 @ sK3)) | ($true != (s__instance @ (sK4 @ sK3) @ s__BusinessTrip)) | ($true != (s__instance @ X0 @ s__Smoking))) )),
  inference(trivial_inequality_removal,[status(thm)],[f69])).
thf(f69,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != (accreln1 @ s__Prohibition @ X2 @ sK3)) | ($true != (s__instance @ (sK4 @ sK3) @ s__BusinessTrip)) | ($true != $true) | ($true != (s__located @ X0 @ X1 @ X2)) | ($true != (s__agent @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X0 @ s__Smoking)) | ($true != (s__instance @ X1 @ s__OfficeBuilding))) )),
  inference(superposition,[status(thm)],[f62,f66])).
thf(f66,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true = (s__agent @ (sK4 @ X1) @ s__Bill @ X1)) | ($true != (s__located @ X3 @ X2 @ X0)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__instance @ X3 @ s__Smoking)) | ($true != (s__agent @ X3 @ s__Bill @ X0)) | ((accreln1 @ s__Prohibition @ X0 @ X1) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f53])).
thf(f62,plain,(
  ( ! [X1 : $i] : (($true != (s__agent @ X1 @ s__Bill @ sK3)) | ($true != (s__instance @ X1 @ s__BusinessTrip))) )),
  inference(cnf_transformation,[status(thm)],[f49])).
% SZS output end Proof for SOT_MQsEBL
% (1722635)------------------------------
% (1722635)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (1722635)Termination reason: Refutation

% (1722635)Memory used [KB]: 5628
% (1722635)Time elapsed: 0.019 s
% (1722635)------------------------------
% (1722635)------------------------------
% (1722626)Success in time 0.032 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_MQsEBL - Vampire---5.0.1 says Theorem - CPU = 0.01 WC = 0.08 
% OUTPUT: SOT_MQsEBL - Vampire---5.0.1 says Refutation - CPU = 0.01 WC = 0.08 

