% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 1e_Obligation_existentialNHops
% SZS output start Proof for 1e_Obligation_existentialNHops
23. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ s__Bill @ X0) & (s__patient @ X2 @ X3 @ X0) & (s__instance @ X3 @ s__PayCheck) & (s__instance @ X2 @ s__Giving)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4] : ((s__origin @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Tax)))) [input(axiom)]
24. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [input(axiom)]
25. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [input(axiom)]
26. ? [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) & (accreln1 @ s__Obligation @ cw @ X0)) => ? [X2] : ((s__origin @ X2 @ s__Bill @ X1) & (s__instance @ X2 @ s__Tax))) [input(conjecture)]
27. ~? [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) & (accreln1 @ s__Obligation @ cw @ X0)) => ? [X2] : ((s__origin @ X2 @ s__Bill @ X1) & (s__instance @ X2 @ s__Tax))) [negated conjecture 26]
30. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ s__Bill @ X0) & (s__patient @ X2 @ X3 @ X0) & (s__instance @ X3 @ s__PayCheck) & (s__instance @ X2 @ s__Giving)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4] : ((s__origin @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Tax)))) [rectify 23]
31. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ s__Bill @ X0)) = $true) & (((s__patient @ X2 @ X3 @ X0)) = $true) & (((s__instance @ X3 @ s__PayCheck)) = $true) & (((s__instance @ X2 @ s__Giving)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)))) [fool elimination 30]
32. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [rectify 24]
33. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [fool elimination 32]
34. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [rectify 25]
35. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) [fool elimination 34]
36. ~? [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) & (accreln1 @ s__Obligation @ cw @ X0)) => ? [X2] : ((s__origin @ X2 @ s__Bill @ X1) & (s__instance @ X2 @ s__Tax))) [rectify 27]
37. ~? [X0 : w,X1 : w] : (((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) => ? [X2] : ((((s__origin @ X2 @ s__Bill @ X1)) = $true) & (((s__instance @ X2 @ s__Tax)) = $true))) [fool elimination 36]
59. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ s__Bill @ X0)) = $true) & (((s__patient @ X2 @ X3 @ X0)) = $true) & (((s__instance @ X3 @ s__PayCheck)) = $true) & (((s__instance @ X2 @ s__Giving)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)))) [flattening 31]
60. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [flattening 33]
61. ! [X0 : w,X1 : w,X2,X3] : ((? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true))) [ennf transformation 59]
62. ! [X0 : w,X1 : w,X2,X3] : (? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true)) [flattening 61]
63. ! [X0 : w,X1 : w,X2,X3] : ((? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true))) [ennf transformation 60]
64. ! [X0 : w,X1 : w,X2,X3] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [flattening 63]
65. ! [X0 : w,X1 : w] : (! [X2] : ((((s__origin @ X2 @ s__Bill @ X1)) != $true) | (((s__instance @ X2 @ s__Tax)) != $true)) & ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true))) [ennf transformation 37]
66. ! [X0 : w,X1 : w] : (! [X2] : ((((s__origin @ X2 @ s__Bill @ X1)) != $true) | (((s__instance @ X2 @ s__Tax)) != $true)) & (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) [flattening 65]
67. ! [X1 : w] : (? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) => (($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK0 @ X1) @ s__Tax))))) [skolem symbol introduction]
68. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK0 @ X1) @ s__Tax)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true)) [skolemisation 62,67]
69. ! [X1 : w] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) => (($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Giving))))) [skolem symbol introduction]
70. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Giving)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [skolemisation 64,69]
71. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) => (($true = ((s__destination @ sK3 @ sK4 @ cw))) & ($true = ((s__agent @ sK3 @ s__Bill @ cw))) & ($true = ((s__instance @ sK4 @ s__Moon))) & ($true = ((s__instance @ sK3 @ s__Flying)))) [skolem symbol introduction]
72. ($true = ((s__destination @ sK3 @ sK4 @ cw))) & ($true = ((s__agent @ sK3 @ s__Bill @ cw))) & ($true = ((s__instance @ sK4 @ s__Moon))) & ($true = ((s__instance @ sK3 @ s__Flying))) [skolemisation 35,71]
95. ($true = ((s__instance @ (sK0 @ X1) @ s__Tax))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) [cnf transformation 68]
96. ($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) [cnf transformation 68]
97. ($true = ((s__instance @ (sK1 @ X1) @ s__Giving))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
98. ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
99. ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
100. ($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
101. ($true = ((s__instance @ sK3 @ s__Flying))) [cnf transformation 72]
102. ($true = ((s__instance @ sK4 @ s__Moon))) [cnf transformation 72]
103. ($true = ((s__agent @ sK3 @ s__Bill @ cw))) [cnf transformation 72]
104. ($true = ((s__destination @ sK3 @ sK4 @ cw))) [cnf transformation 72]
106. (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) [cnf transformation 66]
107. (((s__origin @ X2 @ s__Bill @ X1)) != $true) | (((s__instance @ X2 @ s__Tax)) != $true) [cnf transformation 66]
110. ($true = ((s__instance @ (sK1 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [forward subsumption resolution 97,106]
111. ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [forward subsumption resolution 98,106]
112. ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [forward subsumption resolution 99,106]
113. ($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [forward subsumption resolution 100,106]
114. ($true = ((s__instance @ (sK0 @ X1) @ s__Tax))) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) [forward subsumption resolution 95,106]
115. ($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) [forward subsumption resolution 96,106]
117. 1 <=> ! [X2,X0 : w,X3] : ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true)) [avatar definition]
118. (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) <- (1) [avatar component clause 117]
120. 2 <=> ! [X1 : w] : ($true = ((s__instance @ (sK1 @ X1) @ s__Giving))) [avatar definition]
121. ($true = ((s__instance @ (sK1 @ X1) @ s__Giving))) <- (2) [avatar component clause 120]
122. 1 | 2 [avatar split clause 110,120,117]
124. 3 <=> ! [X1 : w] : ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) [avatar definition]
125. ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) <- (3) [avatar component clause 124]
126. 1 | 3 [avatar split clause 111,124,117]
128. 4 <=> ! [X1 : w] : ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) [avatar definition]
129. ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) <- (4) [avatar component clause 128]
130. 1 | 4 [avatar split clause 112,128,117]
132. 5 <=> ! [X1 : w] : ($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) [avatar definition]
133. ($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) <- (5) [avatar component clause 132]
134. 1 | 5 [avatar split clause 113,132,117]
136. 6 <=> ! [X2,X0 : w,X3] : ((((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true)) [avatar definition]
137. (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) <- (6) [avatar component clause 136]
139. 7 <=> ! [X1 : w] : ($true = ((s__instance @ (sK0 @ X1) @ s__Tax))) [avatar definition]
140. ($true = ((s__instance @ (sK0 @ X1) @ s__Tax))) <- (7) [avatar component clause 139]
141. 6 | 7 [avatar split clause 114,139,136]
143. 8 <=> ! [X1 : w] : ($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) [avatar definition]
144. ($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) <- (8) [avatar component clause 143]
145. 6 | 8 [avatar split clause 115,143,136]
153. ($true != $true) | ($true != ((s__instance @ (sK1 @ X0) @ s__Giving))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__patient @ (sK1 @ X0) @ X1 @ X0))) <- (5, 6) [superposition 137,133]
156. ($true != ((s__patient @ (sK1 @ X0) @ X1 @ X0))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ (sK1 @ X0) @ s__Giving))) <- (5, 6) [trivial inequality removal 153]
157. ($true != ((s__patient @ (sK1 @ X0) @ X1 @ X0))) | ($true != ((s__instance @ X1 @ s__PayCheck))) <- (2, 5, 6) [forward subsumption resolution 156,121]
161. ($true != $true) | ($true != ((s__instance @ (sK2 @ X0) @ s__PayCheck))) <- (2, 4, 5, 6) [superposition 157,129]
164. ($true != ((s__instance @ (sK2 @ X0) @ s__PayCheck))) <- (2, 4, 5, 6) [trivial inequality removal 161]
165. $false <- (2, 3, 4, 5, 6) [forward subsumption resolution 164,125]
166. ~2 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 165]
191. ($true != $true) | ($true != ((s__instance @ sK3 @ s__Flying))) | ($true != ((s__instance @ sK4 @ s__Moon))) | ($true != ((s__agent @ sK3 @ s__Bill @ cw))) <- (1) [superposition 118,104]
192. ($true != ((s__instance @ sK3 @ s__Flying))) | ($true != ((s__instance @ sK4 @ s__Moon))) | ($true != ((s__agent @ sK3 @ s__Bill @ cw))) <- (1) [trivial inequality removal 191]
193. ($true != ((s__instance @ sK4 @ s__Moon))) | ($true != ((s__agent @ sK3 @ s__Bill @ cw))) <- (1) [forward subsumption resolution 192,101]
194. ($true != ((s__agent @ sK3 @ s__Bill @ cw))) <- (1) [forward subsumption resolution 193,102]
195. $false <- (1) [forward subsumption resolution 194,103]
196. ~1 [avatar contradiction clause 195]
201. ($true != $true) | ($true != ((s__instance @ (sK0 @ X0) @ s__Tax))) <- (8) [superposition 107,144]
202. ($true != ((s__instance @ (sK0 @ X0) @ s__Tax))) <- (8) [trivial inequality removal 201]
205. $false <- (7, 8) [forward subsumption resolution 202,140]
206. ~7 | ~8 [avatar contradiction clause 205]
s1. 1 | 2 [sat_conversion 122]
s2. 1 | 3 [sat_conversion 126]
s3. 1 | 4 [sat_conversion 130]
s4. 1 | 5 [sat_conversion 134]
s5. 6 | 7 [sat_conversion 141]
s6. 6 | 8 [sat_conversion 145]
s7. ~6 | ~5 | ~4 | ~3 | ~2 [sat_conversion 166]
s9. ~1 [sat_conversion 196]
s10. ~8 | ~7 [sat_conversion 206]
s11. 5 [rat s4,s9]
s12. 4 [rat s3,s9]
s13. 3 [rat s2,s9]
s14. 2 [rat s1,s9]
s15. ~6 [rat s7,s13,s12,s11,s14]
s16. 8 [rat s6,s15]
s17. 7 [rat s5,s15]
s18. # [rat s10,s17,s16]
207. $false [avatar sat refutation s18]
% SZS output end Proof for 1e_Obligation_existentialNHops
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.146 s
% Peak memory usage: 12 MB
% Instructions burned: 18 (million)
% ------------------------------
% ------------------------------
