% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 1a_Obligation_univConj
% SZS output start Proof for 1a_Obligation_univConj
1. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1) [input(axiom)]
3. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X0 @ X2) & (accreln1 @ s__Obligation @ X0 @ X1)) => (accreln1 @ s__Obligation @ X1 @ X2)) [input(axiom)]
26. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ s__Bill @ X0) & (s__patient @ X2 @ X3 @ X0) & (s__instance @ X3 @ s__PayCheck) & (s__instance @ X2 @ s__Giving)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4] : ((s__origin @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Tax)))) [input(axiom)]
27. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [input(axiom)]
28. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [input(axiom)]
29. ! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1] : ((s__origin @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Tax))) [input(conjecture)]
30. ~! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1] : ((s__origin @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Tax))) [negated conjecture 29]
31. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1) [rectify 1]
32. ! [X0 : w] : ? [X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) [fool elimination 31]
35. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X0 @ X2) & (accreln1 @ s__Obligation @ X0 @ X1)) => (accreln1 @ s__Obligation @ X1 @ X2)) [rectify 3]
36. ! [X0 : w,X1 : w,X2 : w] : (((((accreln1 @ s__Obligation @ X0 @ X2)) = $true) & (((accreln1 @ s__Obligation @ X0 @ X1)) = $true)) => (((accreln1 @ s__Obligation @ X1 @ X2)) = $true)) [fool elimination 35]
39. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ s__Bill @ X0) & (s__patient @ X2 @ X3 @ X0) & (s__instance @ X3 @ s__PayCheck) & (s__instance @ X2 @ s__Giving)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4] : ((s__origin @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Tax)))) [rectify 26]
40. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ s__Bill @ X0)) = $true) & (((s__patient @ X2 @ X3 @ X0)) = $true) & (((s__instance @ X3 @ s__PayCheck)) = $true) & (((s__instance @ X2 @ s__Giving)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)))) [fool elimination 39]
41. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [rectify 27]
42. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [fool elimination 41]
43. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [rectify 28]
44. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) [fool elimination 43]
45. ~! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1] : ((s__origin @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Tax))) [rectify 30]
46. ~! [X0 : w] : ((((accreln1 @ s__Obligation @ cw @ X0)) = $true) => ? [X1] : ((((s__origin @ X1 @ s__Bill @ X0)) = $true) & (((s__instance @ X1 @ s__Tax)) = $true))) [fool elimination 45]
68. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ s__Bill @ X0)) = $true) & (((s__patient @ X2 @ X3 @ X0)) = $true) & (((s__instance @ X3 @ s__PayCheck)) = $true) & (((s__instance @ X2 @ s__Giving)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)))) [flattening 40]
69. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [flattening 42]
72. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Obligation @ X1 @ X2)) = $true) | ((((accreln1 @ s__Obligation @ X0 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true))) [ennf transformation 36]
73. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Obligation @ X1 @ X2)) = $true) | (((accreln1 @ s__Obligation @ X0 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) [flattening 72]
74. ! [X0 : w,X1 : w,X2,X3] : ((? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true))) [ennf transformation 68]
75. ! [X0 : w,X1 : w,X2,X3] : (? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true)) [flattening 74]
76. ! [X0 : w,X1 : w,X2,X3] : ((? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true))) [ennf transformation 69]
77. ! [X0 : w,X1 : w,X2,X3] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [flattening 76]
78. ? [X0 : w] : (! [X1] : ((((s__origin @ X1 @ s__Bill @ X0)) != $true) | (((s__instance @ X1 @ s__Tax)) != $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) [ennf transformation 46]
79. ! [X0 : w] : (? [X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0))))) [skolem symbol introduction]
80. ! [X0 : w] : ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)))) [skolemisation 32,79]
81. ! [X1 : w] : (? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) => (($true = ((s__origin @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Tax))))) [skolem symbol introduction]
82. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__origin @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Tax)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true)) [skolemisation 75,81]
83. ! [X1 : w] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) => (($true = ((s__destination @ (sK2 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK2 @ X1) @ (sK3 @ X1) @ X1))) & ($true = ((s__instance @ (sK3 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK2 @ X1) @ s__Giving))))) [skolem symbol introduction]
84. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK2 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK2 @ X1) @ (sK3 @ X1) @ X1))) & ($true = ((s__instance @ (sK3 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK2 @ X1) @ s__Giving)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [skolemisation 77,83]
85. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) => (($true = ((s__destination @ sK4 @ sK5 @ cw))) & ($true = ((s__agent @ sK4 @ s__Bill @ cw))) & ($true = ((s__instance @ sK5 @ s__Moon))) & ($true = ((s__instance @ sK4 @ s__Flying)))) [skolem symbol introduction]
86. ($true = ((s__destination @ sK4 @ sK5 @ cw))) & ($true = ((s__agent @ sK4 @ s__Bill @ cw))) & ($true = ((s__instance @ sK5 @ s__Moon))) & ($true = ((s__instance @ sK4 @ s__Flying))) [skolemisation 44,85]
87. ? [X0 : w] : (! [X1] : ((((s__origin @ X1 @ s__Bill @ X0)) != $true) | (((s__instance @ X1 @ s__Tax)) != $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) => (! [X1] : (($true != ((s__origin @ X1 @ s__Bill @ sK6))) | (((s__instance @ X1 @ s__Tax)) != $true)) & ($true = ((accreln1 @ s__Obligation @ cw @ sK6)))) [skolem symbol introduction]
88. ! [X1] : (($true != ((s__origin @ X1 @ s__Bill @ sK6))) | (((s__instance @ X1 @ s__Tax)) != $true)) & ($true = ((accreln1 @ s__Obligation @ cw @ sK6))) [skolemisation 78,87]
89. ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)))) [cnf transformation 80]
91. (((accreln1 @ s__Obligation @ X0 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X1 @ X2)) = $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) [cnf transformation 73]
114. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__Tax))) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) [cnf transformation 82]
115. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__origin @ (sK1 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) [cnf transformation 82]
116. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK2 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 84]
117. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 84]
118. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__patient @ (sK2 @ X1) @ (sK3 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 84]
119. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK2 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 84]
120. ($true = ((s__instance @ sK4 @ s__Flying))) [cnf transformation 86]
121. ($true = ((s__instance @ sK5 @ s__Moon))) [cnf transformation 86]
122. ($true = ((s__agent @ sK4 @ s__Bill @ cw))) [cnf transformation 86]
123. ($true = ((s__destination @ sK4 @ sK5 @ cw))) [cnf transformation 86]
124. ($true = ((accreln1 @ s__Obligation @ cw @ sK6))) [cnf transformation 88]
125. ($true != ((s__origin @ X1 @ s__Bill @ sK6))) | (((s__instance @ X1 @ s__Tax)) != $true) [cnf transformation 88]
127. ($true = X0) | ($false = X0) [fool exhaustiveness axiom]
130. ($true != $true) | ($true != ((s__instance @ X0 @ s__Tax))) | ($false = ((s__origin @ X0 @ s__Bill @ sK6))) [superposition 125,127]
132. ($true != ((s__instance @ X0 @ s__Tax))) | ($false = ((s__origin @ X0 @ s__Bill @ sK6))) [trivial inequality removal 130]
856. ($true != $true) | ($true = ((accreln1 @ s__Obligation @ X0 @ sK6))) | (((accreln1 @ s__Obligation @ cw @ X0)) != $true) [superposition 91,124]
880. (((accreln1 @ s__Obligation @ cw @ X0)) != $true) | ($true = ((accreln1 @ s__Obligation @ X0 @ sK6))) [trivial inequality removal 856]
917. ($true != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__Tax))) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [superposition 114,127]
937. (((s__destination @ X2 @ s__Bill @ X0)) != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__Tax))) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [trivial inequality removal 917]
1123. ($true != $true) | ($true = ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Giving))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 116,89]
1149. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Giving))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 1123]
1194. ($true != $true) | ($true = ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__PayCheck))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 117,89]
1220. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__PayCheck))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 1194]
1361. ($true != $true) | ($true = ((s__origin @ (sK1 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [superposition 115,127]
1388. (((s__destination @ X2 @ s__Bill @ X0)) != $true) | ($true = ((s__origin @ (sK1 @ X1) @ s__Bill @ X1))) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [trivial inequality removal 1361]
1442. ($true != $true) | ($true = ((s__destination @ (sK2 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 119,89]
1471. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__destination @ (sK2 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 1442]
1485. ($true != $true) | ($true = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ sK6))) [superposition 880,89]
1514. ($true = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ sK6))) [trivial inequality removal 1485]
1528. ($true != $true) | ($true = ((s__patient @ (sK2 @ (sK0 @ X0)) @ (sK3 @ (sK0 @ X0)) @ (sK0 @ X0)))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 118,89]
1557. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__patient @ (sK2 @ (sK0 @ X0)) @ (sK3 @ (sK0 @ X0)) @ (sK0 @ X0)))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 1528]
2323. ($true != $true) | ($true = ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [superposition 1149,123]
2352. ($true = ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [trivial inequality removal 2323]
2364. ($true = ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [forward subsumption resolution 2352,122]
2365. ($true = ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($true != ((s__instance @ sK4 @ s__Flying))) [forward subsumption resolution 2364,121]
2366. ($true = ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) [forward subsumption resolution 2365,120]
2430. ($true != $true) | ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [superposition 1220,123]
2459. ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [trivial inequality removal 2430]
2471. ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [forward subsumption resolution 2459,122]
2472. ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__instance @ sK4 @ s__Flying))) [forward subsumption resolution 2471,121]
2473. ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) [forward subsumption resolution 2472,120]
2537. ($true != $true) | ($true = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [superposition 1471,123]
2566. ($true = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [trivial inequality removal 2537]
2578. ($true = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [forward subsumption resolution 2566,122]
2579. ($true = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ sK4 @ s__Flying))) [forward subsumption resolution 2578,121]
2580. ($true = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward subsumption resolution 2579,120]
2833. ($true != $true) | ($true = ((s__patient @ (sK2 @ (sK0 @ cw)) @ (sK3 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [superposition 1557,123]
2862. ($true = ((s__patient @ (sK2 @ (sK0 @ cw)) @ (sK3 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [trivial inequality removal 2833]
2874. ($true = ((s__patient @ (sK2 @ (sK0 @ cw)) @ (sK3 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__instance @ sK4 @ s__Flying))) [forward subsumption resolution 2862,122]
2875. ($true = ((s__patient @ (sK2 @ (sK0 @ cw)) @ (sK3 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__instance @ sK4 @ s__Flying))) [forward subsumption resolution 2874,121]
2876. ($true = ((s__patient @ (sK2 @ (sK0 @ cw)) @ (sK3 @ (sK0 @ cw)) @ (sK0 @ cw)))) [forward subsumption resolution 2875,120]
3655. 46 <=> ($true = ((s__instance @ (sK1 @ sK6) @ s__Tax))) [avatar definition]
3656. ($true != ((s__instance @ (sK1 @ sK6) @ s__Tax))) <- (~46) [avatar component clause 3655]
3657. ($true = ((s__instance @ (sK1 @ sK6) @ s__Tax))) <- (46) [avatar component clause 3655]
4124. ($true != $true) | ($false = ((s__origin @ (sK1 @ sK6) @ s__Bill @ sK6))) <- (46) [superposition 132,3657]
4125. ($false = ((s__origin @ (sK1 @ sK6) @ s__Bill @ sK6))) <- (46) [trivial inequality removal 4124]
6838. ($true != $true) | ($true = ((s__instance @ (sK1 @ X2) @ s__Tax))) | ($true != ((s__patient @ X0 @ X3 @ X1))) | (((s__instance @ X3 @ s__PayCheck)) != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | (((accreln1 @ s__Obligation @ X1 @ X2)) = $false) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) [superposition 937,127]
6883. ($true != ((s__patient @ X0 @ X3 @ X1))) | ($true = ((s__instance @ (sK1 @ X2) @ s__Tax))) | (((s__instance @ X3 @ s__PayCheck)) != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | (((accreln1 @ s__Obligation @ X1 @ X2)) = $false) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) [trivial inequality removal 6838]
6929. ($true != $true) | ($true = ((s__origin @ (sK1 @ X2) @ s__Bill @ X2))) | ($true != ((s__patient @ X0 @ X3 @ X1))) | (((s__instance @ X3 @ s__PayCheck)) != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | (((accreln1 @ s__Obligation @ X1 @ X2)) = $false) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) [superposition 1388,127]
6974. ($true != ((s__patient @ X0 @ X3 @ X1))) | ($true = ((s__origin @ (sK1 @ X2) @ s__Bill @ X2))) | (((s__instance @ X3 @ s__PayCheck)) != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | (((accreln1 @ s__Obligation @ X1 @ X2)) = $false) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) [trivial inequality removal 6929]
13376. 314 <=> ! [X0 : w] : (($true = ((s__instance @ (sK1 @ X0) @ s__Tax))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0)))) [avatar definition]
13377. ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($true = ((s__instance @ (sK1 @ X0) @ s__Tax))) <- (314) [avatar component clause 13376]
13380. 315 <=> ! [X0 : w] : (($true = ((s__origin @ (sK1 @ X0) @ s__Bill @ X0))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0)))) [avatar definition]
13381. ($true = ((s__origin @ (sK1 @ X0) @ s__Bill @ X0))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) <- (315) [avatar component clause 13380]
24099. ($true != $true) | ($true = ((s__instance @ (sK1 @ X0) @ s__Tax))) | ($true != ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($false = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [superposition 6883,2876]
24215. ($true = ((s__instance @ (sK1 @ X0) @ s__Tax))) | ($true != ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($false = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [trivial inequality removal 24099]
24238. ($true = ((s__instance @ (sK1 @ X0) @ s__Tax))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($false = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward subsumption resolution 24215,2473]
24260. ($true = ((s__instance @ (sK1 @ X0) @ s__Tax))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($false = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward subsumption resolution 24238,2366]
24262. ($true = $false) | ($true = ((s__instance @ (sK1 @ X0) @ s__Tax))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) [forward demodulation 24260,2580]
24263. ($true = ((s__instance @ (sK1 @ X0) @ s__Tax))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) [trivial inequality removal 24262]
24264. 314 [avatar split clause 24263,13376]
25053. ($true != $true) | ($true = ((s__origin @ (sK1 @ X0) @ s__Bill @ X0))) | ($true != ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($false = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [superposition 6974,2876]
25169. ($true = ((s__origin @ (sK1 @ X0) @ s__Bill @ X0))) | ($true != ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($false = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [trivial inequality removal 25053]
25194. ($true = ((s__origin @ (sK1 @ X0) @ s__Bill @ X0))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__Giving))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($false = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward subsumption resolution 25169,2473]
25217. ($true = ((s__origin @ (sK1 @ X0) @ s__Bill @ X0))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) | ($false = ((s__destination @ (sK2 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward subsumption resolution 25194,2366]
25219. ($true = $false) | ($true = ((s__origin @ (sK1 @ X0) @ s__Bill @ X0))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) [forward demodulation 25217,2580]
25220. ($true = ((s__origin @ (sK1 @ X0) @ s__Bill @ X0))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ X0))) [trivial inequality removal 25219]
25221. 315 [avatar split clause 25220,13380]
84331. ($true = $false) | ($true = ((s__instance @ (sK1 @ sK6) @ s__Tax))) <- (314) [superposition 13377,1514]
84342. ($true = ((s__instance @ (sK1 @ sK6) @ s__Tax))) <- (314) [trivial inequality removal 84331]
90956. ($true = $false) | ($true = ((s__origin @ (sK1 @ sK6) @ s__Bill @ sK6))) <- (315) [superposition 13381,1514]
90959. ($true = ((s__origin @ (sK1 @ sK6) @ s__Bill @ sK6))) <- (315) [trivial inequality removal 90956]
90963. ($true = $false) <- (46, 315) [forward demodulation 90959,4125]
90964. $false <- (46, 315) [trivial inequality removal 90963]
90965. ~46 | ~315 [avatar contradiction clause 90964]
90968. $false <- (~46, 314) [forward subsumption resolution 84342,3656]
90969. 46 | ~314 [avatar contradiction clause 90968]
s777. 314 [sat_conversion 24264]
s798. 315 [sat_conversion 25221]
s1531. ~315 | ~46 [sat_conversion 90965]
s1533. ~314 | 46 [sat_conversion 90969]
s1539. ~46 [rat s1531,s798]
s1540. ~314 [rat s1533,s1539]
s1541. # [rat s777,s1540]
90974. $false [avatar sat refutation s1541]
% SZS output end Proof for 1a_Obligation_univConj
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 52.902 s
% Peak memory usage: 39 MB
% Instructions burned: 429017 (million)
% ------------------------------
% ------------------------------
