% Checking upload ...
% Checker ran ...
% No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/pa8LrHV0cP/SOT_oBUuvd
% (2605536)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_oBUuvd for (1ds/183Mi)
% (2605540)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_oBUuvd for (1ds/2Mi)
% (2605537)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_oBUuvd for (1ds/4Mi)
2605541)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_oBUuvd for (1ds/275Mi)
% (2605539)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_oBUuvd for (1ds/2Mi)
% (2605538)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_oBUuvd for (1ds/27Mi)
% (2605542)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_oBUuvd for (1ds/18Mi)
% (2605540)Refutation not found, incomplete strategy
% (2605540)------------------------------
% (2605540)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2605540)Termination reason: Refutation not found, incomplete strategy


% (2605540)Memory used [KB]: 5500
% (2605540)Time elapsed: 0.004 s
% (2605540)------------------------------
% (2605540)------------------------------
% (2605545)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_oBUuvd for (1ds/3Mi)
% (2605541)First to succeed.
% (2605539)Also succeeded, but the first one will report.
% (2605538)Also succeeded, but the first one will report.
% (2605541)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_oBUuvd
% SZS output start Proof for SOT_oBUuvd
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Prohibition: m).
thf(func_def_2, type, accreln1: m &gt; w &gt; w &gt; $o).
thf(func_def_3, type, accreln2: m &gt; $i &gt; w &gt; w &gt; $o).
thf(func_def_4, type, accreln3: m &gt; $i &gt; $i &gt; w &gt; w &gt; $o).
thf(func_def_6, type, s__instance: $i &gt; $i &gt; $o).
thf(func_def_7, type, s__agent: $i &gt; $i &gt; w &gt; $o).
thf(func_def_8, type, s__patient: $i &gt; $i &gt; w &gt; $o).
thf(func_def_9, type, s__origin: $i &gt; $i &gt; w &gt; $o).
thf(func_def_10, type, s__destination: $i &gt; $i &gt; w &gt; $o).
thf(func_def_11, type, s__located: $i &gt; $i &gt; w &gt; $o).
thf(func_def_20, type, sK0: w &gt; w).
thf(func_def_21, type, sK1: w &gt; $i &gt; $i).
thf(func_def_23, type, sK3: w &gt; $i).
thf(f153,plain,(
  $false),
  inference(subsumption_resolution,[status(thm)],[f152,f55])).
thf(f55,plain,(
  ((s__instance @ sK2 @ s__OfficeBuilding) = $true)),
  inference(cnf_transformation,[status(thm)],[f44])).
thf(f44,plain,(
  ((s__located @ s__Bill @ sK2 @ cw) = $true) & ((s__instance @ sK2 @ s__OfficeBuilding) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f25,f43])).
thf(f43,plain,(
  ? [X0 : $i] : (((s__located @ s__Bill @ X0 @ cw) = $true) & ((s__instance @ X0 @ s__OfficeBuilding) = $true)) =&gt; (((s__located @ s__Bill @ sK2 @ cw) = $true) & ((s__instance @ sK2 @ s__OfficeBuilding) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f25,plain,(
  ? [X0 : $i] : (((s__located @ s__Bill @ X0 @ cw) = $true) & ((s__instance @ X0 @ s__OfficeBuilding) = $true))),
  inference(fool_elimination,[status(thm)],[f24])).
thf(f24,plain,(
  ? [X0 : $i] : ((s__instance @ X0 @ s__OfficeBuilding) & (s__located @ s__Bill @ X0 @ cw))),
  inference(rectify,[status(thm)],[f12])).
thf(f12,axiom,(
  ? [X3 : $i] : ((s__instance @ X3 @ s__OfficeBuilding) & (s__located @ s__Bill @ X3 @ cw))),
  file('/tmp/pa8LrHV0cP/SOT_oBUuvd',fact_bill_at_office)).
thf(f152,plain,(
  ((s__instance @ sK2 @ s__OfficeBuilding) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f151])).
thf(f151,plain,(
  ($true != $true) | ((s__instance @ sK2 @ s__OfficeBuilding) != $true)),
  inference(superposition,[status(thm)],[f145,f56])).
thf(f56,plain,(
  ((s__located @ s__Bill @ sK2 @ cw) = $true)),
  inference(cnf_transformation,[status(thm)],[f44])).
thf(f145,plain,(
  ( ! [X0 : $i] : (((s__located @ s__Bill @ X0 @ cw) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f143])).
thf(f143,plain,(
  ( ! [X0 : $i] : (((s__located @ s__Bill @ X0 @ cw) != $true) | ($true != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(superposition,[status(thm)],[f137,f49])).
thf(f49,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Prohibition @ X0 @ (sK0 @ X0)))) )),
  inference(cnf_transformation,[status(thm)],[f38])).
thf(f38,plain,(
  ! [X0 : w] : ($true = (accreln1 @ s__Prohibition @ X0 @ (sK0 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f29,f37])).
thf(f37,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Prohibition @ X0 @ X1) = $true) =&gt; ($true = (accreln1 @ s__Prohibition @ X0 @ (sK0 @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f29,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Prohibition @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f28])).
thf(f28,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Prohibition @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Prohibition @ X0 @ X1)),
  file('/tmp/pa8LrHV0cP/SOT_oBUuvd',serial_prohibition)).
thf(f137,plain,(
  ( ! [X0 : $i,X1 : w] : (((accreln1 @ s__Prohibition @ X1 @ (sK0 @ cw)) != $true) | ((s__located @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(subsumption_resolution,[status(thm)],[f136,f52])).
thf(f52,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (((s__agent @ (sK1 @ X2 @ X0) @ s__Bill @ X2) = $true) | ((accreln1 @ s__Prohibition @ X1 @ X2) != $true) | ((s__located @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f42])).
thf(f42,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (((accreln1 @ s__Prohibition @ X1 @ X2) != $true) | ((s__located @ s__Bill @ X0 @ X1) != $true) | (((s__instance @ (sK1 @ X2 @ X0) @ s__Smoking) = $true) & ((s__located @ (sK1 @ X2 @ X0) @ X0 @ X2) = $true) & ((s__agent @ (sK1 @ X2 @ X0) @ s__Bill @ X2) = $true)) | ((s__instance @ X0 @ s__OfficeBuilding) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f40,f41])).
thf(f41,plain,(
  ! [X0 : $i,X2 : w] : (? [X3 : $i] : (((s__instance @ X3 @ s__Smoking) = $true) & ($true = (s__located @ X3 @ X0 @ X2)) & ($true = (s__agent @ X3 @ s__Bill @ X2))) =&gt; (((s__instance @ (sK1 @ X2 @ X0) @ s__Smoking) = $true) & ((s__located @ (sK1 @ X2 @ X0) @ X0 @ X2) = $true) & ((s__agent @ (sK1 @ X2 @ X0) @ s__Bill @ X2) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f40,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (((accreln1 @ s__Prohibition @ X1 @ X2) != $true) | ((s__located @ s__Bill @ X0 @ X1) != $true) | ? [X3 : $i] : (((s__instance @ X3 @ s__Smoking) = $true) & ($true = (s__located @ X3 @ X0 @ X2)) & ($true = (s__agent @ X3 @ s__Bill @ X2))) | ((s__instance @ X0 @ s__OfficeBuilding) != $true))),
  inference(rectify,[status(thm)],[f34])).
thf(f34,plain,(
  ! [X0 : $i,X2 : w,X1 : w] : (((accreln1 @ s__Prohibition @ X2 @ X1) != $true) | ($true != (s__located @ s__Bill @ X0 @ X2)) | ? [X3 : $i] : (((s__instance @ X3 @ s__Smoking) = $true) & ((s__located @ X3 @ X0 @ X1) = $true) & ((s__agent @ X3 @ s__Bill @ X1) = $true)) | ((s__instance @ X0 @ s__OfficeBuilding) != $true))),
  inference(flattening,[status(thm)],[f33])).
thf(f33,plain,(
  ! [X1 : w,X2 : w,X0 : $i] : ((? [X3 : $i] : (((s__instance @ X3 @ s__Smoking) = $true) & ((s__located @ X3 @ X0 @ X1) = $true) & ((s__agent @ X3 @ s__Bill @ X1) = $true)) | ((accreln1 @ s__Prohibition @ X2 @ X1) != $true)) | (($true != (s__located @ s__Bill @ X0 @ X2)) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)))),
  inference(ennf_transformation,[status(thm)],[f19])).
thf(f19,plain,(
  ! [X1 : w,X2 : w,X0 : $i] : ((($true = (s__located @ s__Bill @ X0 @ X2)) & ((s__instance @ X0 @ s__OfficeBuilding) = $true)) =&gt; (((accreln1 @ s__Prohibition @ X2 @ X1) = $true) =&gt; ? [X3 : $i] : (((s__instance @ X3 @ s__Smoking) = $true) & ((s__located @ X3 @ X0 @ X1) = $true) & ((s__agent @ X3 @ s__Bill @ X1) = $true))))),
  inference(fool_elimination,[status(thm)],[f18])).
thf(f18,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (((s__located @ s__Bill @ X0 @ X2) & (s__instance @ X0 @ s__OfficeBuilding)) =&gt; ((accreln1 @ s__Prohibition @ X2 @ X1) =&gt; ? [X3 : $i] : ((s__located @ X3 @ X0 @ X1) & (s__instance @ X3 @ s__Smoking) & (s__agent @ X3 @ s__Bill @ X1))))),
  inference(rectify,[status(thm)],[f10])).
thf(f10,axiom,(
  ! [X3 : $i,X1 : w,X0 : w] : (((s__located @ s__Bill @ X3 @ X0) & (s__instance @ X3 @ s__OfficeBuilding)) =&gt; ((accreln1 @ s__Prohibition @ X0 @ X1) =&gt; ? [X4 : $i] : ((s__located @ X4 @ X3 @ X1) & (s__instance @ X4 @ s__Smoking) & (s__agent @ X4 @ s__Bill @ X1))))),
  file('/tmp/pa8LrHV0cP/SOT_oBUuvd',prohibition1_smoking_in_office)).
thf(f136,plain,(
  ( ! [X0 : $i,X1 : w] : (((accreln1 @ s__Prohibition @ X1 @ (sK0 @ cw)) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ($true != (s__agent @ (sK1 @ (sK0 @ cw) @ X0) @ s__Bill @ (sK0 @ cw))) | ((s__located @ s__Bill @ X0 @ X1) != $true)) )),
  inference(subsumption_resolution,[status(thm)],[f135,f54])).
thf(f54,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (((s__instance @ (sK1 @ X2 @ X0) @ s__Smoking) = $true) | ((s__located @ s__Bill @ X0 @ X1) != $true) | ((accreln1 @ s__Prohibition @ X1 @ X2) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f42])).
thf(f135,plain,(
  ( ! [X0 : $i,X1 : w] : (((s__located @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ (sK1 @ (sK0 @ cw) @ X0) @ s__Smoking) != $true) | ((accreln1 @ s__Prohibition @ X1 @ (sK0 @ cw)) != $true) | ($true != (s__agent @ (sK1 @ (sK0 @ cw) @ X0) @ s__Bill @ (sK0 @ cw))) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f134])).
thf(f134,plain,(
  ( ! [X0 : $i,X1 : w] : (((accreln1 @ s__Prohibition @ X1 @ (sK0 @ cw)) != $true) | ($true != $true) | ((s__instance @ (sK1 @ (sK0 @ cw) @ X0) @ s__Smoking) != $true) | ($true != (s__agent @ (sK1 @ (sK0 @ cw) @ X0) @ s__Bill @ (sK0 @ cw))) | ((s__located @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(duplicate_literal_removal,[status(thm)],[f133])).
thf(f133,plain,(
  ( ! [X0 : $i,X1 : w] : (((s__located @ s__Bill @ X0 @ X1) != $true) | ($true != $true) | ((s__instance @ (sK1 @ (sK0 @ cw) @ X0) @ s__Smoking) != $true) | ($true != (s__agent @ (sK1 @ (sK0 @ cw) @ X0) @ s__Bill @ (sK0 @ cw))) | ((accreln1 @ s__Prohibition @ X1 @ (sK0 @ cw)) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(superposition,[status(thm)],[f128,f53])).
thf(f53,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (((s__located @ (sK1 @ X2 @ X0) @ X0 @ X2) = $true) | ((s__located @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Prohibition @ X1 @ X2) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f42])).
thf(f128,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__located @ X0 @ X1 @ (sK0 @ cw))) | ($true != (s__agent @ X0 @ s__Bill @ (sK0 @ cw))) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ((s__instance @ X0 @ s__Smoking) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f125])).
thf(f125,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__instance @ X1 @ s__OfficeBuilding)) | ((s__instance @ X0 @ s__Smoking) != $true) | ($true != (s__located @ X0 @ X1 @ (sK0 @ cw))) | ($true != (s__agent @ X0 @ s__Bill @ (sK0 @ cw))) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f118,f49])).
thf(f118,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((accreln1 @ s__Prohibition @ cw @ X0) != $true) | ((s__located @ X1 @ X2 @ X0) != $true) | ((s__instance @ X1 @ s__Smoking) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f115])).
thf(f115,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != $true) | ((s__instance @ X1 @ s__Smoking) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Prohibition @ cw @ X0) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__located @ X1 @ X2 @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f78,f49])).
thf(f78,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Prohibition @ X1 @ (sK0 @ X0)) != $true) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ($true != (s__located @ X2 @ X3 @ X1)) | ((s__instance @ X2 @ s__Smoking) != $true) | ((s__instance @ X3 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Prohibition @ cw @ X0) != $true)) )),
  inference(subsumption_resolution,[status(thm)],[f76,f57])).
thf(f57,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__located @ X2 @ X3 @ X1)) | ($true = (s__instance @ (sK3 @ X0) @ s__BusinessTrip)) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ X3 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Prohibition @ X1 @ X0) != $true) | ((s__instance @ X2 @ s__Smoking) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f47])).
thf(f47,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ X2 @ s__Smoking) != $true) | ($true != (s__located @ X2 @ X3 @ X1)) | ((s__instance @ X3 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Prohibition @ X1 @ X0) != $true) | (($true = (s__agent @ (sK3 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK3 @ X0) @ s__BusinessTrip))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f45,f46])).
thf(f46,plain,(
  ! [X0 : w] : (? [X4 : $i] : (((s__agent @ X4 @ s__Bill @ X0) = $true) & ((s__instance @ X4 @ s__BusinessTrip) = $true)) =&gt; (($true = (s__agent @ (sK3 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK3 @ X0) @ s__BusinessTrip))))),
  introduced(definition,[],[choice_axiom])).
thf(f45,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ X2 @ s__Smoking) != $true) | ($true != (s__located @ X2 @ X3 @ X1)) | ((s__instance @ X3 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Prohibition @ X1 @ X0) != $true) | ? [X4 : $i] : (((s__agent @ X4 @ s__Bill @ X0) = $true) & ((s__instance @ X4 @ s__BusinessTrip) = $true)))),
  inference(rectify,[status(thm)],[f36])).
thf(f36,plain,(
  ! [X1 : w,X3 : w,X0 : $i,X2 : $i] : (((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__Smoking) != $true) | ((s__located @ X0 @ X2 @ X3) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Prohibition @ X3 @ X1) != $true) | ? [X4 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ((s__instance @ X4 @ s__BusinessTrip) = $true)))),
  inference(flattening,[status(thm)],[f35])).
thf(f35,plain,(
  ! [X1 : w,X3 : w,X0 : $i,X2 : $i] : ((? [X4 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ((s__instance @ X4 @ s__BusinessTrip) = $true)) | ((accreln1 @ s__Prohibition @ X3 @ X1) != $true)) | (((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__located @ X0 @ X2 @ X3) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__instance @ X0 @ s__Smoking) != $true)))),
  inference(ennf_transformation,[status(thm)],[f27])).
thf(f27,plain,(
  ! [X1 : w,X3 : w,X0 : $i,X2 : $i] : ((((s__agent @ X0 @ s__Bill @ X3) = $true) & ((s__located @ X0 @ X2 @ X3) = $true) & ((s__instance @ X2 @ s__OfficeBuilding) = $true) & ((s__instance @ X0 @ s__Smoking) = $true)) =&gt; (((accreln1 @ s__Prohibition @ X3 @ X1) = $true) =&gt; ? [X4 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ((s__instance @ X4 @ s__BusinessTrip) = $true))))),
  inference(fool_elimination,[status(thm)],[f26])).
thf(f26,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__located @ X0 @ X2 @ X3) & (s__instance @ X2 @ s__OfficeBuilding) & (s__agent @ X0 @ s__Bill @ X3) & (s__instance @ X0 @ s__Smoking)) =&gt; ((accreln1 @ s__Prohibition @ X3 @ X1) =&gt; ? [X4 : $i] : ((s__agent @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__BusinessTrip))))),
  inference(rectify,[status(thm)],[f11])).
thf(f11,axiom,(
  ! [X4 : $i,X1 : w,X3 : $i,X0 : w] : (((s__located @ X4 @ X3 @ X0) & (s__instance @ X3 @ s__OfficeBuilding) & (s__agent @ X4 @ s__Bill @ X0) & (s__instance @ X4 @ s__Smoking)) =&gt; ((accreln1 @ s__Prohibition @ X0 @ X1) =&gt; ? [X5 : $i] : ((s__agent @ X5 @ s__Bill @ X1) & (s__instance @ X5 @ s__BusinessTrip))))),
  file('/tmp/pa8LrHV0cP/SOT_oBUuvd',prohibition2_business_trip)).
thf(f76,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((s__instance @ X3 @ s__OfficeBuilding) != $true) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ X2 @ s__Smoking) != $true) | ($true != (s__located @ X2 @ X3 @ X1)) | ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__BusinessTrip) != $true) | ((accreln1 @ s__Prohibition @ cw @ X0) != $true) | ((accreln1 @ s__Prohibition @ X1 @ (sK0 @ X0)) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f75])).
thf(f75,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Prohibition @ X1 @ (sK0 @ X0)) != $true) | ($true != $true) | ((accreln1 @ s__Prohibition @ cw @ X0) != $true) | ((s__instance @ X3 @ s__OfficeBuilding) != $true) | ($true != (s__located @ X2 @ X3 @ X1)) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__BusinessTrip) != $true) | ((s__instance @ X2 @ s__Smoking) != $true)) )),
  inference(superposition,[status(thm)],[f73,f58])).
thf(f58,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true = (s__agent @ (sK3 @ X0) @ s__Bill @ X0)) | ((accreln1 @ s__Prohibition @ X1 @ X0) != $true) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ($true != (s__located @ X2 @ X3 @ X1)) | ((s__instance @ X2 @ s__Smoking) != $true) | ((s__instance @ X3 @ s__OfficeBuilding) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f47])).
thf(f73,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__agent @ X1 @ s__Bill @ (sK0 @ X0)) != $true) | ((accreln1 @ s__Prohibition @ cw @ X0) != $true) | ((s__instance @ X1 @ s__BusinessTrip) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f70])).
thf(f70,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__BusinessTrip) != $true) | ($true != $true) | ((accreln1 @ s__Prohibition @ cw @ X0) != $true) | ((s__agent @ X1 @ s__Bill @ (sK0 @ X0)) != $true)) )),
  inference(superposition,[status(thm)],[f63,f49])).
thf(f63,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Prohibition @ X2 @ X0)) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((accreln1 @ s__Prohibition @ cw @ X2) != $true) | ((s__instance @ X1 @ s__BusinessTrip) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f62])).
thf(f62,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Prohibition @ cw @ X2) != $true) | ((s__instance @ X1 @ s__BusinessTrip) != $true) | ($true != (accreln1 @ s__Prohibition @ X2 @ X0)) | ($true != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f60,f51])).
thf(f51,plain,(
  ( ! [X2 : w,X0 : w,X1 : w] : (((accreln1 @ s__Prohibition @ X2 @ X1) = $true) | ((accreln1 @ s__Prohibition @ X0 @ X1) != $true) | ($true != (accreln1 @ s__Prohibition @ X2 @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f39])).
thf(f39,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (($true != (accreln1 @ s__Prohibition @ X2 @ X0)) | ((accreln1 @ s__Prohibition @ X0 @ X1) != $true) | ((accreln1 @ s__Prohibition @ X2 @ X1) = $true))),
  inference(rectify,[status(thm)],[f32])).
thf(f32,plain,(
  ! [X1 : w,X0 : w,X2 : w] : (((accreln1 @ s__Prohibition @ X2 @ X1) != $true) | ((accreln1 @ s__Prohibition @ X1 @ X0) != $true) | ($true = (accreln1 @ s__Prohibition @ X2 @ X0)))),
  inference(flattening,[status(thm)],[f31])).
thf(f31,plain,(
  ! [X2 : w,X0 : w,X1 : w] : (($true = (accreln1 @ s__Prohibition @ X2 @ X0)) | (((accreln1 @ s__Prohibition @ X1 @ X0) != $true) | ((accreln1 @ s__Prohibition @ X2 @ X1) != $true)))),
  inference(ennf_transformation,[status(thm)],[f21])).
thf(f21,plain,(
  ! [X2 : w,X0 : w,X1 : w] : ((((accreln1 @ s__Prohibition @ X1 @ X0) = $true) & ((accreln1 @ s__Prohibition @ X2 @ X1) = $true)) =&gt; ($true = (accreln1 @ s__Prohibition @ X2 @ X0)))),
  inference(fool_elimination,[status(thm)],[f20])).
thf(f20,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Prohibition @ X1 @ X0) & (accreln1 @ s__Prohibition @ X2 @ X1)) =&gt; (accreln1 @ s__Prohibition @ X2 @ X0))),
  inference(rectify,[status(thm)],[f2])).
thf(f2,axiom,(
  ! [X2 : w,X1 : w,X0 : w] : (((accreln1 @ s__Prohibition @ X1 @ X2) & (accreln1 @ s__Prohibition @ X0 @ X1)) =&gt; (accreln1 @ s__Prohibition @ X0 @ X2))),
  file('/tmp/pa8LrHV0cP/SOT_oBUuvd',transitive_prohibition)).
thf(f60,plain,(
  ( ! [X0 : w,X1 : $i] : (((accreln1 @ s__Prohibition @ cw @ X0) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__BusinessTrip) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f30])).
thf(f30,plain,(
  ! [X0 : w] : (! [X1 : $i] : (((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__BusinessTrip) != $true)) | ((accreln1 @ s__Prohibition @ cw @ X0) != $true))),
  inference(ennf_transformation,[status(thm)],[f17])).
thf(f17,plain,(
  ~? [X0 : w] : (? [X1 : $i] : (((s__agent @ X1 @ s__Bill @ X0) = $true) & ((s__instance @ X1 @ s__BusinessTrip) = $true)) & ((accreln1 @ s__Prohibition @ cw @ X0) = $true))),
  inference(fool_elimination,[status(thm)],[f16])).
thf(f16,plain,(
  ~? [X0 : w] : ((accreln1 @ s__Prohibition @ cw @ X0) & ? [X1 : $i] : ((s__agent @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__BusinessTrip)))),
  inference(rectify,[status(thm)],[f14])).
thf(f14,negated_conjecture,(
  ~? [X6 : w] : ((accreln1 @ s__Prohibition @ cw @ X6) & ? [X5 : $i] : ((s__agent @ X5 @ s__Bill @ X6) & (s__instance @ X5 @ s__BusinessTrip)))),
  inference(negated_conjecture,[status(cth)],[f13])).
thf(f13,conjecture,(
  ? [X6 : w] : ((accreln1 @ s__Prohibition @ cw @ X6) & ? [X5 : $i] : ((s__agent @ X5 @ s__Bill @ X6) & (s__instance @ X5 @ s__BusinessTrip)))),
  file('/tmp/pa8LrHV0cP/SOT_oBUuvd',conjecture_not_go_on_trip)).
% SZS output end Proof for SOT_oBUuvd
% (2605541)------------------------------
% (2605541)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2605541)Termination reason: Refutation

% (2605541)Memory used [KB]: 5628
% (2605541)Time elapsed: 0.023 s
% (2605541)------------------------------
% (2605541)------------------------------
% (2605532)Success in time 0.034 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_oBUuvd - Vampire---5.0.1 says Theorem - CPU = 0.02 WC = 0.07 
% OUTPUT: SOT_oBUuvd - Vampire---5.0.1 says Refutation - CPU = 0.02 WC = 0.07 
