% Checking upload ...
% Checker ran ...
% No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/w6_R_Xckdu/SOT_VU3a7g
% (2603648)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_VU3a7g for (1ds/4Mi)
% (2603650)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_VU3a7g for (1ds/2Mi)
% (2603652)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_VU3a7g for (1ds/275Mi)
% (2603649)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_VU3a7g for (1ds/27Mi)
% (2603651)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_VU3a7g for (1ds/2Mi)
% (2603647)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_VU3a7g for (1ds/183Mi)
% (2603653)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_VU3a7g for (1ds/18Mi)
Permission denied
% (2603651)Refutation not found, incomplete strategy
% (2603651)------------------------------
% (2603651)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2603651)Termination reason: Refutation not found, incomplete strategy


% (2603651)Memory used [KB]: 5628
% (2603651)Time elapsed: 0.024 s
% (2603651)------------------------------
% (2603651)------------------------------
% (2603655)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_VU3a7g for (1ds/3Mi)
% (2603650)First to succeed.
% (2603649)Also succeeded, but the first one will report.
% (2603650)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_VU3a7g
% SZS output start Proof for SOT_VU3a7g
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, accreln1: m &gt; w &gt; w &gt; $o).
thf(func_def_3, type, accreln2: m &gt; $i &gt; w &gt; w &gt; $o).
thf(func_def_4, type, accreln3: m &gt; $i &gt; $i &gt; w &gt; w &gt; $o).
thf(func_def_6, type, s__instance: $i &gt; $i &gt; $o).
thf(func_def_7, type, s__agent: $i &gt; $i &gt; w &gt; $o).
thf(func_def_8, type, s__patient: $i &gt; $i &gt; w &gt; $o).
thf(func_def_9, type, s__origin: $i &gt; $i &gt; w &gt; $o).
thf(func_def_10, type, s__destination: $i &gt; $i &gt; w &gt; $o).
thf(func_def_21, type, sK0: w &gt; $i).
thf(func_def_24, type, sK3: w &gt; $i).
thf(func_def_25, type, sK4: w &gt; $i).
thf(func_def_26, type, sK5: w &gt; w).
thf(f304,plain,(
  $false),
  inference(avatar_sat_refutation,[status(thm)],[f230,f236,f299,f303])).
thf(f303,plain,(
  ~spl6_6 | ~spl6_8),
  inference(avatar_contradiction_clause,[status(thm)],[f302])).
thf(f302,plain,(
  $false | (~spl6_6 | ~spl6_8)),
  inference(subsumption_resolution,[status(thm)],[f301,f235])).
thf(f235,plain,(
  ($true = (s__instance @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Tax)) | ~spl6_8),
  inference(avatar_component_clause,[status(thm)],[f233])).
thf(f233,definition,(
  spl6_8 &lt;=&gt; ($true = (s__instance @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Tax))),
  introduced(definition,[new_symbols(naming,[spl6_8])],[avatar_definition])).
thf(f301,plain,(
  ($true != (s__instance @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Tax)) | ~spl6_6),
  inference(trivial_inequality_removal,[status(thm)],[f300])).
thf(f300,plain,(
  ($true != (s__instance @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Tax)) | ($true != $true) | ~spl6_6),
  inference(superposition,[status(thm)],[f118,f226])).
thf(f226,plain,(
  ($true = (s__origin @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ~spl6_6),
  inference(avatar_component_clause,[status(thm)],[f224])).
thf(f224,definition,(
  spl6_6 &lt;=&gt; ($true = (s__origin @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw))))),
  introduced(definition,[new_symbols(naming,[spl6_6])],[avatar_definition])).
thf(f118,plain,(
  ( ! [X0 : $i] : (($true != (s__origin @ X0 @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != (s__instance @ X0 @ s__Tax))) )),
  inference(trivial_inequality_removal,[status(thm)],[f114])).
thf(f114,plain,(
  ( ! [X0 : $i] : (($true != $true) | ($true != (s__origin @ X0 @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != (s__instance @ X0 @ s__Tax))) )),
  inference(superposition,[status(thm)],[f77,f105])).
thf(f105,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Obligation @ X0 @ (sK5 @ (sK5 @ X0))))) )),
  inference(trivial_inequality_removal,[status(thm)],[f104])).
thf(f104,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true = (accreln1 @ s__Obligation @ X0 @ (sK5 @ (sK5 @ X0))))) )),
  inference(superposition,[status(thm)],[f101,f97])).
thf(f97,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Obligation @ X0 @ (sK5 @ X0)))) )),
  inference(cnf_transformation,[status(thm)],[f62])).
thf(f62,plain,(
  ! [X0 : w] : ($true = (accreln1 @ s__Obligation @ X0 @ (sK5 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f42,f61])).
thf(f61,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true) =&gt; ($true = (accreln1 @ s__Obligation @ X0 @ (sK5 @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f42,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f41])).
thf(f41,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  file('/tmp/w6_R_Xckdu/SOT_VU3a7g',serial_obligation)).
thf(f101,plain,(
  ( ! [X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ X1 @ X0)) | ($true = (accreln1 @ s__Obligation @ X1 @ (sK5 @ X0)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f100])).
thf(f100,plain,(
  ( ! [X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ X1 @ X0)) | ($true = (accreln1 @ s__Obligation @ X1 @ (sK5 @ X0))) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f87,f97])).
thf(f87,plain,(
  ( ! [X2 : w,X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ($true != (accreln1 @ s__Obligation @ X2 @ X1)))),
  inference(flattening,[status(thm)],[f48])).
thf(f48,plain,(
  ! [X2 : w,X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) | (($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)))),
  inference(ennf_transformation,[status(thm)],[f44])).
thf(f44,plain,(
  ! [X2 : w,X1 : w,X0 : w] : ((($true = (accreln1 @ s__Obligation @ X2 @ X1)) & ((accreln1 @ s__Obligation @ X0 @ X2) = $true)) =&gt; ((accreln1 @ s__Obligation @ X0 @ X1) = $true))),
  inference(fool_elimination,[status(thm)],[f43])).
thf(f43,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X2 @ X1) & (accreln1 @ s__Obligation @ X0 @ X2)) =&gt; (accreln1 @ s__Obligation @ X0 @ X1))),
  inference(rectify,[status(thm)],[f2])).
thf(f2,axiom,(
  ! [X0 : w,X2 : w,X1 : w] : (((accreln1 @ s__Obligation @ X1 @ X2) & (accreln1 @ s__Obligation @ X0 @ X1)) =&gt; (accreln1 @ s__Obligation @ X0 @ X2))),
  file('/tmp/w6_R_Xckdu/SOT_VU3a7g',transitive_obligation)).
thf(f77,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__origin @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X1 @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f45])).
thf(f45,plain,(
  ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ! [X1 : $i] : (($true != (s__instance @ X1 @ s__Tax)) | ($true != (s__origin @ X1 @ s__Bill @ X0))))),
  inference(ennf_transformation,[status(thm)],[f32])).
thf(f32,plain,(
  ~? [X0 : w] : (? [X1 : $i] : (($true = (s__instance @ X1 @ s__Tax)) & ($true = (s__origin @ X1 @ s__Bill @ X0))) & ($true = (accreln1 @ s__Obligation @ cw @ X0)))),
  inference(fool_elimination,[status(thm)],[f31])).
thf(f31,plain,(
  ~? [X0 : w] : (? [X1 : $i] : ((s__instance @ X1 @ s__Tax) & (s__origin @ X1 @ s__Bill @ X0)) & (accreln1 @ s__Obligation @ cw @ X0))),
  inference(rectify,[status(thm)],[f29])).
thf(f29,negated_conjecture,(
  ~? [X10 : w] : (? [X5 : $i] : ((s__instance @ X5 @ s__Tax) & (s__origin @ X5 @ s__Bill @ X10)) & (accreln1 @ s__Obligation @ cw @ X10))),
  inference(negated_conjecture,[status(cth)],[f28])).
thf(f28,conjecture,(
  ? [X10 : w] : (? [X5 : $i] : ((s__instance @ X5 @ s__Tax) & (s__origin @ X5 @ s__Bill @ X10)) & (accreln1 @ s__Obligation @ cw @ X10))),
  file('/tmp/w6_R_Xckdu/SOT_VU3a7g',conjecture_he_get_taxed)).
thf(f299,plain,(
  ~spl6_7),
  inference(avatar_contradiction_clause,[status(thm)],[f298])).
thf(f298,plain,(
  $false | ~spl6_7),
  inference(subsumption_resolution,[status(thm)],[f297,f158])).
thf(f158,plain,(
  ($true = (s__instance @ (sK4 @ (sK5 @ cw)) @ s__PayCheck))),
  inference(trivial_inequality_removal,[status(thm)],[f153])).
thf(f153,plain,(
  ($true = (s__instance @ (sK4 @ (sK5 @ cw)) @ s__PayCheck)) | ($true != $true)),
  inference(superposition,[status(thm)],[f110,f97])).
thf(f110,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__instance @ (sK4 @ X0) @ s__PayCheck))) )),
  inference(subsumption_resolution,[status(thm)],[f109,f83])).
thf(f83,plain,(
  ($true = (s__instance @ sK2 @ s__Moon))),
  inference(cnf_transformation,[status(thm)],[f57])).
thf(f57,plain,(
  ($true = (s__instance @ sK2 @ s__Moon)) & ($true = (s__instance @ sK1 @ s__Flying)) & ($true = (s__agent @ sK1 @ s__Bill @ cw)) & ($true = (s__destination @ sK1 @ sK2 @ cw))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f55,f56])).
thf(f56,plain,(
  ? [X0 : $i,X1 : $i] : (($true = (s__instance @ X1 @ s__Moon)) & ($true = (s__instance @ X0 @ s__Flying)) & ($true = (s__agent @ X0 @ s__Bill @ cw)) & ($true = (s__destination @ X0 @ X1 @ cw))) =&gt; (($true = (s__instance @ sK2 @ s__Moon)) & ($true = (s__instance @ sK1 @ s__Flying)) & ($true = (s__agent @ sK1 @ s__Bill @ cw)) & ($true = (s__destination @ sK1 @ sK2 @ cw)))),
  introduced(definition,[],[choice_axiom])).
thf(f55,plain,(
  ? [X0 : $i,X1 : $i] : (($true = (s__instance @ X1 @ s__Moon)) & ($true = (s__instance @ X0 @ s__Flying)) & ($true = (s__agent @ X0 @ s__Bill @ cw)) & ($true = (s__destination @ X0 @ X1 @ cw)))),
  inference(rectify,[status(thm)],[f34])).
thf(f34,plain,(
  ? [X1 : $i,X0 : $i] : (($true = (s__instance @ X0 @ s__Moon)) & ($true = (s__instance @ X1 @ s__Flying)) & ($true = (s__agent @ X1 @ s__Bill @ cw)) & ($true = (s__destination @ X1 @ X0 @ cw)))),
  inference(fool_elimination,[status(thm)],[f33])).
thf(f33,plain,(
  ? [X0 : $i,X1 : $i] : ((s__instance @ X1 @ s__Flying) & (s__destination @ X1 @ X0 @ cw) & (s__agent @ X1 @ s__Bill @ cw) & (s__instance @ X0 @ s__Moon))),
  inference(rectify,[status(thm)],[f27])).
thf(f27,axiom,(
  ? [X7 : $i,X6 : $i] : ((s__instance @ X6 @ s__Flying) & (s__destination @ X6 @ X7 @ cw) & (s__agent @ X6 @ s__Bill @ cw) & (s__instance @ X7 @ s__Moon))),
  file('/tmp/w6_R_Xckdu/SOT_VU3a7g',fact_flies_to_moon)).
thf(f109,plain,(
  ( ! [X0 : w] : (($true = (s__instance @ (sK4 @ X0) @ s__PayCheck)) | ($true != (s__instance @ sK2 @ s__Moon)) | ($true != (accreln1 @ s__Obligation @ cw @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f108,f81])).
thf(f81,plain,(
  ($true = (s__agent @ sK1 @ s__Bill @ cw))),
  inference(cnf_transformation,[status(thm)],[f57])).
thf(f108,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__instance @ sK2 @ s__Moon)) | ($true = (s__instance @ (sK4 @ X0) @ s__PayCheck))) )),
  inference(subsumption_resolution,[status(thm)],[f107,f82])).
thf(f82,plain,(
  ($true = (s__instance @ sK1 @ s__Flying))),
  inference(cnf_transformation,[status(thm)],[f57])).
thf(f107,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ sK1 @ s__Flying)) | ($true = (s__instance @ (sK4 @ X0) @ s__PayCheck)) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true != (s__instance @ sK2 @ s__Moon)) | ($true != (accreln1 @ s__Obligation @ cw @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f106])).
thf(f106,plain,(
  ( ! [X0 : w] : (($true = (s__instance @ (sK4 @ X0) @ s__PayCheck)) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true != $true) | ($true != (s__instance @ sK2 @ s__Moon)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__instance @ sK1 @ s__Flying))) )),
  inference(superposition,[status(thm)],[f88,f80])).
thf(f80,plain,(
  ($true = (s__destination @ sK1 @ sK2 @ cw))),
  inference(cnf_transformation,[status(thm)],[f57])).
thf(f88,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__destination @ X2 @ X3 @ X0)) | ($true = (s__instance @ (sK4 @ X1) @ s__PayCheck)) | ($true != (s__instance @ X2 @ s__Flying)) | ((s__instance @ X3 @ s__Moon) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ($true != (s__agent @ X2 @ s__Bill @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f60])).
thf(f60,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (($true != (s__instance @ X2 @ s__Flying)) | ($true != (s__destination @ X2 @ X3 @ X0)) | (($true = (s__destination @ (sK3 @ X1) @ s__Bill @ X1)) & ($true = (s__instance @ (sK3 @ X1) @ s__Giving)) & ($true = (s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1)) & ($true = (s__instance @ (sK4 @ X1) @ s__PayCheck))) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ X3 @ s__Moon) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4])],[f58,f59])).
thf(f59,plain,(
  ! [X1 : w] : (? [X4 : $i,X5 : $i] : (($true = (s__destination @ X4 @ s__Bill @ X1)) & ($true = (s__instance @ X4 @ s__Giving)) & ($true = (s__patient @ X4 @ X5 @ X1)) & ($true = (s__instance @ X5 @ s__PayCheck))) =&gt; (($true = (s__destination @ (sK3 @ X1) @ s__Bill @ X1)) & ($true = (s__instance @ (sK3 @ X1) @ s__Giving)) & ($true = (s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1)) & ($true = (s__instance @ (sK4 @ X1) @ s__PayCheck))))),
  introduced(definition,[],[choice_axiom])).
thf(f58,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (($true != (s__instance @ X2 @ s__Flying)) | ($true != (s__destination @ X2 @ X3 @ X0)) | ? [X4 : $i,X5 : $i] : (($true = (s__destination @ X4 @ s__Bill @ X1)) & ($true = (s__instance @ X4 @ s__Giving)) & ($true = (s__patient @ X4 @ X5 @ X1)) & ($true = (s__instance @ X5 @ s__PayCheck))) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ X3 @ s__Moon) != $true))),
  inference(rectify,[status(thm)],[f47])).
thf(f47,plain,(
  ! [X3 : w,X0 : w,X2 : $i,X1 : $i] : (($true != (s__instance @ X2 @ s__Flying)) | ($true != (s__destination @ X2 @ X1 @ X3)) | ? [X4 : $i,X5 : $i] : (($true = (s__destination @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X4 @ s__Giving)) & ($true = (s__patient @ X4 @ X5 @ X0)) & ($true = (s__instance @ X5 @ s__PayCheck))) | ($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (accreln1 @ s__Obligation @ X3 @ X0)) | ($true != (s__instance @ X1 @ s__Moon)))),
  inference(flattening,[status(thm)],[f46])).
thf(f46,plain,(
  ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : ((? [X4 : $i,X5 : $i] : (($true = (s__destination @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X4 @ s__Giving)) & ($true = (s__patient @ X4 @ X5 @ X0)) & ($true = (s__instance @ X5 @ s__PayCheck))) | ($true != (accreln1 @ s__Obligation @ X3 @ X0))) | (($true != (s__instance @ X2 @ s__Flying)) | ($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (s__instance @ X1 @ s__Moon)) | ($true != (s__destination @ X2 @ X1 @ X3))))),
  inference(ennf_transformation,[status(thm)],[f36])).
thf(f36,plain,(
  ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : ((($true = (s__instance @ X2 @ s__Flying)) & ($true = (s__agent @ X2 @ s__Bill @ X3)) & ($true = (s__instance @ X1 @ s__Moon)) & ($true = (s__destination @ X2 @ X1 @ X3))) =&gt; (($true = (accreln1 @ s__Obligation @ X3 @ X0)) =&gt; ? [X4 : $i,X5 : $i] : (($true = (s__destination @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X4 @ s__Giving)) & ($true = (s__patient @ X4 @ X5 @ X0)) & ($true = (s__instance @ X5 @ s__PayCheck)))))),
  inference(fool_elimination,[status(thm)],[f35])).
thf(f35,plain,(
  ! [X0 : w,X1 : $i,X2 : $i,X3 : w] : (((s__instance @ X2 @ s__Flying) & (s__destination @ X2 @ X1 @ X3) & (s__instance @ X1 @ s__Moon) & (s__agent @ X2 @ s__Bill @ X3)) =&gt; ((accreln1 @ s__Obligation @ X3 @ X0) =&gt; ? [X4 : $i,X5 : $i] : ((s__instance @ X4 @ s__Giving) & (s__patient @ X4 @ X5 @ X0) & (s__destination @ X4 @ s__Bill @ X0) & (s__instance @ X5 @ s__PayCheck))))),
  inference(rectify,[status(thm)],[f26])).
thf(f26,axiom,(
  ! [X1 : w,X7 : $i,X6 : $i,X0 : w] : (((s__instance @ X6 @ s__Flying) & (s__destination @ X6 @ X7 @ X0) & (s__instance @ X7 @ s__Moon) & (s__agent @ X6 @ s__Bill @ X0)) =&gt; ((accreln1 @ s__Obligation @ X0 @ X1) =&gt; ? [X8 : $i,X9 : $i] : ((s__instance @ X8 @ s__Giving) & (s__patient @ X8 @ X9 @ X1) & (s__destination @ X8 @ s__Bill @ X1) & (s__instance @ X9 @ s__PayCheck))))),
  file('/tmp/w6_R_Xckdu/SOT_VU3a7g',obligation2_must_be_paid)).
thf(f297,plain,(
  ($true != (s__instance @ (sK4 @ (sK5 @ cw)) @ s__PayCheck)) | ~spl6_7),
  inference(trivial_inequality_removal,[status(thm)],[f296])).
thf(f296,plain,(
  ($true != $true) | ($true != (s__instance @ (sK4 @ (sK5 @ cw)) @ s__PayCheck)) | ~spl6_7),
  inference(superposition,[status(thm)],[f229,f188])).
thf(f188,plain,(
  ($true = (s__patient @ (sK3 @ (sK5 @ cw)) @ (sK4 @ (sK5 @ cw)) @ (sK5 @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f182])).
thf(f182,plain,(
  ($true = (s__patient @ (sK3 @ (sK5 @ cw)) @ (sK4 @ (sK5 @ cw)) @ (sK5 @ cw))) | ($true != $true)),
  inference(superposition,[status(thm)],[f165,f97])).
thf(f165,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f164,f81])).
thf(f164,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0)) | ($true != (s__agent @ sK1 @ s__Bill @ cw))) )),
  inference(subsumption_resolution,[status(thm)],[f163,f83])).
thf(f163,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ sK2 @ s__Moon)) | ($true = (s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__agent @ sK1 @ s__Bill @ cw))) )),
  inference(subsumption_resolution,[status(thm)],[f162,f82])).
thf(f162,plain,(
  ( ! [X0 : w] : (($true = (s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0)) | ($true != (s__instance @ sK1 @ s__Flying)) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true != (s__instance @ sK2 @ s__Moon)) | ($true != (accreln1 @ s__Obligation @ cw @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f161])).
thf(f161,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__instance @ sK1 @ s__Flying)) | ($true = (s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0)) | ($true != (s__instance @ sK2 @ s__Moon)) | ($true != (s__agent @ sK1 @ s__Bill @ cw))) )),
  inference(superposition,[status(thm)],[f89,f80])).
thf(f89,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__destination @ X2 @ X3 @ X0)) | ((s__instance @ X3 @ s__Moon) != $true) | ($true = (s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ($true != (s__instance @ X2 @ s__Flying)) | ($true != (s__agent @ X2 @ s__Bill @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f60])).
thf(f229,plain,(
  ( ! [X0 : $i] : (($true != (s__patient @ (sK3 @ (sK5 @ cw)) @ X0 @ (sK5 @ cw))) | ($true != (s__instance @ X0 @ s__PayCheck))) ) | ~spl6_7),
  inference(avatar_component_clause,[status(thm)],[f228])).
thf(f228,definition,(
  spl6_7 &lt;=&gt; ! [X0 : $i] : (($true != (s__instance @ X0 @ s__PayCheck)) | ($true != (s__patient @ (sK3 @ (sK5 @ cw)) @ X0 @ (sK5 @ cw))))),
  introduced(definition,[new_symbols(naming,[spl6_7])],[avatar_definition])).
thf(f236,plain,(
  spl6_8 | spl6_7),
  inference(avatar_split_clause,[status(thm)],[f231,f228,f233])).
thf(f231,plain,(
  ( ! [X0 : $i] : (($true != (s__patient @ (sK3 @ (sK5 @ cw)) @ X0 @ (sK5 @ cw))) | ($true != (s__instance @ X0 @ s__PayCheck)) | ($true = (s__instance @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Tax))) )),
  inference(subsumption_resolution,[status(thm)],[f198,f171])).
thf(f171,plain,(
  ($true = (s__instance @ (sK3 @ (sK5 @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f166])).
thf(f166,plain,(
  ($true != $true) | ($true = (s__instance @ (sK3 @ (sK5 @ cw)) @ s__Giving))),
  inference(superposition,[status(thm)],[f123,f97])).
thf(f123,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__instance @ (sK3 @ X0) @ s__Giving))) )),
  inference(subsumption_resolution,[status(thm)],[f122,f81])).
thf(f122,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__instance @ (sK3 @ X0) @ s__Giving)) | ($true != (s__agent @ sK1 @ s__Bill @ cw))) )),
  inference(subsumption_resolution,[status(thm)],[f121,f82])).
thf(f121,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__instance @ sK1 @ s__Flying)) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true = (s__instance @ (sK3 @ X0) @ s__Giving))) )),
  inference(subsumption_resolution,[status(thm)],[f120,f83])).
thf(f120,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__instance @ sK2 @ s__Moon)) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true = (s__instance @ (sK3 @ X0) @ s__Giving)) | ($true != (s__instance @ sK1 @ s__Flying))) )),
  inference(trivial_inequality_removal,[status(thm)],[f119])).
thf(f119,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true != (s__instance @ sK1 @ s__Flying)) | ($true != $true) | ($true = (s__instance @ (sK3 @ X0) @ s__Giving)) | ($true != (s__instance @ sK2 @ s__Moon))) )),
  inference(superposition,[status(thm)],[f90,f80])).
thf(f90,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__destination @ X2 @ X3 @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ X3 @ s__Moon) != $true) | ($true = (s__instance @ (sK3 @ X1) @ s__Giving)) | ($true != (s__instance @ X2 @ s__Flying)) | ($true != (s__agent @ X2 @ s__Bill @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f60])).
thf(f198,plain,(
  ( ! [X0 : $i] : (($true = (s__instance @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Tax)) | ($true != (s__patient @ (sK3 @ (sK5 @ cw)) @ X0 @ (sK5 @ cw))) | ($true != (s__instance @ X0 @ s__PayCheck)) | ($true != (s__instance @ (sK3 @ (sK5 @ cw)) @ s__Giving))) )),
  inference(trivial_inequality_removal,[status(thm)],[f191])).
thf(f191,plain,(
  ( ! [X0 : $i] : (($true != $true) | ($true = (s__instance @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Tax)) | ($true != (s__patient @ (sK3 @ (sK5 @ cw)) @ X0 @ (sK5 @ cw))) | ($true != (s__instance @ X0 @ s__PayCheck)) | ($true != (s__instance @ (sK3 @ (sK5 @ cw)) @ s__Giving))) )),
  inference(superposition,[status(thm)],[f103,f178])).
thf(f178,plain,(
  ($true = (s__destination @ (sK3 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f174])).
thf(f174,plain,(
  ($true != $true) | ($true = (s__destination @ (sK3 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw)))),
  inference(superposition,[status(thm)],[f152,f97])).
thf(f152,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__destination @ (sK3 @ X0) @ s__Bill @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f151,f81])).
thf(f151,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__destination @ (sK3 @ X0) @ s__Bill @ X0)) | ($true != (s__agent @ sK1 @ s__Bill @ cw))) )),
  inference(subsumption_resolution,[status(thm)],[f150,f82])).
thf(f150,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ sK1 @ s__Flying)) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__destination @ (sK3 @ X0) @ s__Bill @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f149,f83])).
thf(f149,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ sK2 @ s__Moon)) | ($true = (s__destination @ (sK3 @ X0) @ s__Bill @ X0)) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__instance @ sK1 @ s__Flying))) )),
  inference(trivial_inequality_removal,[status(thm)],[f148])).
thf(f148,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ sK1 @ s__Flying)) | ($true != $true) | ($true != (s__agent @ sK1 @ s__Bill @ cw)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__destination @ (sK3 @ X0) @ s__Bill @ X0)) | ($true != (s__instance @ sK2 @ s__Moon))) )),
  inference(superposition,[status(thm)],[f91,f80])).
thf(f91,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__destination @ X2 @ X3 @ X0)) | ($true = (s__destination @ (sK3 @ X1) @ s__Bill @ X1)) | ((s__instance @ X3 @ s__Moon) != $true) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ($true != (s__instance @ X2 @ s__Flying)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f60])).
thf(f103,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X2 @ s__PayCheck)) | ($true != (s__instance @ X1 @ s__Giving)) | ($true = (s__instance @ (sK0 @ (sK5 @ X0)) @ s__Tax)) | ($true != (s__patient @ X1 @ X2 @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f102])).
thf(f102,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != $true) | ($true = (s__instance @ (sK0 @ (sK5 @ X0)) @ s__Tax)) | ($true != (s__patient @ X1 @ X2 @ X0)) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X2 @ s__PayCheck)) | ($true != (s__instance @ X1 @ s__Giving))) )),
  inference(superposition,[status(thm)],[f69,f97])).
thf(f69,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Obligation @ X2 @ X0)) | ($true != (s__patient @ X3 @ X1 @ X2)) | ((s__instance @ X3 @ s__Giving) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)) | ($true != (s__destination @ X3 @ s__Bill @ X2)) | ($true = (s__instance @ (sK0 @ X0) @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f54])).
thf(f54,plain,(
  ! [X0 : w,X1 : $i,X2 : w,X3 : $i] : ((($true = (s__origin @ (sK0 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK0 @ X0) @ s__Tax))) | ($true != (accreln1 @ s__Obligation @ X2 @ X0)) | ($true != (s__patient @ X3 @ X1 @ X2)) | ($true != (s__destination @ X3 @ s__Bill @ X2)) | ((s__instance @ X3 @ s__Giving) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f52,f53])).
thf(f53,plain,(
  ! [X0 : w] : (? [X4 : $i] : (($true = (s__origin @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X4 @ s__Tax))) =&gt; (($true = (s__origin @ (sK0 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK0 @ X0) @ s__Tax))))),
  introduced(definition,[],[choice_axiom])).
thf(f52,plain,(
  ! [X0 : w,X1 : $i,X2 : w,X3 : $i] : (? [X4 : $i] : (($true = (s__origin @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X4 @ s__Tax))) | ($true != (accreln1 @ s__Obligation @ X2 @ X0)) | ($true != (s__patient @ X3 @ X1 @ X2)) | ($true != (s__destination @ X3 @ s__Bill @ X2)) | ((s__instance @ X3 @ s__Giving) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)))),
  inference(rectify,[status(thm)],[f51])).
thf(f51,plain,(
  ! [X1 : w,X3 : $i,X2 : w,X0 : $i] : (? [X4 : $i] : (($true = (s__origin @ X4 @ s__Bill @ X1)) & ($true = (s__instance @ X4 @ s__Tax))) | ($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ($true != (s__patient @ X0 @ X3 @ X2)) | ($true != (s__destination @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X0 @ s__Giving)) | ($true != (s__instance @ X3 @ s__PayCheck)))),
  inference(flattening,[status(thm)],[f50])).
thf(f50,plain,(
  ! [X2 : w,X0 : $i,X3 : $i,X1 : w] : ((? [X4 : $i] : (($true = (s__origin @ X4 @ s__Bill @ X1)) & ($true = (s__instance @ X4 @ s__Tax))) | ($true != (accreln1 @ s__Obligation @ X2 @ X1))) | (($true != (s__patient @ X0 @ X3 @ X2)) | ($true != (s__instance @ X3 @ s__PayCheck)) | ($true != (s__instance @ X0 @ s__Giving)) | ($true != (s__destination @ X0 @ s__Bill @ X2))))),
  inference(ennf_transformation,[status(thm)],[f40])).
thf(f40,plain,(
  ! [X2 : w,X0 : $i,X3 : $i,X1 : w] : ((($true = (s__patient @ X0 @ X3 @ X2)) & ($true = (s__instance @ X3 @ s__PayCheck)) & ($true = (s__instance @ X0 @ s__Giving)) & ($true = (s__destination @ X0 @ s__Bill @ X2))) =&gt; (($true = (accreln1 @ s__Obligation @ X2 @ X1)) =&gt; ? [X4 : $i] : (($true = (s__origin @ X4 @ s__Bill @ X1)) & ($true = (s__instance @ X4 @ s__Tax)))))),
  inference(fool_elimination,[status(thm)],[f39])).
thf(f39,plain,(
  ! [X0 : $i,X1 : w,X2 : w,X3 : $i] : (((s__instance @ X3 @ s__PayCheck) & (s__instance @ X0 @ s__Giving) & (s__patient @ X0 @ X3 @ X2) & (s__destination @ X0 @ s__Bill @ X2)) =&gt; ((accreln1 @ s__Obligation @ X2 @ X1) =&gt; ? [X4 : $i] : ((s__origin @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Tax))))),
  inference(rectify,[status(thm)],[f25])).
thf(f25,axiom,(
  ! [X3 : $i,X1 : w,X0 : w,X4 : $i] : (((s__instance @ X4 @ s__PayCheck) & (s__instance @ X3 @ s__Giving) & (s__patient @ X3 @ X4 @ X0) & (s__destination @ X3 @ s__Bill @ X0)) =&gt; ((accreln1 @ s__Obligation @ X0 @ X1) =&gt; ? [X5 : $i] : ((s__origin @ X5 @ s__Bill @ X1) & (s__instance @ X5 @ s__Tax))))),
  file('/tmp/w6_R_Xckdu/SOT_VU3a7g',obligation1_bill_taxed)).
thf(f230,plain,(
  spl6_6 | spl6_7),
  inference(avatar_split_clause,[status(thm)],[f222,f228,f224])).
thf(f222,plain,(
  ( ! [X0 : $i] : (($true != (s__instance @ X0 @ s__PayCheck)) | ($true = (s__origin @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != (s__patient @ (sK3 @ (sK5 @ cw)) @ X0 @ (sK5 @ cw)))) )),
  inference(subsumption_resolution,[status(thm)],[f199,f171])).
thf(f199,plain,(
  ( ! [X0 : $i] : (($true != (s__instance @ X0 @ s__PayCheck)) | ($true != (s__instance @ (sK3 @ (sK5 @ cw)) @ s__Giving)) | ($true = (s__origin @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != (s__patient @ (sK3 @ (sK5 @ cw)) @ X0 @ (sK5 @ cw)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f190])).
thf(f190,plain,(
  ( ! [X0 : $i] : (($true != (s__instance @ X0 @ s__PayCheck)) | ($true != (s__instance @ (sK3 @ (sK5 @ cw)) @ s__Giving)) | ($true != $true) | ($true = (s__origin @ (sK0 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != (s__patient @ (sK3 @ (sK5 @ cw)) @ X0 @ (sK5 @ cw)))) )),
  inference(superposition,[status(thm)],[f137,f178])).
thf(f137,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X1 @ s__Giving)) | ($true != (s__instance @ X2 @ s__PayCheck)) | ($true = (s__origin @ (sK0 @ (sK5 @ X0)) @ s__Bill @ (sK5 @ X0))) | ($true != (s__patient @ X1 @ X2 @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f132])).
thf(f132,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != $true) | ($true != (s__patient @ X1 @ X2 @ X0)) | ($true != (s__instance @ X1 @ s__Giving)) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true = (s__origin @ (sK0 @ (sK5 @ X0)) @ s__Bill @ (sK5 @ X0))) | ($true != (s__instance @ X2 @ s__PayCheck))) )),
  inference(superposition,[status(thm)],[f70,f97])).
thf(f70,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Obligation @ X2 @ X0)) | ($true != (s__patient @ X3 @ X1 @ X2)) | ($true != (s__destination @ X3 @ s__Bill @ X2)) | ($true = (s__origin @ (sK0 @ X0) @ s__Bill @ X0)) | ((s__instance @ X3 @ s__Giving) != $true) | ($true != (s__instance @ X1 @ s__PayCheck))) )),
  inference(cnf_transformation,[status(thm)],[f54])).
% SZS output end Proof for SOT_VU3a7g
% (2603650)------------------------------
% (2603650)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2603650)Termination reason: Refutation

% (2603650)Memory used [KB]: 5884
% (2603650)Time elapsed: 0.054 s
% (2603650)------------------------------
% (2603650)------------------------------
% (2603643)Success in time 0.145 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_VU3a7g - Vampire---5.0.1 says Theorem - CPU = 0.02 WC = 0.27 
% OUTPUT: SOT_VU3a7g - Vampire---5.0.1 says Refutation - CPU = 0.02 WC = 0.27 
