 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/FKSbjIJp9N/SOT_UKiqcb
Permission denied
% (3746869)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_UKiqcb for (1ds/275Mi)
% (3746866)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_UKiqcb for (1ds/27Mi)
% (3746870)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_UKiqcb for (1ds/18Mi)
% (3746867)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_UKiqcb for (1ds/2Mi)
% (3746864)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_UKiqcb for (1ds/183Mi)
% (3746868)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_UKiqcb for (1ds/2Mi)
% (3746868)Refutation not found, incomplete strategy
% (3746868)------------------------------
% (3746868)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (3746868)Termination reason: Refutation not found, incomplete strategy


% (3746868)Memory used [KB]: 5628
% (3746868)Time elapsed: 0.005 s
% (3746868)------------------------------
% (3746868)------------------------------
% (3746869)Refutation not found, incomplete strategy
% (3746869)------------------------------
% (3746869)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (3746869)Termination reason: Refutation not found, incomplete strategy


% (3746869)Memory used [KB]: 5628
% (3746869)Time elapsed: 0.011 s
% (3746869)------------------------------
% (3746869)------------------------------
% (3746871)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_UKiqcb for (1ds/3Mi)
% (3746864)First to succeed.
% (3746872)lrs+1002_1:1_cnfonf=lazy_not_be_gen:hud=14:prag=on:sp=weighted_frequency:tnu=1:i=37:si=on:rtra=on_0 on SOT_UKiqcb for (1ds/37Mi)
% (3746864)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_UKiqcb
% SZS output start Proof for SOT_UKiqcb
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m > w > w > $o).
thf(func_def_5, type, accreln2: m > $i > w > w > $o).
thf(func_def_6, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_8, type, s__instance: $i > $i > $o).
thf(func_def_9, type, s__agent: $i > $i > w > $o).
thf(func_def_10, type, s__patient: $i > $i > w > $o).
thf(func_def_11, type, s__destination: $i > $i > w > $o).
thf(func_def_12, type, s__located: $i > $i > w > $o).
thf(func_def_24, type, sK0: w > $i).
thf(func_def_25, type, sK1: w > $i).
thf(func_def_26, type, sK2: w > w).
thf(func_def_28, type, sK4: w > $i).
thf(func_def_29, type, sK5: w > $i).
thf(func_def_30, type, sK6: w > $i).
thf(func_def_31, type, sK7: w > $i).
thf(f261,plain,(
  $false),
  inference(avatar_sat_refutation,[status(thm)],[f144,f148,f152,f156,f241,f260])).
thf(f260,plain,(
  ~spl8_1 | ~spl8_3 | ~spl8_4 | ~spl8_5),
  inference(avatar_contradiction_clause,[status(thm)],[f259])).
thf(f259,plain,(
  $false | (~spl8_1 | ~spl8_3 | ~spl8_4 | ~spl8_5)),
  inference(subsumption_resolution,[status(thm)],[f258,f151])).
thf(f151,plain,(
  ( ! [X1 : w] : (((s__instance @ (sK0 @ X1) @ s__Keeping) = $true)) ) | ~spl8_4),
  inference(avatar_component_clause,[status(thm)],[f150])).
thf(f150,definition,(
  spl8_4 <=> ! [X1 : w] : ((s__instance @ (sK0 @ X1) @ s__Keeping) = $true)),
  introduced(definition,[new_symbols(naming,[spl8_4])],[avatar_definition])).
thf(f258,plain,(
  ((s__instance @ (sK0 @ (sK2 @ cw)) @ s__Keeping) != $true) | (~spl8_1 | ~spl8_3 | ~spl8_5)),
  inference(subsumption_resolution,[status(thm)],[f257,f155])).
thf(f155,plain,(
  ( ! [X1 : w] : (((s__instance @ (sK1 @ X1) @ s__ParkingLot) = $true)) ) | ~spl8_5),
  inference(avatar_component_clause,[status(thm)],[f154])).
thf(f154,definition,(
  spl8_5 <=> ! [X1 : w] : ((s__instance @ (sK1 @ X1) @ s__ParkingLot) = $true)),
  introduced(definition,[new_symbols(naming,[spl8_5])],[avatar_definition])).
thf(f257,plain,(
  ((s__instance @ (sK1 @ (sK2 @ cw)) @ s__ParkingLot) != $true) | ((s__instance @ (sK0 @ (sK2 @ cw)) @ s__Keeping) != $true) | (~spl8_1 | ~spl8_3)),
  inference(subsumption_resolution,[status(thm)],[f251,f147])).
thf(f147,plain,(
  ( ! [X1 : w] : (((s__agent @ (sK0 @ X1) @ s__Bill @ X1) = $true)) ) | ~spl8_3),
  inference(avatar_component_clause,[status(thm)],[f146])).
thf(f146,definition,(
  spl8_3 <=> ! [X1 : w] : ((s__agent @ (sK0 @ X1) @ s__Bill @ X1) = $true)),
  introduced(definition,[new_symbols(naming,[spl8_3])],[avatar_definition])).
thf(f251,plain,(
  ($true != (s__agent @ (sK0 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ((s__instance @ (sK1 @ (sK2 @ cw)) @ s__ParkingLot) != $true) | ((s__instance @ (sK0 @ (sK2 @ cw)) @ s__Keeping) != $true) | ~spl8_1),
  inference(trivial_inequality_removal,[status(thm)],[f249])).
thf(f249,plain,(
  ((s__instance @ (sK1 @ (sK2 @ cw)) @ s__ParkingLot) != $true) | ($true != $true) | ((s__instance @ (sK0 @ (sK2 @ cw)) @ s__Keeping) != $true) | ($true != (s__agent @ (sK0 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ~spl8_1),
  inference(superposition,[status(thm)],[f186,f140])).
thf(f140,plain,(
  ( ! [X1 : w] : (((s__located @ (sK0 @ X1) @ (sK1 @ X1) @ X1) = $true)) ) | ~spl8_1),
  inference(avatar_component_clause,[status(thm)],[f139])).
thf(f139,definition,(
  spl8_1 <=> ! [X1 : w] : ((s__located @ (sK0 @ X1) @ (sK1 @ X1) @ X1) = $true)),
  introduced(definition,[new_symbols(naming,[spl8_1])],[avatar_definition])).
thf(f186,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__located @ X0 @ X1 @ (sK2 @ cw))) | ((s__agent @ X0 @ s__Bill @ (sK2 @ cw)) != $true) | ((s__instance @ X0 @ s__Keeping) != $true) | ((s__instance @ X1 @ s__ParkingLot) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f184])).
thf(f184,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__located @ X0 @ X1 @ (sK2 @ cw))) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ((s__agent @ X0 @ s__Bill @ (sK2 @ cw)) != $true) | ((s__instance @ X0 @ s__Keeping) != $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f181,f95])).
thf(f95,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ X0 @ (sK2 @ X0)) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f75])).
thf(f75,plain,(
  ! [X0 : w] : ((accreln1 @ s__Obligation @ X0 @ (sK2 @ X0)) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f56,f74])).
thf(f74,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ((accreln1 @ s__Obligation @ X0 @ (sK2 @ X0)) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f56,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f55])).
thf(f55,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  inference(rectify,[status(thm)],[f4])).
thf(f4,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  file('/tmp/FKSbjIJp9N/SOT_UKiqcb',serial_obligation)).
thf(f181,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ((s__located @ X2 @ X1 @ X0) != $true) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ($true != (s__instance @ X2 @ s__Keeping))) )),
  inference(trivial_inequality_removal,[status(thm)],[f180])).
thf(f180,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__located @ X2 @ X1 @ X0) != $true) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ($true != (s__instance @ X2 @ s__Keeping)) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ($true != $true) | ($true != (accreln1 @ s__Obligation @ cw @ X0))) )),
  inference(superposition,[status(thm)],[f127,f99])).
thf(f99,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Permission @ X1 @ X0) = $true) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f76])).
thf(f76,plain,(
  ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ((accreln1 @ s__Permission @ X1 @ X0) = $true))),
  inference(rectify,[status(thm)],[f65])).
thf(f65,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((accreln1 @ s__Permission @ X0 @ X1) = $true))),
  inference(ennf_transformation,[status(thm)],[f54])).
thf(f54,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ((accreln1 @ s__Permission @ X0 @ X1) = $true))),
  inference(fool_elimination,[status(thm)],[f53])).
thf(f53,plain,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  inference(rectify,[status(thm)],[f6])).
thf(f6,axiom,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  file('/tmp/FKSbjIJp9N/SOT_UKiqcb',bridge_obligation_permission)).
thf(f127,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Permission @ cw @ X0)) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ((s__located @ X2 @ X1 @ X0) != $true) | ($true != (s__instance @ X2 @ s__Keeping)) | ($true != (s__agent @ X2 @ s__Bill @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f85])).
thf(f85,plain,(
  ! [X0 : w] : (($true != (accreln1 @ s__Permission @ cw @ X0)) | ! [X1 : $i,X2 : $i] : (((s__instance @ X1 @ s__ParkingLot) != $true) | ((s__located @ X2 @ X1 @ X0) != $true) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ($true != (s__instance @ X2 @ s__Keeping))))),
  inference(rectify,[status(thm)],[f70])).
thf(f70,plain,(
  ! [X0 : w] : (($true != (accreln1 @ s__Permission @ cw @ X0)) | ! [X2 : $i,X1 : $i] : (((s__instance @ X2 @ s__ParkingLot) != $true) | ((s__located @ X1 @ X2 @ X0) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__Keeping) != $true)))),
  inference(ennf_transformation,[status(thm)],[f48])).
thf(f48,plain,(
  ~? [X0 : w] : (? [X1 : $i,X2 : $i] : (((s__instance @ X2 @ s__ParkingLot) = $true) & ((s__agent @ X1 @ s__Bill @ X0) = $true) & ((s__located @ X1 @ X2 @ X0) = $true) & ((s__instance @ X1 @ s__Keeping) = $true)) & ($true = (accreln1 @ s__Permission @ cw @ X0)))),
  inference(fool_elimination,[status(thm)],[f47])).
thf(f47,plain,(
  ~? [X0 : w] : ((accreln1 @ s__Permission @ cw @ X0) & ? [X1 : $i,X2 : $i] : ((s__instance @ X2 @ s__ParkingLot) & (s__located @ X1 @ X2 @ X0) & (s__instance @ X1 @ s__Keeping) & (s__agent @ X1 @ s__Bill @ X0)))),
  inference(rectify,[status(thm)],[f41])).
thf(f41,negated_conjecture,(
  ~? [X10 : w] : ((accreln1 @ s__Permission @ cw @ X10) & ? [X8 : $i,X9 : $i] : ((s__instance @ X9 @ s__ParkingLot) & (s__located @ X8 @ X9 @ X10) & (s__instance @ X8 @ s__Keeping) & (s__agent @ X8 @ s__Bill @ X10)))),
  inference(negated_conjecture,[status(cth)],[f40])).
thf(f40,conjecture,(
  ? [X10 : w] : ((accreln1 @ s__Permission @ cw @ X10) & ? [X8 : $i,X9 : $i] : ((s__instance @ X9 @ s__ParkingLot) & (s__located @ X8 @ X9 @ X10) & (s__instance @ X8 @ s__Keeping) & (s__agent @ X8 @ s__Bill @ X10)))),
  file('/tmp/FKSbjIJp9N/SOT_UKiqcb',conjecture_is_permitted)).
thf(f241,plain,(
  ~spl8_2),
  inference(avatar_contradiction_clause,[status(thm)],[f240])).
thf(f240,plain,(
  $false | ~spl8_2),
  inference(subsumption_resolution,[status(thm)],[f239,f101])).
thf(f101,plain,(
  ($true = (s__instance @ sK3 @ s__PayCheck))),
  inference(cnf_transformation,[status(thm)],[f78])).
thf(f78,plain,(
  ((s__destination @ sK3 @ s__Bill @ cw) = $true) & ($true = (s__instance @ sK3 @ s__PayCheck))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f60,f77])).
thf(f77,plain,(
  ? [X0 : $i] : (((s__destination @ X0 @ s__Bill @ cw) = $true) & ((s__instance @ X0 @ s__PayCheck) = $true)) => (((s__destination @ sK3 @ s__Bill @ cw) = $true) & ($true = (s__instance @ sK3 @ s__PayCheck)))),
  introduced(definition,[],[choice_axiom])).
thf(f60,plain,(
  ? [X0 : $i] : (((s__destination @ X0 @ s__Bill @ cw) = $true) & ((s__instance @ X0 @ s__PayCheck) = $true))),
  inference(fool_elimination,[status(thm)],[f59])).
thf(f59,plain,(
  ? [X0 : $i] : ((s__instance @ X0 @ s__PayCheck) & (s__destination @ X0 @ s__Bill @ cw))),
  inference(rectify,[status(thm)],[f39])).
thf(f39,axiom,(
  ? [X3 : $i] : ((s__instance @ X3 @ s__PayCheck) & (s__destination @ X3 @ s__Bill @ cw))),
  file('/tmp/FKSbjIJp9N/SOT_UKiqcb',optionA_bill_gets_paycheck)).
thf(f239,plain,(
  ($true != (s__instance @ sK3 @ s__PayCheck)) | ~spl8_2),
  inference(trivial_inequality_removal,[status(thm)],[f238])).
thf(f238,plain,(
  ($true != (s__instance @ sK3 @ s__PayCheck)) | ($true != $true) | ~spl8_2),
  inference(superposition,[status(thm)],[f236,f102])).
thf(f102,plain,(
  ((s__destination @ sK3 @ s__Bill @ cw) = $true)),
  inference(cnf_transformation,[status(thm)],[f78])).
thf(f236,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true)) ) | ~spl8_2),
  inference(trivial_inequality_removal,[status(thm)],[f234])).
thf(f234,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true)) ) | ~spl8_2),
  inference(superposition,[status(thm)],[f232,f95])).
thf(f232,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ (sK2 @ X0) @ X2) != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true)) ) | ~spl8_2),
  inference(trivial_inequality_removal,[status(thm)],[f230])).
thf(f230,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != $true) | ((accreln1 @ s__Obligation @ (sK2 @ X0) @ X2) != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true)) ) | ~spl8_2),
  inference(superposition,[status(thm)],[f229,f95])).
thf(f229,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ((s__destination @ X3 @ s__Bill @ X2) != $true) | ((s__instance @ X3 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true)) ) | ~spl8_2),
  inference(subsumption_resolution,[status(thm)],[f228,f116])).
thf(f116,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (($true = (s__agent @ (sK7 @ X1) @ s__Bill @ X1)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true) | ((s__destination @ X2 @ s__Bill @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f84])).
thf(f84,plain,(
  ! [X0 : w,X1 : w,X2 : $i] : (((s__destination @ X2 @ s__Bill @ X0) != $true) | (($true = (s__agent @ (sK7 @ X1) @ s__Bill @ X1)) & ($true = (s__instance @ (sK6 @ X1) @ s__Tax)) & ((s__patient @ (sK7 @ X1) @ (sK6 @ X1) @ X1) = $true) & ((s__instance @ (sK7 @ X1) @ s__Giving) = $true)) | ((s__instance @ X2 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f82,f83])).
thf(f83,plain,(
  ! [X1 : w] : (? [X3 : $i,X4 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X3 @ s__Tax)) & ((s__patient @ X4 @ X3 @ X1) = $true) & ((s__instance @ X4 @ s__Giving) = $true)) => (($true = (s__agent @ (sK7 @ X1) @ s__Bill @ X1)) & ($true = (s__instance @ (sK6 @ X1) @ s__Tax)) & ((s__patient @ (sK7 @ X1) @ (sK6 @ X1) @ X1) = $true) & ((s__instance @ (sK7 @ X1) @ s__Giving) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f82,plain,(
  ! [X0 : w,X1 : w,X2 : $i] : (((s__destination @ X2 @ s__Bill @ X0) != $true) | ? [X3 : $i,X4 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X3 @ s__Tax)) & ((s__patient @ X4 @ X3 @ X1) = $true) & ((s__instance @ X4 @ s__Giving) = $true)) | ((s__instance @ X2 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true))),
  inference(rectify,[status(thm)],[f64])).
thf(f64,plain,(
  ! [X2 : w,X1 : w,X0 : $i] : (((s__destination @ X0 @ s__Bill @ X2) != $true) | ? [X3 : $i,X4 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X3 @ s__Tax)) & ((s__patient @ X4 @ X3 @ X1) = $true) & ((s__instance @ X4 @ s__Giving) = $true)) | ((s__instance @ X0 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X2 @ X1) != $true))),
  inference(flattening,[status(thm)],[f63])).
thf(f63,plain,(
  ! [X2 : w,X1 : w,X0 : $i] : ((? [X3 : $i,X4 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X3 @ s__Tax)) & ((s__patient @ X4 @ X3 @ X1) = $true) & ((s__instance @ X4 @ s__Giving) = $true)) | ((accreln1 @ s__Obligation @ X2 @ X1) != $true)) | (((s__instance @ X0 @ s__PayCheck) != $true) | ((s__destination @ X0 @ s__Bill @ X2) != $true)))),
  inference(ennf_transformation,[status(thm)],[f50])).
thf(f50,plain,(
  ! [X2 : w,X1 : w,X0 : $i] : ((((s__instance @ X0 @ s__PayCheck) = $true) & ((s__destination @ X0 @ s__Bill @ X2) = $true)) => (((accreln1 @ s__Obligation @ X2 @ X1) = $true) => ? [X3 : $i,X4 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X3 @ s__Tax)) & ((s__patient @ X4 @ X3 @ X1) = $true) & ((s__instance @ X4 @ s__Giving) = $true))))),
  inference(fool_elimination,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (((s__destination @ X0 @ s__Bill @ X2) & (s__instance @ X0 @ s__PayCheck)) => ((accreln1 @ s__Obligation @ X2 @ X1) => ? [X3 : $i,X4 : $i] : ((s__instance @ X3 @ s__Tax) & (s__instance @ X4 @ s__Giving) & (s__agent @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X3 @ X1))))),
  inference(rectify,[status(thm)],[f36])).
thf(f36,axiom,(
  ! [X3 : $i,X1 : w,X0 : w] : (((s__destination @ X3 @ s__Bill @ X0) & (s__instance @ X3 @ s__PayCheck)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X5 : $i,X4 : $i] : ((s__instance @ X5 @ s__Tax) & (s__instance @ X4 @ s__Giving) & (s__agent @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1))))),
  file('/tmp/FKSbjIJp9N/SOT_UKiqcb',obligation1_pay_taxes)).
thf(f228,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ((s__instance @ X3 @ s__PayCheck) != $true) | ($true != (s__agent @ (sK7 @ X0) @ s__Bill @ X0)) | ((s__destination @ X3 @ s__Bill @ X2) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true)) ) | ~spl8_2),
  inference(subsumption_resolution,[status(thm)],[f227,f113])).
thf(f113,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ (sK7 @ X1) @ s__Giving) = $true) | ((s__destination @ X2 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f84])).
thf(f227,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : w] : (((s__instance @ X3 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__destination @ X3 @ s__Bill @ X2) != $true) | ($true != (s__agent @ (sK7 @ X0) @ s__Bill @ X0)) | ((s__instance @ (sK7 @ X0) @ s__Giving) != $true) | ((accreln1 @ s__Obligation @ X2 @ X0) != $true)) ) | ~spl8_2),
  inference(subsumption_resolution,[status(thm)],[f226,f115])).
thf(f115,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ($true = (s__instance @ (sK6 @ X1) @ s__Tax)) | ((s__destination @ X2 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f84])).
thf(f226,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : w] : (($true != (s__agent @ (sK7 @ X0) @ s__Bill @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ (sK6 @ X0) @ s__Tax) != $true) | ((s__instance @ (sK7 @ X0) @ s__Giving) != $true) | ((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ((s__destination @ X3 @ s__Bill @ X2) != $true) | ((s__instance @ X3 @ s__PayCheck) != $true)) ) | ~spl8_2),
  inference(trivial_inequality_removal,[status(thm)],[f225])).
thf(f225,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ((s__instance @ (sK6 @ X0) @ s__Tax) != $true) | ((s__destination @ X3 @ s__Bill @ X2) != $true) | ((s__instance @ (sK7 @ X0) @ s__Giving) != $true) | ($true != (s__agent @ (sK7 @ X0) @ s__Bill @ X0)) | ((s__instance @ X3 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ($true != $true)) ) | ~spl8_2),
  inference(superposition,[status(thm)],[f224,f114])).
thf(f114,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (((s__patient @ (sK7 @ X1) @ (sK6 @ X1) @ X1) = $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__destination @ X2 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f84])).
thf(f224,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__patient @ X2 @ X3 @ X1)) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ($true != (s__instance @ X3 @ s__Tax))) ) | ~spl8_2),
  inference(subsumption_resolution,[status(thm)],[f223,f107])).
thf(f107,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (((s__agent @ (sK4 @ X2) @ s__Bill @ X2) = $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__patient @ X0 @ X1 @ X3) != $true) | ((s__instance @ X0 @ s__Giving) != $true) | ((s__agent @ X0 @ s__Bill @ X3) != $true) | ((accreln1 @ s__Obligation @ X3 @ X2) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f81])).
thf(f81,plain,(
  ! [X0 : $i,X1 : $i,X2 : w,X3 : w] : (((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__Giving) != $true) | ((s__patient @ X0 @ X1 @ X3) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((accreln1 @ s__Obligation @ X3 @ X2) != $true) | (((s__instance @ (sK5 @ X2) @ s__OfficeBuilding) = $true) & ((s__agent @ (sK4 @ X2) @ s__Bill @ X2) = $true) & ((s__instance @ (sK4 @ X2) @ s__Driving) = $true) & ((s__destination @ (sK4 @ X2) @ (sK5 @ X2) @ X2) = $true)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f79,f80])).
thf(f80,plain,(
  ! [X2 : w] : (? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__OfficeBuilding) = $true) & ((s__agent @ X4 @ s__Bill @ X2) = $true) & ((s__instance @ X4 @ s__Driving) = $true) & ($true = (s__destination @ X4 @ X5 @ X2))) => (((s__instance @ (sK5 @ X2) @ s__OfficeBuilding) = $true) & ((s__agent @ (sK4 @ X2) @ s__Bill @ X2) = $true) & ((s__instance @ (sK4 @ X2) @ s__Driving) = $true) & ((s__destination @ (sK4 @ X2) @ (sK5 @ X2) @ X2) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f79,plain,(
  ! [X0 : $i,X1 : $i,X2 : w,X3 : w] : (((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__Giving) != $true) | ((s__patient @ X0 @ X1 @ X3) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((accreln1 @ s__Obligation @ X3 @ X2) != $true) | ? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__OfficeBuilding) = $true) & ((s__agent @ X4 @ s__Bill @ X2) = $true) & ((s__instance @ X4 @ s__Driving) = $true) & ($true = (s__destination @ X4 @ X5 @ X2))))),
  inference(rectify,[status(thm)],[f69])).
thf(f69,plain,(
  ! [X2 : $i,X0 : $i,X3 : w,X1 : w] : (((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__patient @ X2 @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__Tax) != $true) | ((accreln1 @ s__Obligation @ X1 @ X3) != $true) | ? [X5 : $i,X4 : $i] : (((s__instance @ X4 @ s__OfficeBuilding) = $true) & ((s__agent @ X5 @ s__Bill @ X3) = $true) & ($true = (s__instance @ X5 @ s__Driving)) & ((s__destination @ X5 @ X4 @ X3) = $true)))),
  inference(flattening,[status(thm)],[f68])).
thf(f68,plain,(
  ! [X0 : $i,X3 : w,X2 : $i,X1 : w] : ((? [X5 : $i,X4 : $i] : (((s__instance @ X4 @ s__OfficeBuilding) = $true) & ((s__agent @ X5 @ s__Bill @ X3) = $true) & ($true = (s__instance @ X5 @ s__Driving)) & ((s__destination @ X5 @ X4 @ X3) = $true)) | ((accreln1 @ s__Obligation @ X1 @ X3) != $true)) | (((s__instance @ X2 @ s__Giving) != $true) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__patient @ X2 @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__Tax) != $true)))),
  inference(ennf_transformation,[status(thm)],[f52])).
thf(f52,plain,(
  ! [X0 : $i,X3 : w,X2 : $i,X1 : w] : ((((s__instance @ X2 @ s__Giving) = $true) & ((s__agent @ X2 @ s__Bill @ X1) = $true) & ((s__patient @ X2 @ X0 @ X1) = $true) & ((s__instance @ X0 @ s__Tax) = $true)) => (((accreln1 @ s__Obligation @ X1 @ X3) = $true) => ? [X5 : $i,X4 : $i] : (((s__instance @ X4 @ s__OfficeBuilding) = $true) & ((s__agent @ X5 @ s__Bill @ X3) = $true) & ($true = (s__instance @ X5 @ s__Driving)) & ((s__destination @ X5 @ X4 @ X3) = $true))))),
  inference(fool_elimination,[status(thm)],[f51])).
thf(f51,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__patient @ X2 @ X0 @ X1) & (s__instance @ X2 @ s__Giving) & (s__agent @ X2 @ s__Bill @ X1) & (s__instance @ X0 @ s__Tax)) => ((accreln1 @ s__Obligation @ X1 @ X3) => ? [X4 : $i,X5 : $i] : ((s__destination @ X5 @ X4 @ X3) & (s__instance @ X4 @ s__OfficeBuilding) & (s__agent @ X5 @ s__Bill @ X3) & (s__instance @ X5 @ s__Driving))))),
  inference(rectify,[status(thm)],[f37])).
thf(f37,axiom,(
  ! [X5 : $i,X0 : w,X4 : $i,X1 : w] : (((s__patient @ X4 @ X5 @ X0) & (s__instance @ X4 @ s__Giving) & (s__agent @ X4 @ s__Bill @ X0) & (s__instance @ X5 @ s__Tax)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X7 : $i,X6 : $i] : ((s__destination @ X6 @ X7 @ X1) & (s__instance @ X7 @ s__OfficeBuilding) & (s__agent @ X6 @ s__Bill @ X1) & (s__instance @ X6 @ s__Driving))))),
  file('/tmp/FKSbjIJp9N/SOT_UKiqcb',obligation2_drive_to_office)).
thf(f223,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__patient @ X2 @ X3 @ X1)) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ($true != (s__instance @ X3 @ s__Tax)) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__agent @ (sK4 @ X0) @ s__Bill @ X0) != $true)) ) | ~spl8_2),
  inference(subsumption_resolution,[status(thm)],[f222,f106])).
thf(f106,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (((accreln1 @ s__Obligation @ X3 @ X2) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__Giving) != $true) | ((s__instance @ (sK4 @ X2) @ s__Driving) = $true) | ((s__patient @ X0 @ X1 @ X3) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f81])).
thf(f222,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (s__patient @ X2 @ X3 @ X1)) | ((s__instance @ (sK4 @ X0) @ s__Driving) != $true) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ($true != (s__instance @ X3 @ s__Tax)) | ((s__agent @ (sK4 @ X0) @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__Giving) != $true)) ) | ~spl8_2),
  inference(subsumption_resolution,[status(thm)],[f221,f108])).
thf(f108,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (((accreln1 @ s__Obligation @ X3 @ X2) != $true) | ((s__patient @ X0 @ X1 @ X3) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__instance @ (sK5 @ X2) @ s__OfficeBuilding) = $true) | ((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__Giving) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f81])).
thf(f221,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((s__instance @ (sK5 @ X0) @ s__OfficeBuilding) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ($true != (s__instance @ X3 @ s__Tax)) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ((s__instance @ (sK4 @ X0) @ s__Driving) != $true) | ((s__agent @ (sK4 @ X0) @ s__Bill @ X0) != $true) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ($true != (s__patient @ X2 @ X3 @ X1))) ) | ~spl8_2),
  inference(trivial_inequality_removal,[status(thm)],[f220])).
thf(f220,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((s__agent @ X2 @ s__Bill @ X1) != $true) | ($true != (s__patient @ X2 @ X3 @ X1)) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ((s__instance @ (sK4 @ X0) @ s__Driving) != $true) | ((s__instance @ (sK5 @ X0) @ s__OfficeBuilding) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ($true != (s__instance @ X3 @ s__Tax)) | ($true != $true) | ((s__agent @ (sK4 @ X0) @ s__Bill @ X0) != $true)) ) | ~spl8_2),
  inference(superposition,[status(thm)],[f143,f105])).
thf(f105,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (((s__destination @ (sK4 @ X2) @ (sK5 @ X2) @ X2) = $true) | ((accreln1 @ s__Obligation @ X3 @ X2) != $true) | ((s__instance @ X0 @ s__Giving) != $true) | ((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__patient @ X0 @ X1 @ X3) != $true) | ((s__instance @ X1 @ s__Tax) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f81])).
thf(f143,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i] : (((s__destination @ X2 @ X0 @ X3) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) ) | ~spl8_2),
  inference(avatar_component_clause,[status(thm)],[f142])).
thf(f142,definition,(
  spl8_2 <=> ! [X2 : $i,X0 : $i,X3 : w] : (((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((s__destination @ X2 @ X0 @ X3) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true))),
  introduced(definition,[new_symbols(naming,[spl8_2])],[avatar_definition])).
thf(f156,plain,(
  spl8_2 | spl8_5),
  inference(avatar_split_clause,[status(thm)],[f93,f154,f142])).
thf(f93,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (((s__instance @ X2 @ s__Driving) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ((s__destination @ X2 @ X0 @ X3) != $true) | ((s__instance @ (sK1 @ X1) @ s__ParkingLot) = $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f73])).
thf(f73,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : ((((accreln1 @ s__Permission @ X3 @ X1) = $true) & (((s__instance @ (sK1 @ X1) @ s__ParkingLot) = $true) & ((s__located @ (sK0 @ X1) @ (sK1 @ X1) @ X1) = $true) & ((s__agent @ (sK0 @ X1) @ s__Bill @ X1) = $true) & ((s__instance @ (sK0 @ X1) @ s__Keeping) = $true))) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((s__destination @ X2 @ X0 @ X3) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f71,f72])).
thf(f72,plain,(
  ! [X1 : w] : (? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__ParkingLot) = $true) & ((s__located @ X4 @ X5 @ X1) = $true) & ((s__agent @ X4 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X4 @ s__Keeping))) => (((s__instance @ (sK1 @ X1) @ s__ParkingLot) = $true) & ((s__located @ (sK0 @ X1) @ (sK1 @ X1) @ X1) = $true) & ((s__agent @ (sK0 @ X1) @ s__Bill @ X1) = $true) & ((s__instance @ (sK0 @ X1) @ s__Keeping) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f71,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : ((((accreln1 @ s__Permission @ X3 @ X1) = $true) & ? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__ParkingLot) = $true) & ((s__located @ X4 @ X5 @ X1) = $true) & ((s__agent @ X4 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X4 @ s__Keeping)))) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((s__destination @ X2 @ X0 @ X3) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true))),
  inference(rectify,[status(thm)],[f67])).
thf(f67,plain,(
  ! [X3 : $i,X0 : w,X2 : $i,X1 : w] : ((((accreln1 @ s__Permission @ X1 @ X0) = $true) & ? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__ParkingLot) = $true) & ((s__located @ X4 @ X5 @ X0) = $true) & ((s__agent @ X4 @ s__Bill @ X0) = $true) & ($true = (s__instance @ X4 @ s__Keeping)))) | ((s__instance @ X3 @ s__OfficeBuilding) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((s__destination @ X2 @ X3 @ X1) != $true) | ((s__agent @ X2 @ s__Bill @ X1) != $true))),
  inference(flattening,[status(thm)],[f66])).
thf(f66,plain,(
  ! [X0 : w,X1 : w,X3 : $i,X2 : $i] : ((((accreln1 @ s__Permission @ X1 @ X0) = $true) & ? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__ParkingLot) = $true) & ((s__located @ X4 @ X5 @ X0) = $true) & ((s__agent @ X4 @ s__Bill @ X0) = $true) & ($true = (s__instance @ X4 @ s__Keeping)))) | (((s__destination @ X2 @ X3 @ X1) != $true) | ((s__agent @ X2 @ s__Bill @ X1) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((s__instance @ X3 @ s__OfficeBuilding) != $true)))),
  inference(ennf_transformation,[status(thm)],[f58])).
thf(f58,plain,(
  ! [X0 : w,X1 : w,X3 : $i,X2 : $i] : ((((s__destination @ X2 @ X3 @ X1) = $true) & ((s__agent @ X2 @ s__Bill @ X1) = $true) & ((s__instance @ X2 @ s__Driving) = $true) & ((s__instance @ X3 @ s__OfficeBuilding) = $true)) => (((accreln1 @ s__Permission @ X1 @ X0) = $true) & ? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__ParkingLot) = $true) & ((s__located @ X4 @ X5 @ X0) = $true) & ((s__agent @ X4 @ s__Bill @ X0) = $true) & ($true = (s__instance @ X4 @ s__Keeping)))))),
  inference(fool_elimination,[status(thm)],[f57])).
thf(f57,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (((s__agent @ X2 @ s__Bill @ X1) & (s__destination @ X2 @ X3 @ X1) & (s__instance @ X3 @ s__OfficeBuilding) & (s__instance @ X2 @ s__Driving)) => (? [X4 : $i,X5 : $i] : ((s__instance @ X4 @ s__Keeping) & (s__located @ X4 @ X5 @ X0) & (s__instance @ X5 @ s__ParkingLot) & (s__agent @ X4 @ s__Bill @ X0)) & (accreln1 @ s__Permission @ X1 @ X0)))),
  inference(rectify,[status(thm)],[f38])).
thf(f38,axiom,(
  ! [X1 : w,X0 : w,X6 : $i,X7 : $i] : (((s__agent @ X6 @ s__Bill @ X0) & (s__destination @ X6 @ X7 @ X0) & (s__instance @ X7 @ s__OfficeBuilding) & (s__instance @ X6 @ s__Driving)) => (? [X8 : $i,X9 : $i] : ((s__instance @ X8 @ s__Keeping) & (s__located @ X8 @ X9 @ X1) & (s__instance @ X9 @ s__ParkingLot) & (s__agent @ X8 @ s__Bill @ X1)) & (accreln1 @ s__Permission @ X0 @ X1)))),
  file('/tmp/FKSbjIJp9N/SOT_UKiqcb',permission_park_in_lot)).
thf(f152,plain,(
  spl8_2 | spl8_4),
  inference(avatar_split_clause,[status(thm)],[f90,f150,f142])).
thf(f90,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (((s__instance @ (sK0 @ X1) @ s__Keeping) = $true) | ((s__destination @ X2 @ X0 @ X3) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f73])).
thf(f148,plain,(
  spl8_3 | spl8_2),
  inference(avatar_split_clause,[status(thm)],[f91,f142,f146])).
thf(f91,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (((s__agent @ (sK0 @ X1) @ s__Bill @ X1) = $true) | ((s__destination @ X2 @ X0 @ X3) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f73])).
thf(f144,plain,(
  spl8_1 | spl8_2),
  inference(avatar_split_clause,[status(thm)],[f92,f142,f139])).
thf(f92,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (((s__located @ (sK0 @ X1) @ (sK1 @ X1) @ X1) = $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ((s__destination @ X2 @ X0 @ X3) != $true) | ((s__instance @ X2 @ s__Driving) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f73])).
% SZS output end Proof for SOT_UKiqcb
% (3746864)------------------------------
% (3746864)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (3746864)Termination reason: Refutation

% (3746864)Memory used [KB]: 5884
% (3746864)Time elapsed: 0.025 s
% (3746864)------------------------------
% (3746864)------------------------------
% (3746860)Success in time 0.038 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_UKiqcb - Vampire---5.0.1 says Theorem - CPU = 0.01 WC = 0.08 
% OUTPUT: SOT_UKiqcb - Vampire---5.0.1 says Refutation - CPU = 0.01 WC = 0.08 

