% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 2a_Permission_chained_univConj
% SZS output start Proof for 2a_Permission_chained_univConj
3. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Permission @ X0 @ X2) & (accreln1 @ s__Permission @ X0 @ X1)) => (accreln1 @ s__Permission @ X1 @ X2)) [input(axiom)]
20. ! [X0 : w,X1 : w] : ! [X2] : (((s__possesses @ s__Bill @ X2 @ X0) & (s__instance @ X2 @ s__DriversLicense)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X3,X4] : ((s__located @ X3 @ X4 @ X1) & (s__agent @ X3 @ s__Bill @ X1) & (s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking)))) [input(axiom)]
21. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__located @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__ParkingLot) & (s__instance @ X2 @ s__Parking)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X4] : ((s__destination @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Reimbursement)))) [input(axiom)]
22. ? [X0] : ((s__possesses @ s__Bill @ X0 @ cw) & (s__instance @ X0 @ s__DriversLicense)) [input(axiom)]
23. ! [X0 : w] : ((accreln1 @ s__Permission @ cw @ X0) => ? [X1] : ((s__destination @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Reimbursement))) [input(conjecture)]
24. ~! [X0 : w] : ((accreln1 @ s__Permission @ cw @ X0) => ? [X1] : ((s__destination @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Reimbursement))) [negated conjecture 23]
29. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Permission @ X0 @ X2) & (accreln1 @ s__Permission @ X0 @ X1)) => (accreln1 @ s__Permission @ X1 @ X2)) [rectify 3]
30. ! [X0 : w,X1 : w,X2 : w] : (((((accreln1 @ s__Permission @ X0 @ X2)) = $true) & (((accreln1 @ s__Permission @ X0 @ X1)) = $true)) => (((accreln1 @ s__Permission @ X1 @ X2)) = $true)) [fool elimination 29]
33. ! [X0 : w,X1 : w] : ! [X2] : (((s__possesses @ s__Bill @ X2 @ X0) & (s__instance @ X2 @ s__DriversLicense)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X3,X4] : ((s__located @ X3 @ X4 @ X1) & (s__agent @ X3 @ s__Bill @ X1) & (s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking)))) [rectify 20]
34. ! [X0 : w,X1 : w] : ! [X2] : (((((s__possesses @ s__Bill @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__DriversLicense)) = $true)) => ((((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)))) [fool elimination 33]
35. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__located @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__ParkingLot) & (s__instance @ X2 @ s__Parking)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X4] : ((s__destination @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Reimbursement)))) [rectify 21]
36. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__located @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__ParkingLot)) = $true) & (((s__instance @ X2 @ s__Parking)) = $true)) => ((((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)))) [fool elimination 35]
37. ? [X0] : ((s__possesses @ s__Bill @ X0 @ cw) & (s__instance @ X0 @ s__DriversLicense)) [rectify 22]
38. ? [X0] : ((((s__possesses @ s__Bill @ X0 @ cw)) = $true) & (((s__instance @ X0 @ s__DriversLicense)) = $true)) [fool elimination 37]
39. ~! [X0 : w] : ((accreln1 @ s__Permission @ cw @ X0) => ? [X1] : ((s__destination @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Reimbursement))) [rectify 24]
40. ~! [X0 : w] : ((((accreln1 @ s__Permission @ cw @ X0)) = $true) => ? [X1] : ((((s__destination @ X1 @ s__Bill @ X0)) = $true) & (((s__instance @ X1 @ s__Reimbursement)) = $true))) [fool elimination 39]
56. ! [X0 : w,X1 : w,X2] : (((((s__possesses @ s__Bill @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__DriversLicense)) = $true)) => ((((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)))) [flattening 34]
57. ! [X0 : w,X1 : w,X2,X3] : (((((s__located @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__ParkingLot)) = $true) & (((s__instance @ X2 @ s__Parking)) = $true)) => ((((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)))) [flattening 36]
60. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Permission @ X1 @ X2)) = $true) | ((((accreln1 @ s__Permission @ X0 @ X2)) != $true) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true))) [ennf transformation 30]
61. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Permission @ X1 @ X2)) = $true) | (((accreln1 @ s__Permission @ X0 @ X2)) != $true) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true)) [flattening 60]
62. ! [X0 : w,X1 : w,X2] : ((? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true)) | ((((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true))) [ennf transformation 56]
63. ! [X0 : w,X1 : w,X2] : (? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true)) [flattening 62]
64. ! [X0 : w,X1 : w,X2,X3] : ((? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true)) | ((((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true))) [ennf transformation 57]
65. ! [X0 : w,X1 : w,X2,X3] : (? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true)) [flattening 64]
66. ? [X0 : w] : (! [X1] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__instance @ X1 @ s__Reimbursement)) != $true)) & (((accreln1 @ s__Permission @ cw @ X0)) = $true)) [ennf transformation 40]
69. ! [X1 : w] : (? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)) => (($true = ((s__located @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__ParkingLot))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Parking))))) [skolem symbol introduction]
70. ! [X0 : w,X1 : w,X2] : ((($true = ((s__located @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__ParkingLot))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Parking)))) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true)) [skolemisation 63,69]
71. ! [X1 : w] : (? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)) => (($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Reimbursement))))) [skolem symbol introduction]
72. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Reimbursement)))) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true)) [skolemisation 65,71]
73. ? [X0] : ((((s__possesses @ s__Bill @ X0 @ cw)) = $true) & (((s__instance @ X0 @ s__DriversLicense)) = $true)) => (($true = ((s__possesses @ s__Bill @ sK4 @ cw))) & ($true = ((s__instance @ sK4 @ s__DriversLicense)))) [skolem symbol introduction]
74. ($true = ((s__possesses @ s__Bill @ sK4 @ cw))) & ($true = ((s__instance @ sK4 @ s__DriversLicense))) [skolemisation 38,73]
75. ? [X0 : w] : (! [X1] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__instance @ X1 @ s__Reimbursement)) != $true)) & (((accreln1 @ s__Permission @ cw @ X0)) = $true)) => (! [X1] : (($true != ((s__destination @ X1 @ s__Bill @ sK5))) | (((s__instance @ X1 @ s__Reimbursement)) != $true)) & ($true = ((accreln1 @ s__Permission @ cw @ sK5)))) [skolem symbol introduction]
76. ! [X1] : (($true != ((s__destination @ X1 @ s__Bill @ sK5))) | (((s__instance @ X1 @ s__Reimbursement)) != $true)) & ($true = ((accreln1 @ s__Permission @ cw @ sK5))) [skolemisation 66,75]
79. (((accreln1 @ s__Permission @ X0 @ X2)) != $true) | (((accreln1 @ s__Permission @ X1 @ X2)) = $true) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) [cnf transformation 61]
96. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__Parking))) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true) [cnf transformation 70]
97. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK2 @ X1) @ s__ParkingLot))) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true) [cnf transformation 70]
98. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true) [cnf transformation 70]
99. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__located @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true) [cnf transformation 70]
100. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Reimbursement))) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true) [cnf transformation 72]
101. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true) [cnf transformation 72]
102. ($true = ((s__instance @ sK4 @ s__DriversLicense))) [cnf transformation 74]
103. ($true = ((s__possesses @ s__Bill @ sK4 @ cw))) [cnf transformation 74]
104. ($true = ((accreln1 @ s__Permission @ cw @ sK5))) [cnf transformation 76]
105. ($true != ((s__destination @ X1 @ s__Bill @ sK5))) | (((s__instance @ X1 @ s__Reimbursement)) != $true) [cnf transformation 76]
107. ($true = X0) | ($false = X0) [fool exhaustiveness axiom]
110. ($true != $true) | ($true != ((s__instance @ X0 @ s__Reimbursement))) | ($false = ((s__destination @ X0 @ s__Bill @ sK5))) [superposition 105,107]
112. ($true != ((s__instance @ X0 @ s__Reimbursement))) | ($false = ((s__destination @ X0 @ s__Bill @ sK5))) [trivial inequality removal 110]
537. ($true != $true) | ($true = ((accreln1 @ s__Permission @ X0 @ sK5))) | (((accreln1 @ s__Permission @ cw @ X0)) != $true) [superposition 79,104]
555. (((accreln1 @ s__Permission @ cw @ X0)) != $true) | ($true = ((accreln1 @ s__Permission @ X0 @ sK5))) [trivial inequality removal 537]
583. ($true != $true) | ($true = ((s__instance @ (sK1 @ sK5) @ s__Parking))) | (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) [superposition 96,104]
601. ($true = ((s__instance @ (sK1 @ sK5) @ s__Parking))) | (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) [trivial inequality removal 583]
609. 4 <=> ! [X0] : ((((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true)) [avatar definition]
610. (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) <- (4) [avatar component clause 609]
612. 5 <=> ($true = ((s__instance @ (sK1 @ sK5) @ s__Parking))) [avatar definition]
614. ($true = ((s__instance @ (sK1 @ sK5) @ s__Parking))) <- (5) [avatar component clause 612]
615. 4 | 5 [avatar split clause 601,612,609]
638. ($true != $true) | ($true = ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) | (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) [superposition 97,104]
656. ($true = ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) | (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) [trivial inequality removal 638]
664. 6 <=> ($true = ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) [avatar definition]
666. ($true = ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) <- (6) [avatar component clause 664]
667. 4 | 6 [avatar split clause 656,664,609]
668. ($true != $true) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) <- (4) [superposition 610,103]
684. ($true != ((s__instance @ sK4 @ s__DriversLicense))) <- (4) [trivial inequality removal 668]
689. $false <- (4) [forward subsumption resolution 684,102]
690. ~4 [avatar contradiction clause 689]
691. ($true != $true) | ($true = ((s__agent @ (sK1 @ sK5) @ s__Bill @ sK5))) | (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) [superposition 98,104]
709. ($true = ((s__agent @ (sK1 @ sK5) @ s__Bill @ sK5))) | (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) [trivial inequality removal 691]
717. 7 <=> ($true = ((s__agent @ (sK1 @ sK5) @ s__Bill @ sK5))) [avatar definition]
719. ($true = ((s__agent @ (sK1 @ sK5) @ s__Bill @ sK5))) <- (7) [avatar component clause 717]
720. 4 | 7 [avatar split clause 709,717,609]
743. ($true != $true) | ($true = ((s__located @ (sK1 @ sK5) @ (sK2 @ sK5) @ sK5))) | (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) [superposition 99,104]
761. ($true = ((s__located @ (sK1 @ sK5) @ (sK2 @ sK5) @ sK5))) | (((s__possesses @ s__Bill @ X0 @ cw)) != $true) | (((s__instance @ X0 @ s__DriversLicense)) != $true) [trivial inequality removal 743]
769. 8 <=> ($true = ((s__located @ (sK1 @ sK5) @ (sK2 @ sK5) @ sK5))) [avatar definition]
771. ($true = ((s__located @ (sK1 @ sK5) @ (sK2 @ sK5) @ sK5))) <- (8) [avatar component clause 769]
772. 4 | 8 [avatar split clause 761,769,609]
797. ($true != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Reimbursement))) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true) | (((accreln1 @ s__Permission @ X0 @ X1)) = $false) [superposition 100,107]
811. (((s__located @ X2 @ X3 @ X0)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Reimbursement))) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true) | (((accreln1 @ s__Permission @ X0 @ X1)) = $false) [trivial inequality removal 797]
824. 10 <=> ($true = ((s__instance @ (sK3 @ sK5) @ s__Reimbursement))) [avatar definition]
825. ($true != ((s__instance @ (sK3 @ sK5) @ s__Reimbursement))) <- (~10) [avatar component clause 824]
826. ($true = ((s__instance @ (sK3 @ sK5) @ s__Reimbursement))) <- (10) [avatar component clause 824]
841. ($true != $true) | ($false = ((s__destination @ (sK3 @ sK5) @ s__Bill @ sK5))) <- (10) [superposition 112,826]
842. ($false = ((s__destination @ (sK3 @ sK5) @ s__Bill @ sK5))) <- (10) [trivial inequality removal 841]
854. ($true != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true) | (((accreln1 @ s__Permission @ X0 @ X1)) = $false) [superposition 101,107]
868. (((s__located @ X2 @ X3 @ X0)) != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true) | (((accreln1 @ s__Permission @ X0 @ X1)) = $false) [trivial inequality removal 854]
924. ($true != $true) | ($true = ((accreln1 @ s__Permission @ sK5 @ sK5))) [superposition 555,104]
942. ($true = ((accreln1 @ s__Permission @ sK5 @ sK5))) [trivial inequality removal 924]
1926. ($true != $true) | ($true = ((s__instance @ (sK3 @ X0) @ s__Reimbursement))) | ($true != ((s__agent @ (sK1 @ sK5) @ s__Bill @ sK5))) | ($true != ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Parking))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (8) [superposition 811,771]
1942. ($true = ((s__instance @ (sK3 @ X0) @ s__Reimbursement))) | ($true != ((s__agent @ (sK1 @ sK5) @ s__Bill @ sK5))) | ($true != ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Parking))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (8) [trivial inequality removal 1926]
1949. ($true = ((s__instance @ (sK3 @ X0) @ s__Reimbursement))) | ($true != ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Parking))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (7, 8) [forward subsumption resolution 1942,719]
1950. ($true = ((s__instance @ (sK3 @ X0) @ s__Reimbursement))) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Parking))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (6, 7, 8) [forward subsumption resolution 1949,666]
1951. ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) | ($true = ((s__instance @ (sK3 @ X0) @ s__Reimbursement))) <- (5, 6, 7, 8) [forward subsumption resolution 1950,614]
2000. ($true != $true) | ($true = ((s__destination @ (sK3 @ X0) @ s__Bill @ X0))) | ($true != ((s__agent @ (sK1 @ sK5) @ s__Bill @ sK5))) | ($true != ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Parking))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (8) [superposition 868,771]
2016. ($true = ((s__destination @ (sK3 @ X0) @ s__Bill @ X0))) | ($true != ((s__agent @ (sK1 @ sK5) @ s__Bill @ sK5))) | ($true != ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Parking))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (8) [trivial inequality removal 2000]
2023. ($true = ((s__destination @ (sK3 @ X0) @ s__Bill @ X0))) | ($true != ((s__instance @ (sK2 @ sK5) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Parking))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (7, 8) [forward subsumption resolution 2016,719]
2024. ($true = ((s__destination @ (sK3 @ X0) @ s__Bill @ X0))) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Parking))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (6, 7, 8) [forward subsumption resolution 2023,666]
2025. ($true = ((s__destination @ (sK3 @ X0) @ s__Bill @ X0))) | ($false = ((accreln1 @ s__Permission @ sK5 @ X0))) <- (5, 6, 7, 8) [forward subsumption resolution 2024,614]
2051. ($true = $false) | ($true = ((s__instance @ (sK3 @ sK5) @ s__Reimbursement))) <- (5, 6, 7, 8) [superposition 1951,942]
2089. ($true = ((s__instance @ (sK3 @ sK5) @ s__Reimbursement))) <- (5, 6, 7, 8) [trivial inequality removal 2051]
2455. ($true = $false) | ($false = ((accreln1 @ s__Permission @ sK5 @ sK5))) <- (5, 6, 7, 8, 10) [superposition 842,2025]
2460. ($false = ((accreln1 @ s__Permission @ sK5 @ sK5))) <- (5, 6, 7, 8, 10) [trivial inequality removal 2455]
2469. ($true = $false) <- (5, 6, 7, 8, 10) [superposition 2460,942]
2512. $false <- (5, 6, 7, 8, 10) [trivial inequality removal 2469]
2513. ~5 | ~6 | ~7 | ~8 | ~10 [avatar contradiction clause 2512]
2559. $false <- (5, 6, 7, 8, ~10) [forward subsumption resolution 2089,825]
2560. ~5 | ~6 | ~7 | ~8 | 10 [avatar contradiction clause 2559]
s12. 4 | 5 [sat_conversion 615]
s13. 4 | 6 [sat_conversion 667]
s14. ~4 [sat_conversion 690]
s15. 4 | 7 [sat_conversion 720]
s16. 4 | 8 [sat_conversion 772]
s32. ~10 | ~8 | ~7 | ~6 | ~5 [sat_conversion 2513]
s41. ~8 | ~7 | ~6 | ~5 | 10 [sat_conversion 2560]
s59. 8 [rat s16,s14]
s60. 7 [rat s15,s14]
s61. 6 [rat s13,s14]
s62. 5 [rat s12,s14]
s63. 10 [rat s41,s61,s60,s59,s62]
s64. # [rat s32,s61,s59,s60,s62,s63]
2587. $false [avatar sat refutation s64]
% SZS output end Proof for 2a_Permission_chained_univConj
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.295 s
% Peak memory usage: 13 MB
% Instructions burned: 327 (million)
% ------------------------------
% ------------------------------
