% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 1d_Obligation_nHops
% SZS output start Proof for 1d_Obligation_nHops
23. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ s__Bill @ X0) & (s__patient @ X2 @ X3 @ X0) & (s__instance @ X3 @ s__PayCheck) & (s__instance @ X2 @ s__Giving)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4] : ((s__origin @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Tax)))) [input(axiom)]
24. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [input(axiom)]
25. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [input(axiom)]
26. ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) & (accreln1 @ s__Obligation @ cw @ X0)) => ? [X2] : ((s__origin @ X2 @ s__Bill @ X1) & (s__instance @ X2 @ s__Tax))) [input(conjecture)]
27. ~! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) & (accreln1 @ s__Obligation @ cw @ X0)) => ? [X2] : ((s__origin @ X2 @ s__Bill @ X1) & (s__instance @ X2 @ s__Tax))) [negated conjecture 26]
30. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ s__Bill @ X0) & (s__patient @ X2 @ X3 @ X0) & (s__instance @ X3 @ s__PayCheck) & (s__instance @ X2 @ s__Giving)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4] : ((s__origin @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Tax)))) [rectify 23]
31. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ s__Bill @ X0)) = $true) & (((s__patient @ X2 @ X3 @ X0)) = $true) & (((s__instance @ X3 @ s__PayCheck)) = $true) & (((s__instance @ X2 @ s__Giving)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)))) [fool elimination 30]
32. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [rectify 24]
33. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [fool elimination 32]
34. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [rectify 25]
35. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) [fool elimination 34]
36. ~! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) & (accreln1 @ s__Obligation @ cw @ X0)) => ? [X2] : ((s__origin @ X2 @ s__Bill @ X1) & (s__instance @ X2 @ s__Tax))) [rectify 27]
37. ~! [X0 : w,X1 : w] : (((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) => ? [X2] : ((((s__origin @ X2 @ s__Bill @ X1)) = $true) & (((s__instance @ X2 @ s__Tax)) = $true))) [fool elimination 36]
59. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ s__Bill @ X0)) = $true) & (((s__patient @ X2 @ X3 @ X0)) = $true) & (((s__instance @ X3 @ s__PayCheck)) = $true) & (((s__instance @ X2 @ s__Giving)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)))) [flattening 31]
60. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [flattening 33]
61. ! [X0 : w,X1 : w,X2,X3] : ((? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true))) [ennf transformation 59]
62. ! [X0 : w,X1 : w,X2,X3] : (? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true)) [flattening 61]
63. ! [X0 : w,X1 : w,X2,X3] : ((? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true))) [ennf transformation 60]
64. ! [X0 : w,X1 : w,X2,X3] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [flattening 63]
65. ? [X0 : w,X1 : w] : (! [X2] : ((((s__origin @ X2 @ s__Bill @ X1)) != $true) | (((s__instance @ X2 @ s__Tax)) != $true)) & ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true))) [ennf transformation 37]
66. ? [X0 : w,X1 : w] : (! [X2] : ((((s__origin @ X2 @ s__Bill @ X1)) != $true) | (((s__instance @ X2 @ s__Tax)) != $true)) & (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) [flattening 65]
67. ! [X1 : w] : (? [X4] : ((((s__origin @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Tax)) = $true)) => (($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK0 @ X1) @ s__Tax))))) [skolem symbol introduction]
68. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK0 @ X1) @ s__Tax)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true)) [skolemisation 62,67]
69. ! [X1 : w] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) => (($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Giving))))) [skolem symbol introduction]
70. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Giving)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [skolemisation 64,69]
71. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) => (($true = ((s__destination @ sK3 @ sK4 @ cw))) & ($true = ((s__agent @ sK3 @ s__Bill @ cw))) & ($true = ((s__instance @ sK4 @ s__Moon))) & ($true = ((s__instance @ sK3 @ s__Flying)))) [skolem symbol introduction]
72. ($true = ((s__destination @ sK3 @ sK4 @ cw))) & ($true = ((s__agent @ sK3 @ s__Bill @ cw))) & ($true = ((s__instance @ sK4 @ s__Moon))) & ($true = ((s__instance @ sK3 @ s__Flying))) [skolemisation 35,71]
73. ? [X0 : w,X1 : w] : (! [X2] : ((((s__origin @ X2 @ s__Bill @ X1)) != $true) | (((s__instance @ X2 @ s__Tax)) != $true)) & (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) => (! [X2] : (($true != ((s__origin @ X2 @ s__Bill @ sK6))) | (((s__instance @ X2 @ s__Tax)) != $true)) & ($true = ((accreln1 @ s__Obligation @ sK5 @ sK6))) & ($true = ((accreln1 @ s__Obligation @ cw @ sK5)))) [skolem symbol introduction]
74. ! [X2] : (($true != ((s__origin @ X2 @ s__Bill @ sK6))) | (((s__instance @ X2 @ s__Tax)) != $true)) & ($true = ((accreln1 @ s__Obligation @ sK5 @ sK6))) & ($true = ((accreln1 @ s__Obligation @ cw @ sK5))) [skolemisation 66,73]
97. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK0 @ X1) @ s__Tax))) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) [cnf transformation 68]
98. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__origin @ (sK0 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ s__Bill @ X0)) != $true) | (((s__patient @ X2 @ X3 @ X0)) != $true) | (((s__instance @ X3 @ s__PayCheck)) != $true) | (((s__instance @ X2 @ s__Giving)) != $true) [cnf transformation 68]
99. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
100. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK2 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
101. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
102. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK1 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
103. ($true = ((s__instance @ sK3 @ s__Flying))) [cnf transformation 72]
104. ($true = ((s__instance @ sK4 @ s__Moon))) [cnf transformation 72]
105. ($true = ((s__agent @ sK3 @ s__Bill @ cw))) [cnf transformation 72]
106. ($true = ((s__destination @ sK3 @ sK4 @ cw))) [cnf transformation 72]
107. ($true = ((accreln1 @ s__Obligation @ cw @ sK5))) [cnf transformation 74]
108. ($true = ((accreln1 @ s__Obligation @ sK5 @ sK6))) [cnf transformation 74]
109. ($true != ((s__origin @ X2 @ s__Bill @ sK6))) | (((s__instance @ X2 @ s__Tax)) != $true) [cnf transformation 74]
778. ($true != $true) | ($true = ((s__instance @ (sK0 @ sK6) @ s__Tax))) | ($true != ((s__destination @ X0 @ s__Bill @ sK5))) | ($true != ((s__patient @ X0 @ X1 @ sK5))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) [superposition 97,108]
800. ($true = ((s__instance @ (sK0 @ sK6) @ s__Tax))) | ($true != ((s__destination @ X0 @ s__Bill @ sK5))) | ($true != ((s__patient @ X0 @ X1 @ sK5))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) [trivial inequality removal 778]
862. ($true != $true) | ($true = ((s__instance @ (sK1 @ sK5) @ s__Giving))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 99,107]
886. ($true = ((s__instance @ (sK1 @ sK5) @ s__Giving))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 862]
905. 20 <=> ! [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true)) [avatar definition]
906. (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) <- (20) [avatar component clause 905]
908. 21 <=> ($true = ((s__instance @ (sK1 @ sK5) @ s__Giving))) [avatar definition]
910. ($true = ((s__instance @ (sK1 @ sK5) @ s__Giving))) <- (21) [avatar component clause 908]
911. 20 | 21 [avatar split clause 886,908,905]
942. ($true != $true) | ($true = ((s__instance @ (sK2 @ sK5) @ s__PayCheck))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 100,107]
966. ($true = ((s__instance @ (sK2 @ sK5) @ s__PayCheck))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 942]
1128. ($true != $true) | ($true = ((s__origin @ (sK0 @ sK6) @ s__Bill @ sK6))) | ($true != ((s__destination @ X0 @ s__Bill @ sK5))) | ($true != ((s__patient @ X0 @ X1 @ sK5))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) [superposition 98,108]
1154. ($true = ((s__origin @ (sK0 @ sK6) @ s__Bill @ sK6))) | ($true != ((s__destination @ X0 @ s__Bill @ sK5))) | ($true != ((s__patient @ X0 @ X1 @ sK5))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) [trivial inequality removal 1128]
1190. ($true != $true) | ($true = ((s__destination @ (sK1 @ sK5) @ s__Bill @ sK5))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 102,107]
1218. ($true = ((s__destination @ (sK1 @ sK5) @ s__Bill @ sK5))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 1190]
1358. ($true != $true) | ($true = ((s__patient @ (sK1 @ sK5) @ (sK2 @ sK5) @ sK5))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 101,107]
1389. ($true = ((s__patient @ (sK1 @ sK5) @ (sK2 @ sK5) @ sK5))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 1358]
1611. ($true != $true) | ($true != ((s__instance @ sK3 @ s__Flying))) | ($true != ((s__instance @ sK4 @ s__Moon))) | ($true != ((s__agent @ sK3 @ s__Bill @ cw))) <- (20) [superposition 906,106]
1640. ($true != ((s__instance @ sK3 @ s__Flying))) | ($true != ((s__instance @ sK4 @ s__Moon))) | ($true != ((s__agent @ sK3 @ s__Bill @ cw))) <- (20) [trivial inequality removal 1611]
1645. ($true != ((s__instance @ sK4 @ s__Moon))) | ($true != ((s__agent @ sK3 @ s__Bill @ cw))) <- (20) [forward subsumption resolution 1640,103]
1646. ($true != ((s__agent @ sK3 @ s__Bill @ cw))) <- (20) [forward subsumption resolution 1645,104]
1647. $false <- (20) [forward subsumption resolution 1646,105]
1648. ~20 [avatar contradiction clause 1647]
1650. 22 <=> ($true = ((s__instance @ (sK2 @ sK5) @ s__PayCheck))) [avatar definition]
1652. ($true = ((s__instance @ (sK2 @ sK5) @ s__PayCheck))) <- (22) [avatar component clause 1650]
1653. 20 | 22 [avatar split clause 966,1650,905]
1655. 23 <=> ($true = ((s__destination @ (sK1 @ sK5) @ s__Bill @ sK5))) [avatar definition]
1657. ($true = ((s__destination @ (sK1 @ sK5) @ s__Bill @ sK5))) <- (23) [avatar component clause 1655]
1658. 20 | 23 [avatar split clause 1218,1655,905]
1660. 24 <=> ($true = ((s__patient @ (sK1 @ sK5) @ (sK2 @ sK5) @ sK5))) [avatar definition]
1662. ($true = ((s__patient @ (sK1 @ sK5) @ (sK2 @ sK5) @ sK5))) <- (24) [avatar component clause 1660]
1663. 20 | 24 [avatar split clause 1389,1660,905]
3520. 62 <=> ! [X0,X1] : (($true != ((s__destination @ X0 @ s__Bill @ sK5))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__patient @ X0 @ X1 @ sK5)))) [avatar definition]
3521. ($true != ((s__destination @ X0 @ s__Bill @ sK5))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__patient @ X0 @ X1 @ sK5))) <- (62) [avatar component clause 3520]
3523. 63 <=> ($true = ((s__instance @ (sK0 @ sK6) @ s__Tax))) [avatar definition]
3525. ($true = ((s__instance @ (sK0 @ sK6) @ s__Tax))) <- (63) [avatar component clause 3523]
3526. 62 | 63 [avatar split clause 800,3523,3520]
3612. 80 <=> ($true = ((s__origin @ (sK0 @ sK6) @ s__Bill @ sK6))) [avatar definition]
3614. ($true = ((s__origin @ (sK0 @ sK6) @ s__Bill @ sK6))) <- (80) [avatar component clause 3612]
3615. 62 | 80 [avatar split clause 1154,3612,3520]
4779. ($true != $true) | ($true != ((s__instance @ (sK1 @ sK5) @ s__Giving))) | ($true != ((s__instance @ X0 @ s__PayCheck))) | ($true != ((s__patient @ (sK1 @ sK5) @ X0 @ sK5))) <- (23, 62) [superposition 3521,1657]
4808. ($true != ((s__instance @ (sK1 @ sK5) @ s__Giving))) | ($true != ((s__instance @ X0 @ s__PayCheck))) | ($true != ((s__patient @ (sK1 @ sK5) @ X0 @ sK5))) <- (23, 62) [trivial inequality removal 4779]
4817. ($true != ((s__patient @ (sK1 @ sK5) @ X0 @ sK5))) | ($true != ((s__instance @ X0 @ s__PayCheck))) <- (21, 23, 62) [forward subsumption resolution 4808,910]
7193. ($true != $true) | ($true != ((s__instance @ (sK2 @ sK5) @ s__PayCheck))) <- (21, 23, 24, 62) [superposition 4817,1662]
7240. ($true != ((s__instance @ (sK2 @ sK5) @ s__PayCheck))) <- (21, 23, 24, 62) [trivial inequality removal 7193]
7258. $false <- (21, 22, 23, 24, 62) [forward subsumption resolution 7240,1652]
7259. ~21 | ~22 | ~23 | ~24 | ~62 [avatar contradiction clause 7258]
7402. ($true != $true) | ($true != ((s__instance @ (sK0 @ sK6) @ s__Tax))) <- (80) [superposition 109,3614]
7411. ($true != ((s__instance @ (sK0 @ sK6) @ s__Tax))) <- (80) [trivial inequality removal 7402]
7430. $false <- (63, 80) [forward subsumption resolution 7411,3525]
7431. ~63 | ~80 [avatar contradiction clause 7430]
s35. 20 | 21 [sat_conversion 911]
s36. ~20 [sat_conversion 1648]
s37. 20 | 22 [sat_conversion 1653]
s38. 20 | 23 [sat_conversion 1658]
s39. 20 | 24 [sat_conversion 1663]
s101. 62 | 63 [sat_conversion 3526]
s126. 62 | 80 [sat_conversion 3615]
s313. ~62 | ~24 | ~23 | ~22 | ~21 [sat_conversion 7259]
s323. ~80 | ~63 [sat_conversion 7431]
s325. 24 [rat s39,s36]
s326. 23 [rat s38,s36]
s327. 22 [rat s37,s36]
s328. 21 [rat s35,s36]
s329. ~62 [rat s313,s325,s327,s326,s328]
s331. 80 [rat s126,s329]
s332. 63 [rat s101,s329]
s339. # [rat s323,s332,s331]
7462. $false [avatar sat refutation s339]
% SZS output end Proof for 1d_Obligation_nHops
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.410 s
% Peak memory usage: 15 MB
% Instructions burned: 2085 (million)
% ------------------------------
% ------------------------------
