 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/dbsAtWB5gp/SOT_U1H4ec
% (2781800)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_U1H4ec for (1ds/275Mi)
% (2781795)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_U1H4ec for (1ds/183Mi)
% (2781799)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_U1H4ec for (1ds/2Mi)
% (2781796)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_U1H4ec for (1ds/4Mi)
% (2781797)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_U1H4ec for (1ds/27Mi)
% (2781801)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_U1H4ec for (1ds/18Mi)
% (2781798)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_U1H4ec for (1ds/2Mi)
% (2781800)Refutation not found, incomplete strategy
% (2781800)------------------------------
% (2781800)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2781800)Termination reason: Refutation not found, incomplete strategy


% (2781800)Memory used [KB]: 5628
% (2781800)Time elapsed: 0.004 s
% (2781800)------------------------------
% (2781800)------------------------------
% (2781799)Refutation not found, incomplete strategy
% (2781799)------------------------------
% (2781799)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2781799)Termination reason: Refutation not found, incomplete strategy


% (2781799)Memory used [KB]: 5628
% (2781799)Time elapsed: 0.006 s
% (2781799)------------------------------
% (2781799)------------------------------
% (2781795)First to succeed.
% (2781795)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_U1H4ec
% SZS output start Proof for SOT_U1H4ec
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m > w > w > $o).
thf(func_def_5, type, accreln2: m > $i > w > w > $o).
thf(func_def_6, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_8, type, s__instance: $i > $i > $o).
thf(func_def_9, type, s__agent: $i > $i > w > $o).
thf(func_def_10, type, s__destination: $i > $i > w > $o).
thf(func_def_11, type, s__located: $i > $i > w > $o).
thf(func_def_12, type, s__origin: $i > $i > w > $o).
thf(func_def_22, type, sK1: w > $i > $i).
thf(func_def_23, type, sK2: w > $i > $i).
thf(func_def_24, type, sK3: w > w).
thf(func_def_25, type, sK4: w > w).
thf(f156,plain,(
  $false),
  inference(trivial_inequality_removal,[status(thm)],[f154])).
thf(f154,plain,(
  ($true != $true)),
  inference(superposition,[status(thm)],[f153,f99])).
thf(f99,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Obligation @ X0 @ (sK4 @ X0)))) )),
  inference(cnf_transformation,[status(thm)],[f70])).
thf(f70,plain,(
  ! [X0 : w] : ($true = (accreln1 @ s__Obligation @ X0 @ (sK4 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f37,f69])).
thf(f69,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ($true = (accreln1 @ s__Obligation @ X0 @ (sK4 @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f37,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f36])).
thf(f36,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  file('/tmp/dbsAtWB5gp/SOT_U1H4ec',serial_obligation)).
thf(f153,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f152])).
thf(f152,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true != (accreln1 @ s__Obligation @ cw @ X0))) )),
  inference(superposition,[status(thm)],[f151,f95])).
thf(f95,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Permission @ X1 @ X0) = $true) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f67])).
thf(f67,plain,(
  ! [X0 : w,X1 : w] : (((accreln1 @ s__Permission @ X1 @ X0) = $true) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true))),
  inference(rectify,[status(thm)],[f46])).
thf(f46,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Permission @ X0 @ X1) = $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true))),
  inference(ennf_transformation,[status(thm)],[f33])).
thf(f33,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ((accreln1 @ s__Permission @ X0 @ X1) = $true))),
  inference(fool_elimination,[status(thm)],[f32])).
thf(f32,plain,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  inference(rectify,[status(thm)],[f8])).
thf(f8,axiom,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  file('/tmp/dbsAtWB5gp/SOT_U1H4ec',bridge1_O_P)).
thf(f151,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Permission @ cw @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f150,f76])).
thf(f76,plain,(
  ($true = (s__instance @ sK0 @ s__OfficeBuilding))),
  inference(cnf_transformation,[status(thm)],[f57])).
thf(f57,plain,(
  ($true = (s__located @ s__Bill @ sK0 @ cw)) & ($true = (s__instance @ sK0 @ s__OfficeBuilding))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f27,f56])).
thf(f56,plain,(
  ? [X0 : $i] : (($true = (s__located @ s__Bill @ X0 @ cw)) & ($true = (s__instance @ X0 @ s__OfficeBuilding))) => (($true = (s__located @ s__Bill @ sK0 @ cw)) & ($true = (s__instance @ sK0 @ s__OfficeBuilding)))),
  introduced(definition,[],[choice_axiom])).
thf(f27,plain,(
  ? [X0 : $i] : (($true = (s__located @ s__Bill @ X0 @ cw)) & ($true = (s__instance @ X0 @ s__OfficeBuilding)))),
  inference(fool_elimination,[status(thm)],[f26])).
thf(f26,plain,(
  ? [X0 : $i] : ((s__located @ s__Bill @ X0 @ cw) & (s__instance @ X0 @ s__OfficeBuilding))),
  inference(rectify,[status(thm)],[f22])).
thf(f22,axiom,(
  ? [X3 : $i] : ((s__located @ s__Bill @ X3 @ cw) & (s__instance @ X3 @ s__OfficeBuilding))),
  file('/tmp/dbsAtWB5gp/SOT_U1H4ec',optA_bill_at_office)).
thf(f150,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ sK0 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Permission @ cw @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f149])).
thf(f149,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true != (s__instance @ sK0 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Permission @ cw @ X0))) )),
  inference(duplicate_literal_removal,[status(thm)],[f148])).
thf(f148,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true != (s__instance @ sK0 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Permission @ cw @ X0)) | ($true != (accreln1 @ s__Permission @ cw @ X0))) )),
  inference(superposition,[status(thm)],[f143,f85])).
thf(f85,plain,(
  ( ! [X0 : $i,X1 : w] : (($true = (s__instance @ (sK2 @ X1 @ X0) @ s__Smoking)) | ($true != (s__instance @ X0 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Permission @ cw @ X1))) )),
  inference(cnf_transformation,[status(thm)],[f63])).
thf(f63,plain,(
  ! [X0 : $i,X1 : w] : (($true != (s__instance @ X0 @ s__OfficeBuilding)) | (($true = (s__instance @ (sK2 @ X1 @ X0) @ s__Smoking)) & ($true = (s__agent @ (sK2 @ X1 @ X0) @ s__Bill @ X1)) & ($true = (s__located @ (sK2 @ X1 @ X0) @ X0 @ X1))) | ($true != (accreln1 @ s__Permission @ cw @ X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f61,f62])).
thf(f62,plain,(
  ! [X0 : $i,X1 : w] : (? [X2 : $i] : (($true = (s__instance @ X2 @ s__Smoking)) & ((s__agent @ X2 @ s__Bill @ X1) = $true) & ($true = (s__located @ X2 @ X0 @ X1))) => (($true = (s__instance @ (sK2 @ X1 @ X0) @ s__Smoking)) & ($true = (s__agent @ (sK2 @ X1 @ X0) @ s__Bill @ X1)) & ($true = (s__located @ (sK2 @ X1 @ X0) @ X0 @ X1))))),
  introduced(definition,[],[choice_axiom])).
thf(f61,plain,(
  ! [X0 : $i,X1 : w] : (($true != (s__instance @ X0 @ s__OfficeBuilding)) | ? [X2 : $i] : (($true = (s__instance @ X2 @ s__Smoking)) & ((s__agent @ X2 @ s__Bill @ X1) = $true) & ($true = (s__located @ X2 @ X0 @ X1))) | ($true != (accreln1 @ s__Permission @ cw @ X1)))),
  inference(rectify,[status(thm)],[f55])).
thf(f55,plain,(
  ! [X1 : $i,X0 : w] : (($true != (s__instance @ X1 @ s__OfficeBuilding)) | ? [X2 : $i] : (($true = (s__instance @ X2 @ s__Smoking)) & ($true = (s__agent @ X2 @ s__Bill @ X0)) & ((s__located @ X2 @ X1 @ X0) = $true)) | ($true != (accreln1 @ s__Permission @ cw @ X0)))),
  inference(ennf_transformation,[status(thm)],[f45])).
thf(f45,plain,(
  ~? [X0 : w,X1 : $i] : (($true = (s__instance @ X1 @ s__OfficeBuilding)) & ~? [X2 : $i] : (($true = (s__instance @ X2 @ s__Smoking)) & ($true = (s__agent @ X2 @ s__Bill @ X0)) & ((s__located @ X2 @ X1 @ X0) = $true)) & ($true = (accreln1 @ s__Permission @ cw @ X0)))),
  inference(fool_elimination,[status(thm)],[f44])).
thf(f44,plain,(
  ~? [X0 : w,X1 : $i] : ((accreln1 @ s__Permission @ cw @ X0) & (s__instance @ X1 @ s__OfficeBuilding) & ~? [X2 : $i] : ((s__agent @ X2 @ s__Bill @ X0) & (s__located @ X2 @ X1 @ X0) & (s__instance @ X2 @ s__Smoking)))),
  inference(rectify,[status(thm)],[f24])).
thf(f24,negated_conjecture,(
  ~? [X5 : w,X3 : $i] : ((accreln1 @ s__Permission @ cw @ X5) & (s__instance @ X3 @ s__OfficeBuilding) & ~? [X4 : $i] : ((s__agent @ X4 @ s__Bill @ X5) & (s__located @ X4 @ X3 @ X5) & (s__instance @ X4 @ s__Smoking)))),
  inference(negated_conjecture,[status(cth)],[f23])).
thf(f23,conjecture,(
  ? [X5 : w,X3 : $i] : ((accreln1 @ s__Permission @ cw @ X5) & (s__instance @ X3 @ s__OfficeBuilding) & ~? [X4 : $i] : ((s__agent @ X4 @ s__Bill @ X5) & (s__located @ X4 @ X3 @ X5) & (s__instance @ X4 @ s__Smoking)))),
  file('/tmp/dbsAtWB5gp/SOT_U1H4ec',conjecture_permitted_not_smoke)).
thf(f143,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK2 @ X0 @ sK0) @ s__Smoking)) | ($true != (accreln1 @ s__Permission @ cw @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f142,f76])).
thf(f142,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Permission @ cw @ X0)) | ($true != (s__instance @ (sK2 @ X0 @ sK0) @ s__Smoking)) | ($true != (s__instance @ sK0 @ s__OfficeBuilding))) )),
  inference(trivial_inequality_removal,[status(thm)],[f141])).
thf(f141,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ sK0 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Permission @ cw @ X0)) | ($true != (s__instance @ (sK2 @ X0 @ sK0) @ s__Smoking)) | ($true != $true)) )),
  inference(duplicate_literal_removal,[status(thm)],[f140])).
thf(f140,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Permission @ cw @ X0)) | ($true != $true) | ($true != (s__instance @ sK0 @ s__OfficeBuilding)) | ($true != (s__instance @ (sK2 @ X0 @ sK0) @ s__Smoking)) | ($true != (accreln1 @ s__Permission @ cw @ X0))) )),
  inference(superposition,[status(thm)],[f139,f84])).
thf(f84,plain,(
  ( ! [X0 : $i,X1 : w] : (($true = (s__agent @ (sK2 @ X1 @ X0) @ s__Bill @ X1)) | ($true != (s__instance @ X0 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Permission @ cw @ X1))) )),
  inference(cnf_transformation,[status(thm)],[f63])).
thf(f139,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK2 @ X0 @ sK0) @ s__Bill @ X0)) | ($true != (accreln1 @ s__Permission @ cw @ X0)) | ($true != (s__instance @ (sK2 @ X0 @ sK0) @ s__Smoking))) )),
  inference(subsumption_resolution,[status(thm)],[f126,f76])).
thf(f126,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK2 @ X0 @ sK0) @ s__Smoking)) | ($true != (s__agent @ (sK2 @ X0 @ sK0) @ s__Bill @ X0)) | ($true != (accreln1 @ s__Permission @ cw @ X0)) | ($true != (s__instance @ sK0 @ s__OfficeBuilding))) )),
  inference(trivial_inequality_removal,[status(thm)],[f124])).
thf(f124,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ sK0 @ s__OfficeBuilding)) | ($true != (s__instance @ (sK2 @ X0 @ sK0) @ s__Smoking)) | ($true != $true) | ($true != (accreln1 @ s__Permission @ cw @ X0)) | ($true != (s__agent @ (sK2 @ X0 @ sK0) @ s__Bill @ X0))) )),
  inference(superposition,[status(thm)],[f122,f83])).
thf(f83,plain,(
  ( ! [X0 : $i,X1 : w] : (($true = (s__located @ (sK2 @ X1 @ X0) @ X0 @ X1)) | ($true != (accreln1 @ s__Permission @ cw @ X1)) | ($true != (s__instance @ X0 @ s__OfficeBuilding))) )),
  inference(cnf_transformation,[status(thm)],[f63])).
thf(f122,plain,(
  ( ! [X0 : $i,X1 : w] : (($true != (s__located @ X0 @ sK0 @ X1)) | ($true != (s__instance @ X0 @ s__Smoking)) | ($true != (s__agent @ X0 @ s__Bill @ X1))) )),
  inference(subsumption_resolution,[status(thm)],[f121,f76])).
thf(f121,plain,(
  ( ! [X0 : $i,X1 : w] : (($true != (s__located @ X0 @ sK0 @ X1)) | ($true != (s__instance @ sK0 @ s__OfficeBuilding)) | ($true != (s__agent @ X0 @ s__Bill @ X1)) | ($true != (s__instance @ X0 @ s__Smoking))) )),
  inference(trivial_inequality_removal,[status(thm)],[f120])).
thf(f120,plain,(
  ( ! [X0 : $i,X1 : w] : (($true != (s__agent @ X0 @ s__Bill @ X1)) | ($true != (s__instance @ X0 @ s__Smoking)) | ($true != (s__located @ X0 @ sK0 @ X1)) | ($true != $true) | ($true != (s__instance @ sK0 @ s__OfficeBuilding))) )),
  inference(superposition,[status(thm)],[f119,f77])).
thf(f77,plain,(
  ($true = (s__located @ s__Bill @ sK0 @ cw))),
  inference(cnf_transformation,[status(thm)],[f57])).
thf(f119,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != (s__located @ s__Bill @ X0 @ cw)) | ($true != (s__agent @ X1 @ s__Bill @ X2)) | ($true != (s__instance @ X1 @ s__Smoking)) | ($true != (s__located @ X1 @ X0 @ X2)) | ($true != (s__instance @ X0 @ s__OfficeBuilding))) )),
  inference(trivial_inequality_removal,[status(thm)],[f116])).
thf(f116,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != (s__instance @ X0 @ s__OfficeBuilding)) | ($true != (s__located @ s__Bill @ X0 @ cw)) | ($true != (s__located @ X1 @ X0 @ X2)) | ($true != $true) | ($true != (s__agent @ X1 @ s__Bill @ X2)) | ($true != (s__instance @ X1 @ s__Smoking))) )),
  inference(superposition,[status(thm)],[f115,f99])).
thf(f115,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Obligation @ cw @ (sK4 @ X0))) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__located @ X2 @ X1 @ X3)) | ($true != (s__instance @ X2 @ s__Smoking)) | ($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (s__located @ s__Bill @ X1 @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f114])).
thf(f114,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__instance @ X2 @ s__Smoking)) | ($true != (s__located @ X2 @ X1 @ X3)) | ($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != $true) | ($true != (s__located @ s__Bill @ X1 @ X0)) | ($true != (accreln1 @ s__Obligation @ cw @ (sK4 @ X0))) | ($true != (s__instance @ X1 @ s__OfficeBuilding))) )),
  inference(superposition,[status(thm)],[f113,f95])).
thf(f113,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Permission @ cw @ (sK4 @ X0))) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__located @ s__Bill @ X1 @ X0)) | ($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (s__located @ X2 @ X1 @ X3)) | ($true != (s__instance @ X2 @ s__Smoking))) )),
  inference(subsumption_resolution,[status(thm)],[f112,f84])).
thf(f112,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (s__located @ s__Bill @ X1 @ X0)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__instance @ X2 @ s__Smoking)) | ($true != (s__located @ X2 @ X1 @ X3)) | ($true != (accreln1 @ s__Permission @ cw @ (sK4 @ X0))) | ($true != (s__agent @ (sK2 @ (sK4 @ X0) @ X1) @ s__Bill @ (sK4 @ X0)))) )),
  inference(subsumption_resolution,[status(thm)],[f107,f85])).
thf(f107,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__instance @ (sK2 @ (sK4 @ X0) @ X1) @ s__Smoking)) | ($true != (accreln1 @ s__Permission @ cw @ (sK4 @ X0))) | ($true != (s__agent @ (sK2 @ (sK4 @ X0) @ X1) @ s__Bill @ (sK4 @ X0))) | ($true != (s__instance @ X2 @ s__Smoking)) | ($true != (s__located @ X2 @ X1 @ X3)) | ($true != (s__located @ s__Bill @ X1 @ X0)) | ($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (s__instance @ X1 @ s__OfficeBuilding))) )),
  inference(trivial_inequality_removal,[status(thm)],[f106])).
thf(f106,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__located @ s__Bill @ X1 @ X0)) | ($true != (s__instance @ (sK2 @ (sK4 @ X0) @ X1) @ s__Smoking)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__agent @ (sK2 @ (sK4 @ X0) @ X1) @ s__Bill @ (sK4 @ X0))) | ($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (s__located @ X2 @ X1 @ X3)) | ($true != (s__instance @ X2 @ s__Smoking)) | ($true != $true) | ($true != (accreln1 @ s__Permission @ cw @ (sK4 @ X0)))) )),
  inference(duplicate_literal_removal,[status(thm)],[f104])).
thf(f104,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__located @ s__Bill @ X1 @ X0)) | ($true != (s__agent @ (sK2 @ (sK4 @ X0) @ X1) @ s__Bill @ (sK4 @ X0))) | ($true != (s__located @ X2 @ X1 @ X3)) | ($true != (s__instance @ X2 @ s__Smoking)) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Permission @ cw @ (sK4 @ X0))) | ($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != $true) | ($true != (s__instance @ X1 @ s__OfficeBuilding)) | ($true != (s__instance @ (sK2 @ (sK4 @ X0) @ X1) @ s__Smoking))) )),
  inference(superposition,[status(thm)],[f102,f83])).
thf(f102,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : $i,X4 : w] : (($true != (s__located @ X1 @ X2 @ (sK4 @ X0))) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__located @ s__Bill @ X2 @ X0)) | ($true != (s__located @ X3 @ X2 @ X4)) | ($true != (s__agent @ X3 @ s__Bill @ X4)) | ($true != (s__instance @ X3 @ s__Smoking)) | ($true != (s__instance @ X1 @ s__Smoking)) | ($true != (s__agent @ X1 @ s__Bill @ (sK4 @ X0)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f100])).
thf(f100,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : $i,X4 : w] : (($true != $true) | ($true != (s__agent @ X3 @ s__Bill @ X4)) | ($true != (s__located @ X1 @ X2 @ (sK4 @ X0))) | ($true != (s__agent @ X1 @ s__Bill @ (sK4 @ X0))) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__located @ X3 @ X2 @ X4)) | ($true != (s__instance @ X1 @ s__Smoking)) | ($true != (s__instance @ X3 @ s__Smoking)) | ($true != (s__located @ s__Bill @ X2 @ X0))) )),
  inference(superposition,[status(thm)],[f88,f99])).
thf(f88,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : w,X4 : $i,X5 : $i] : (($true != (accreln1 @ s__Obligation @ X3 @ X1)) | ($true != (s__agent @ X5 @ s__Bill @ X1)) | ($true != (s__located @ s__Bill @ X2 @ X3)) | ($true != (s__located @ X5 @ X2 @ X1)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__instance @ X5 @ s__Smoking)) | ($true != (s__located @ X4 @ X2 @ X0)) | ($true != (s__agent @ X4 @ s__Bill @ X0)) | ((s__instance @ X4 @ s__Smoking) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f64])).
thf(f64,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : w] : ((! [X4 : $i] : (($true != (s__located @ X4 @ X2 @ X0)) | ((s__instance @ X4 @ s__Smoking) != $true) | ($true != (s__agent @ X4 @ s__Bill @ X0))) & ($true = (accreln1 @ s__Prohibition @ X3 @ X0))) | ($true != (accreln1 @ s__Obligation @ X3 @ X1)) | ! [X5 : $i] : (($true != (s__instance @ X5 @ s__Smoking)) | ($true != (s__located @ X5 @ X2 @ X1)) | ($true != (s__agent @ X5 @ s__Bill @ X1))) | ($true != (s__located @ s__Bill @ X2 @ X3)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)))),
  inference(rectify,[status(thm)],[f50])).
thf(f50,plain,(
  ! [X3 : w,X1 : w,X2 : $i,X0 : w] : ((! [X4 : $i] : (((s__located @ X4 @ X2 @ X3) != $true) | ((s__instance @ X4 @ s__Smoking) != $true) | ($true != (s__agent @ X4 @ s__Bill @ X3))) & ((accreln1 @ s__Prohibition @ X0 @ X3) = $true)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ! [X5 : $i] : (($true != (s__instance @ X5 @ s__Smoking)) | ($true != (s__located @ X5 @ X2 @ X1)) | ($true != (s__agent @ X5 @ s__Bill @ X1))) | ($true != (s__located @ s__Bill @ X2 @ X0)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)))),
  inference(flattening,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X0 : w,X3 : w,X2 : $i,X1 : w] : ((! [X5 : $i] : (($true != (s__instance @ X5 @ s__Smoking)) | ($true != (s__located @ X5 @ X2 @ X1)) | ($true != (s__agent @ X5 @ s__Bill @ X1))) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true)) | (($true != (s__located @ s__Bill @ X2 @ X0)) | (! [X4 : $i] : (((s__located @ X4 @ X2 @ X3) != $true) | ((s__instance @ X4 @ s__Smoking) != $true) | ($true != (s__agent @ X4 @ s__Bill @ X3))) & ((accreln1 @ s__Prohibition @ X0 @ X3) = $true)) | ($true != (s__instance @ X2 @ s__OfficeBuilding))))),
  inference(ennf_transformation,[status(thm)],[f39])).
thf(f39,plain,(
  ! [X0 : w,X3 : w,X2 : $i,X1 : w] : ((($true = (s__located @ s__Bill @ X2 @ X0)) & (((accreln1 @ s__Prohibition @ X0 @ X3) = $true) => ? [X4 : $i] : (((s__instance @ X4 @ s__Smoking) = $true) & ((s__located @ X4 @ X2 @ X3) = $true) & ($true = (s__agent @ X4 @ s__Bill @ X3)))) & ($true = (s__instance @ X2 @ s__OfficeBuilding))) => (((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ~? [X5 : $i] : (($true = (s__agent @ X5 @ s__Bill @ X1)) & ($true = (s__located @ X5 @ X2 @ X1)) & ($true = (s__instance @ X5 @ s__Smoking)))))),
  inference(fool_elimination,[status(thm)],[f38])).
thf(f38,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : w] : (((s__instance @ X2 @ s__OfficeBuilding) & (s__located @ s__Bill @ X2 @ X0) & ((accreln1 @ s__Prohibition @ X0 @ X3) => ? [X4 : $i] : ((s__located @ X4 @ X2 @ X3) & (s__instance @ X4 @ s__Smoking) & (s__agent @ X4 @ s__Bill @ X3)))) => ((accreln1 @ s__Obligation @ X0 @ X1) => ~? [X5 : $i] : ((s__located @ X5 @ X2 @ X1) & (s__instance @ X5 @ s__Smoking) & (s__agent @ X5 @ s__Bill @ X1))))),
  inference(rectify,[status(thm)],[f20])).
thf(f20,axiom,(
  ! [X0 : w,X2 : w,X3 : $i,X1 : w] : (((s__instance @ X3 @ s__OfficeBuilding) & (s__located @ s__Bill @ X3 @ X0) & ((accreln1 @ s__Prohibition @ X0 @ X1) => ? [X4 : $i] : ((s__located @ X4 @ X3 @ X1) & (s__instance @ X4 @ s__Smoking) & (s__agent @ X4 @ s__Bill @ X1)))) => ((accreln1 @ s__Obligation @ X0 @ X2) => ~? [X4 : $i] : ((s__located @ X4 @ X3 @ X2) & (s__instance @ X4 @ s__Smoking) & (s__agent @ X4 @ s__Bill @ X2))))),
  file('/tmp/dbsAtWB5gp/SOT_U1H4ec',bridge4_F_O_hardcoded)).
% SZS output end Proof for SOT_U1H4ec
% (2781795)------------------------------
% (2781795)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2781795)Termination reason: Refutation

% (2781795)Memory used [KB]: 5756
% (2781795)Time elapsed: 0.015 s
% (2781795)------------------------------
% (2781795)------------------------------
% (2781791)Success in time 0.027 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_U1H4ec - Vampire---5.0.1 says Theorem - CPU = 0.00 WC = 0.07 
% OUTPUT: SOT_U1H4ec - Vampire---5.0.1 says Refutation - CPU = 0.00 WC = 0.07 

