% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 1g_Obligation_UO1_existConj
% SZS output start Proof for 1g_Obligation_UO1_existConj
1. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1) [input(axiom)]
2. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X1 @ X2) & (accreln1 @ s__Obligation @ X0 @ X1)) => (accreln1 @ s__Obligation @ X0 @ X2)) [input(axiom)]
19. ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X2,X3] : ((s__destination @ X2 @ X3 @ X1) & (s__agent @ X2 @ s__Bill @ X1) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying))) [input(axiom)]
20. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [input(axiom)]
22. ? [X0 : w] : (? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving)) & (accreln1 @ s__Obligation @ cw @ X0)) [input(conjecture)]
23. ~? [X0 : w] : (? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving)) & (accreln1 @ s__Obligation @ cw @ X0)) [negated conjecture 22]
24. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1) [rectify 1]
25. ! [X0 : w] : ? [X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) [fool elimination 24]
26. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X1 @ X2) & (accreln1 @ s__Obligation @ X0 @ X1)) => (accreln1 @ s__Obligation @ X0 @ X2)) [rectify 2]
27. ! [X0 : w,X1 : w,X2 : w] : (((((accreln1 @ s__Obligation @ X1 @ X2)) = $true) & (((accreln1 @ s__Obligation @ X0 @ X1)) = $true)) => (((accreln1 @ s__Obligation @ X0 @ X2)) = $true)) [fool elimination 26]
30. ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X2,X3] : ((s__destination @ X2 @ X3 @ X1) & (s__agent @ X2 @ s__Bill @ X1) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying))) [rectify 19]
31. ! [X0 : w,X1 : w] : ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X2,X3] : ((((s__destination @ X2 @ X3 @ X1)) = $true) & (((s__agent @ X2 @ s__Bill @ X1)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true))) [fool elimination 30]
32. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [rectify 20]
33. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [fool elimination 32]
36. ~? [X0 : w] : (? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving)) & (accreln1 @ s__Obligation @ cw @ X0)) [rectify 23]
37. ~? [X0 : w] : (? [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) = $true) & (((s__patient @ X1 @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__PayCheck)) = $true) & (((s__instance @ X1 @ s__Giving)) = $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) [fool elimination 36]
53. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [flattening 33]
54. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Obligation @ X0 @ X2)) = $true) | ((((accreln1 @ s__Obligation @ X1 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true))) [ennf transformation 27]
55. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Obligation @ X0 @ X2)) = $true) | (((accreln1 @ s__Obligation @ X1 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) [flattening 54]
56. ! [X0 : w,X1 : w] : (? [X2,X3] : ((((s__destination @ X2 @ X3 @ X1)) = $true) & (((s__agent @ X2 @ s__Bill @ X1)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) [ennf transformation 31]
57. ! [X0 : w,X1 : w,X2,X3] : ((? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true))) [ennf transformation 53]
58. ! [X0 : w,X1 : w,X2,X3] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [flattening 57]
59. ! [X0 : w] : (! [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) | (((accreln1 @ s__Obligation @ cw @ X0)) != $true)) [ennf transformation 37]
60. ! [X0 : w] : (? [X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0))))) [skolem symbol introduction]
61. ! [X0 : w] : ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)))) [skolemisation 25,60]
62. ! [X1 : w] : (? [X2,X3] : ((((s__destination @ X2 @ X3 @ X1)) = $true) & (((s__agent @ X2 @ s__Bill @ X1)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => (($true = ((s__destination @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__Moon))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Flying))))) [skolem symbol introduction]
63. ! [X0 : w,X1 : w] : ((($true = ((s__destination @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__Moon))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Flying)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) [skolemisation 56,62]
64. ! [X1 : w] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) => (($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) & ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))))) [skolem symbol introduction]
65. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) & ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Giving)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [skolemisation 58,64]
68. ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)))) [cnf transformation 61]
69. (((accreln1 @ s__Obligation @ X1 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X2)) = $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) [cnf transformation 55]
86. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__Flying))) [cnf transformation 63]
87. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK2 @ X1) @ s__Moon))) [cnf transformation 63]
88. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) [cnf transformation 63]
89. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) [cnf transformation 63]
90. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 65]
91. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 65]
92. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 65]
93. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 65]
98. (((accreln1 @ s__Obligation @ cw @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | (((s__destination @ X1 @ s__Bill @ X0)) != $true) [cnf transformation 59]
100. ($true = X0) | ($false = X0) [fool exhaustiveness axiom]
104. ($true != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | (((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((accreln1 @ s__Obligation @ cw @ X0)) = $false) [superposition 98,100]
105. (((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((accreln1 @ s__Obligation @ cw @ X0)) = $false) [trivial inequality removal 104]
106. ($true != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | ($true != ((s__patient @ X0 @ X2 @ X1))) | ($false = ((accreln1 @ s__Obligation @ cw @ X1))) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) [superposition 105,100]
107. ($true != ((s__patient @ X0 @ X2 @ X1))) | ($true != ((s__instance @ X0 @ s__Giving))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | ($false = ((accreln1 @ s__Obligation @ cw @ X1))) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) [trivial inequality removal 106]
240. ($true != $true) | ($true = ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [superposition 86,68]
247. ($true = ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [trivial inequality removal 240]
257. ($true != $true) | ($true = ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) [superposition 87,68]
264. ($true = ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) [trivial inequality removal 257]
272. ($true != $true) | ($true = ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) [superposition 88,68]
279. ($true = ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) [trivial inequality removal 272]
287. ($true != $true) | ($true = ((s__destination @ (sK1 @ (sK0 @ X0)) @ (sK2 @ (sK0 @ X0)) @ (sK0 @ X0)))) [superposition 89,68]
294. ($true = ((s__destination @ (sK1 @ (sK0 @ X0)) @ (sK2 @ (sK0 @ X0)) @ (sK0 @ X0)))) [trivial inequality removal 287]
306. ($true != $true) | ($true = ((accreln1 @ s__Obligation @ X1 @ (sK0 @ X0)))) | ($true != ((accreln1 @ s__Obligation @ X1 @ X0))) [superposition 69,68]
313. ($true != ((accreln1 @ s__Obligation @ X1 @ X0))) | ($true = ((accreln1 @ s__Obligation @ X1 @ (sK0 @ X0)))) [trivial inequality removal 306]
752. ($true != $true) | ($true = ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__Giving))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 90,68]
768. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__Giving))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 752]
832. ($true != $true) | ($true = ((s__instance @ (sK4 @ (sK0 @ X0)) @ s__PayCheck))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 91,68]
848. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__instance @ (sK4 @ (sK0 @ X0)) @ s__PayCheck))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 832]
1321. ($true != $true) | ($true = ((s__destination @ (sK3 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 93,68]
1343. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__destination @ (sK3 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 1321]
1384. ($true != $true) | ($true = ((s__patient @ (sK3 @ (sK0 @ X0)) @ (sK4 @ (sK0 @ X0)) @ (sK0 @ X0)))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 92,68]
1406. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__patient @ (sK3 @ (sK0 @ X0)) @ (sK4 @ (sK0 @ X0)) @ (sK0 @ X0)))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 1384]
2340. ($true != $true) | ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ (sK0 @ X0))))) [superposition 313,68]
2372. ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ (sK0 @ X0))))) [trivial inequality removal 2340]
2712. ($true != $true) | ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Giving))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [superposition 768,294]
2745. ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Giving))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [trivial inequality removal 2712]
2758. ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Giving))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [forward subsumption resolution 2745,279]
2761. ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Giving))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [forward subsumption resolution 2758,264]
2772. ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Giving))) [forward subsumption resolution 2761,247]
2809. ($true != $true) | ($true = ((s__instance @ (sK4 @ (sK0 @ (sK0 @ X0))) @ s__PayCheck))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [superposition 848,294]
2842. ($true = ((s__instance @ (sK4 @ (sK0 @ (sK0 @ X0))) @ s__PayCheck))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [trivial inequality removal 2809]
2864. ($true = ((s__instance @ (sK4 @ (sK0 @ (sK0 @ X0))) @ s__PayCheck))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [forward subsumption resolution 2842,279]
2874. ($true = ((s__instance @ (sK4 @ (sK0 @ (sK0 @ X0))) @ s__PayCheck))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [forward subsumption resolution 2864,264]
2876. ($true = ((s__instance @ (sK4 @ (sK0 @ (sK0 @ X0))) @ s__PayCheck))) [forward subsumption resolution 2874,247]
3211. ($true != $true) | ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [superposition 1343,294]
3244. ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [trivial inequality removal 3211]
3257. ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [forward subsumption resolution 3244,279]
3259. ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [forward subsumption resolution 3257,264]
3261. ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) [forward subsumption resolution 3259,247]
6496. ($true != $true) | ($true = ((s__patient @ (sK3 @ (sK0 @ (sK0 @ X0))) @ (sK4 @ (sK0 @ (sK0 @ X0))) @ (sK0 @ (sK0 @ X0))))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [superposition 1406,294]
6529. ($true = ((s__patient @ (sK3 @ (sK0 @ (sK0 @ X0))) @ (sK4 @ (sK0 @ (sK0 @ X0))) @ (sK0 @ (sK0 @ X0))))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [trivial inequality removal 6496]
6542. ($true = ((s__patient @ (sK3 @ (sK0 @ (sK0 @ X0))) @ (sK4 @ (sK0 @ (sK0 @ X0))) @ (sK0 @ (sK0 @ X0))))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [forward subsumption resolution 6529,279]
6544. ($true = ((s__patient @ (sK3 @ (sK0 @ (sK0 @ X0))) @ (sK4 @ (sK0 @ (sK0 @ X0))) @ (sK0 @ (sK0 @ X0))))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [forward subsumption resolution 6542,264]
6546. ($true = ((s__patient @ (sK3 @ (sK0 @ (sK0 @ X0))) @ (sK4 @ (sK0 @ (sK0 @ X0))) @ (sK0 @ (sK0 @ X0))))) [forward subsumption resolution 6544,247]
14813. ($true != $true) | ($true != ((s__instance @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Giving))) | ($true != ((s__instance @ (sK4 @ (sK0 @ (sK0 @ X0))) @ s__PayCheck))) | ($false = ((accreln1 @ s__Obligation @ cw @ (sK0 @ (sK0 @ X0))))) | ($false = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) [superposition 107,6546]
14814. ($true != ((s__instance @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Giving))) | ($true != ((s__instance @ (sK4 @ (sK0 @ (sK0 @ X0))) @ s__PayCheck))) | ($false = ((accreln1 @ s__Obligation @ cw @ (sK0 @ (sK0 @ X0))))) | ($false = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) [trivial inequality removal 14813]
14826. ($true != ((s__instance @ (sK4 @ (sK0 @ (sK0 @ X0))) @ s__PayCheck))) | ($false = ((accreln1 @ s__Obligation @ cw @ (sK0 @ (sK0 @ X0))))) | ($false = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) [forward subsumption resolution 14814,2772]
14859. ($false = ((accreln1 @ s__Obligation @ cw @ (sK0 @ (sK0 @ X0))))) | ($false = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ X0))) @ s__Bill @ (sK0 @ (sK0 @ X0))))) [forward subsumption resolution 14826,2876]
14876. ($true = $false) | ($false = ((accreln1 @ s__Obligation @ cw @ (sK0 @ (sK0 @ X0))))) [forward demodulation 14859,3261]
14877. ($false = ((accreln1 @ s__Obligation @ cw @ (sK0 @ (sK0 @ X0))))) [trivial inequality removal 14876]
25808. ($true = $false) [superposition 2372,14877]
25824. $false [trivial inequality removal 25808]
% SZS output end Proof for 1g_Obligation_UO1_existConj
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 3.729 s
% Peak memory usage: 20 MB
% Instructions burned: 26190 (million)
% ------------------------------
% ------------------------------
