 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/O7yWjfFj05/SOT_BE6gJR
% (771444)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_BE6gJR for (1ds/18Mi)
% (771440)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_BE6gJR for (1ds/27Mi)
% (771439)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_BE6gJR for (1ds/4Mi)
% (771442)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_BE6gJR for (1ds/2Mi)
% (771443)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_BE6gJR for (1ds/275Mi)
% (771438)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_BE6gJR for (1ds/183Mi)
% (771441)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_BE6gJR for (1ds/2Mi)
% (771442)Refutation not found, incomplete strategy
% (771442)------------------------------
% (771442)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (771442)Termination reason: Refutation not found, incomplete strategy


% (771442)Memory used [KB]: 5628
% (771442)Time elapsed: 0.004 s
% (771442)------------------------------
% (771442)------------------------------
% (771440)First to succeed.
% (771441)Also succeeded, but the first one will report.
% (771438)Also succeeded, but the first one will report.
% (771440)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_BE6gJR
% SZS output start Proof for SOT_BE6gJR
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, s__confersObligation: m).
thf(func_def_5, type, accreln1: m > w > w > $o).
thf(func_def_6, type, accreln2: m > $i > w > w > $o).
thf(func_def_7, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_8, type, s__highestAuth: $i > $o).
thf(func_def_10, type, s__instance: $i > $i > $o).
thf(func_def_11, type, s__agent: $i > $i > w > $o).
thf(func_def_12, type, s__patient: $i > $i > w > $o).
thf(func_def_13, type, s__origin: $i > $i > w > $o).
thf(func_def_14, type, s__destination: $i > $i > w > $o).
thf(func_def_27, type, vEPSILON: !>[X0: $tType]:((X0 > $o) > X0)).
thf(func_def_30, type, sK0: w > $i).
thf(func_def_33, type, sK3: $i > $i > w > w).
thf(f149,plain,(
  $false),
  inference(avatar_sat_refutation,[status(thm)],[f127,f134,f145,f148])).
thf(f148,plain,(
  ~spl4_1 | ~spl4_3),
  inference(avatar_split_clause,[status(thm)],[f147,f142,f121])).
thf(f121,definition,(
  spl4_1 <=> ((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Tax) = $true)),
  introduced(definition,[new_symbols(naming,[spl4_1])],[avatar_definition])).
thf(f142,definition,(
  spl4_3 <=> ((s__origin @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK3 @ s__USG @ s__Bill @ cw)) = $true)),
  introduced(definition,[new_symbols(naming,[spl4_3])],[avatar_definition])).
thf(f147,plain,(
  ((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Tax) != $true) | ~spl4_3),
  inference(trivial_inequality_removal,[status(thm)],[f146])).
thf(f146,plain,(
  ((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Tax) != $true) | ($true != $true) | ~spl4_3),
  inference(superposition,[status(thm)],[f111,f144])).
thf(f144,plain,(
  ((s__origin @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK3 @ s__USG @ s__Bill @ cw)) = $true) | ~spl4_3),
  inference(avatar_component_clause,[status(thm)],[f142])).
thf(f111,plain,(
  ( ! [X0 : $i] : (((s__origin @ X0 @ s__Bill @ (sK3 @ s__USG @ s__Bill @ cw)) != $true) | ((s__instance @ X0 @ s__Tax) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f110])).
thf(f110,plain,(
  ( ! [X0 : $i] : (((s__instance @ X0 @ s__Tax) != $true) | ((s__origin @ X0 @ s__Bill @ (sK3 @ s__USG @ s__Bill @ cw)) != $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f106,f109])).
thf(f109,plain,(
  ( ! [X0 : $i,X1 : w] : (((accreln3 @ s__confersObligation @ s__USG @ X0 @ X1 @ (sK3 @ s__USG @ X0 @ X1)) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f108])).
thf(f108,plain,(
  ( ! [X0 : $i,X1 : w] : (((accreln3 @ s__confersObligation @ s__USG @ X0 @ X1 @ (sK3 @ s__USG @ X0 @ X1)) = $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f91,f94])).
thf(f94,plain,(
  ((s__highestAuth @ s__USG) = $true)),
  inference(cnf_transformation,[status(thm)],[f40])).
thf(f40,plain,(
  ((s__highestAuth @ s__USG) = $true)),
  inference(fool_elimination,[status(thm)],[f39])).
thf(f39,plain,(
  (s__highestAuth @ s__USG)),
  inference(rectify,[status(thm)],[f31])).
thf(f31,axiom,(
  (s__highestAuth @ s__USG)),
  file('/tmp/O7yWjfFj05/SOT_BE6gJR',optA_usg_over_ca)).
thf(f91,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__highestAuth @ X2) != $true) | ((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ (sK3 @ X2 @ X1 @ X0)) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f70])).
thf(f70,plain,(
  ! [X0 : w,X1 : $i,X2 : $i] : (((s__highestAuth @ X2) != $true) | ((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ (sK3 @ X2 @ X1 @ X0)) = $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f62,f69])).
thf(f69,plain,(
  ! [X0 : w,X1 : $i,X2 : $i] : (? [X3 : w] : ((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ X3) = $true) => ((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ (sK3 @ X2 @ X1 @ X0)) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f62,plain,(
  ! [X0 : w,X1 : $i,X2 : $i] : (((s__highestAuth @ X2) != $true) | ? [X3 : w] : ((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ X3) = $true))),
  inference(ennf_transformation,[status(thm)],[f44])).
thf(f44,plain,(
  ! [X1 : $i,X2 : $i,X0 : w] : (((s__highestAuth @ X2) = $true) => ? [X3 : w] : ((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ X3) = $true))),
  inference(fool_elimination,[status(thm)],[f43])).
thf(f43,plain,(
  ! [X0 : w,X1 : $i,X2 : $i] : ((s__highestAuth @ X2) => ? [X3 : w] : (accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ X3))),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w,X2 : $i,X1 : $i] : ((s__highestAuth @ X1) => ? [X3 : w] : (accreln3 @ s__confersObligation @ X1 @ X2 @ X0 @ X3))),
  file('/tmp/O7yWjfFj05/SOT_BE6gJR',serial_confersObligation)).
thf(f106,plain,(
  ( ! [X0 : w,X1 : $i] : (((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0) != $true) | ((s__origin @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__Tax) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f55])).
thf(f55,plain,(
  ! [X0 : w] : (! [X1 : $i] : (((s__instance @ X1 @ s__Tax) != $true) | ((s__origin @ X1 @ s__Bill @ X0) != $true)) | ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0) != $true))),
  inference(ennf_transformation,[status(thm)],[f46])).
thf(f46,plain,(
  ~? [X0 : w] : (? [X1 : $i] : (((s__origin @ X1 @ s__Bill @ X0) = $true) & ((s__instance @ X1 @ s__Tax) = $true)) & ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0) = $true))),
  inference(fool_elimination,[status(thm)],[f45])).
thf(f45,plain,(
  ~? [X0 : w] : (? [X1 : $i] : ((s__instance @ X1 @ s__Tax) & (s__origin @ X1 @ s__Bill @ X0)) & (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0))),
  inference(rectify,[status(thm)],[f33])).
thf(f33,negated_conjecture,(
  ~? [X8 : w] : (? [X7 : $i] : ((s__instance @ X7 @ s__Tax) & (s__origin @ X7 @ s__Bill @ X8)) & (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X8))),
  inference(negated_conjecture,[status(cth)],[f32])).
thf(f32,conjecture,(
  ? [X8 : w] : (? [X7 : $i] : ((s__instance @ X7 @ s__Tax) & (s__origin @ X7 @ s__Bill @ X8)) & (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X8))),
  file('/tmp/O7yWjfFj05/SOT_BE6gJR',conjecture_bill_taxed)).
thf(f145,plain,(
  spl4_3 | spl4_2),
  inference(avatar_split_clause,[status(thm)],[f140,f125,f142])).
thf(f125,definition,(
  spl4_2 <=> ! [X0 : $i] : (((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true))),
  introduced(definition,[new_symbols(naming,[spl4_2])],[avatar_definition])).
thf(f140,plain,(
  ( ! [X0 : $i] : (((s__instance @ X0 @ s__PayCheck) != $true) | ((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__origin @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK3 @ s__USG @ s__Bill @ cw)) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f139])).
thf(f139,plain,(
  ( ! [X0 : $i] : (((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true) | ($true != $true) | ((s__origin @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK3 @ s__USG @ s__Bill @ cw)) = $true)) )),
  inference(forward_demodulation,[status(thm)],[f138,f89])).
thf(f89,plain,(
  ((s__instance @ sK2 @ s__Giving) = $true)),
  inference(cnf_transformation,[status(thm)],[f68])).
thf(f68,plain,(
  ((s__instance @ sK1 @ s__PayCheck) = $true) & ((s__instance @ sK2 @ s__Giving) = $true) & ($true = (s__destination @ sK2 @ s__Bill @ cw)) & ((s__patient @ sK2 @ sK1 @ cw) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f52,f67])).
thf(f67,plain,(
  ? [X0 : $i,X1 : $i] : (((s__instance @ X0 @ s__PayCheck) = $true) & ((s__instance @ X1 @ s__Giving) = $true) & ((s__destination @ X1 @ s__Bill @ cw) = $true) & ((s__patient @ X1 @ X0 @ cw) = $true)) => (((s__instance @ sK1 @ s__PayCheck) = $true) & ((s__instance @ sK2 @ s__Giving) = $true) & ($true = (s__destination @ sK2 @ s__Bill @ cw)) & ((s__patient @ sK2 @ sK1 @ cw) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f52,plain,(
  ? [X0 : $i,X1 : $i] : (((s__instance @ X0 @ s__PayCheck) = $true) & ((s__instance @ X1 @ s__Giving) = $true) & ((s__destination @ X1 @ s__Bill @ cw) = $true) & ((s__patient @ X1 @ X0 @ cw) = $true))),
  inference(fool_elimination,[status(thm)],[f51])).
thf(f51,plain,(
  ? [X0 : $i,X1 : $i] : ((s__instance @ X0 @ s__PayCheck) & (s__instance @ X1 @ s__Giving) & (s__destination @ X1 @ s__Bill @ cw) & (s__patient @ X1 @ X0 @ cw))),
  inference(rectify,[status(thm)],[f30])).
thf(f30,axiom,(
  ? [X6 : $i,X5 : $i] : ((s__instance @ X6 @ s__PayCheck) & (s__instance @ X5 @ s__Giving) & (s__destination @ X5 @ s__Bill @ cw) & (s__patient @ X5 @ X6 @ cw))),
  file('/tmp/O7yWjfFj05/SOT_BE6gJR',fact_bill_gets_paycheck)).
thf(f138,plain,(
  ( ! [X0 : $i] : (((s__origin @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK3 @ s__USG @ s__Bill @ cw)) = $true) | ((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__instance @ sK2 @ s__Giving) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f137])).
thf(f137,plain,(
  ( ! [X0 : $i] : (($true != $true) | ((s__origin @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK3 @ s__USG @ s__Bill @ cw)) = $true) | ((s__instance @ X0 @ s__PayCheck) != $true) | ((s__instance @ sK2 @ s__Giving) != $true) | ((s__patient @ sK2 @ X0 @ cw) != $true)) )),
  inference(superposition,[status(thm)],[f131,f88])).
thf(f88,plain,(
  ($true = (s__destination @ sK2 @ s__Bill @ cw))),
  inference(cnf_transformation,[status(thm)],[f68])).
thf(f131,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__origin @ (sK0 @ (sK3 @ s__USG @ s__Bill @ X0)) @ s__Bill @ (sK3 @ s__USG @ s__Bill @ X0)) = $true) | ((s__patient @ X1 @ X2 @ X0) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f130])).
thf(f130,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__instance @ X2 @ s__PayCheck) != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true != $true) | ((s__origin @ (sK0 @ (sK3 @ s__USG @ s__Bill @ X0)) @ s__Bill @ (sK3 @ s__USG @ s__Bill @ X0)) = $true) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__patient @ X1 @ X2 @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f84,f109])).
thf(f84,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3) != $true) | ((s__origin @ (sK0 @ X3) @ s__Bill @ X3) = $true) | ((s__destination @ X0 @ s__Bill @ X2) != $true) | ((s__instance @ X0 @ s__Giving) != $true) | ((s__patient @ X0 @ X1 @ X2) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f66])).
thf(f66,plain,(
  ! [X0 : $i,X1 : $i,X2 : w,X3 : w] : (((s__patient @ X0 @ X1 @ X2) != $true) | ((s__instance @ X0 @ s__Giving) != $true) | ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | (((s__instance @ (sK0 @ X3) @ s__Tax) = $true) & ((s__origin @ (sK0 @ X3) @ s__Bill @ X3) = $true)) | ((s__destination @ X0 @ s__Bill @ X2) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f64,f65])).
thf(f65,plain,(
  ! [X3 : w] : (? [X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__origin @ X4 @ s__Bill @ X3) = $true)) => (((s__instance @ (sK0 @ X3) @ s__Tax) = $true) & ((s__origin @ (sK0 @ X3) @ s__Bill @ X3) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f64,plain,(
  ! [X0 : $i,X1 : $i,X2 : w,X3 : w] : (((s__patient @ X0 @ X1 @ X2) != $true) | ((s__instance @ X0 @ s__Giving) != $true) | ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ? [X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__origin @ X4 @ s__Bill @ X3) = $true)) | ((s__destination @ X0 @ s__Bill @ X2) != $true))),
  inference(rectify,[status(thm)],[f61])).
thf(f61,plain,(
  ! [X2 : $i,X0 : $i,X3 : w,X1 : w] : (((s__patient @ X2 @ X0 @ X3) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X3 @ X1) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true) | ? [X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__origin @ X4 @ s__Bill @ X1) = $true)) | ((s__destination @ X2 @ s__Bill @ X3) != $true))),
  inference(flattening,[status(thm)],[f60])).
thf(f60,plain,(
  ! [X3 : w,X1 : w,X2 : $i,X0 : $i] : ((? [X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__origin @ X4 @ s__Bill @ X1) = $true)) | ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X3 @ X1) != $true)) | (((s__patient @ X2 @ X0 @ X3) != $true) | ((s__destination @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true)))),
  inference(ennf_transformation,[status(thm)],[f50])).
thf(f50,plain,(
  ! [X3 : w,X1 : w,X2 : $i,X0 : $i] : ((((s__patient @ X2 @ X0 @ X3) = $true) & ((s__destination @ X2 @ s__Bill @ X3) = $true) & ((s__instance @ X2 @ s__Giving) = $true) & ((s__instance @ X0 @ s__PayCheck) = $true)) => (((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X3 @ X1) = $true) => ? [X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__origin @ X4 @ s__Bill @ X1) = $true))))),
  inference(fool_elimination,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__instance @ X2 @ s__Giving) & (s__patient @ X2 @ X0 @ X3) & (s__destination @ X2 @ s__Bill @ X3) & (s__instance @ X0 @ s__PayCheck)) => ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X3 @ X1) => ? [X4 : $i] : ((s__instance @ X4 @ s__Tax) & (s__origin @ X4 @ s__Bill @ X1))))),
  inference(rectify,[status(thm)],[f28])).
thf(f28,axiom,(
  ! [X6 : $i,X3 : w,X5 : $i,X0 : w] : (((s__instance @ X5 @ s__Giving) & (s__patient @ X5 @ X6 @ X0) & (s__destination @ X5 @ s__Bill @ X0) & (s__instance @ X6 @ s__PayCheck)) => ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X0 @ X3) => ? [X7 : $i] : ((s__instance @ X7 @ s__Tax) & (s__origin @ X7 @ s__Bill @ X3))))),
  file('/tmp/O7yWjfFj05/SOT_BE6gJR',obligation1_usg_bill_taxed)).
thf(f134,plain,(
  ~spl4_2),
  inference(avatar_contradiction_clause,[status(thm)],[f133])).
thf(f133,plain,(
  $false | ~spl4_2),
  inference(trivial_inequality_removal,[status(thm)],[f132])).
thf(f132,plain,(
  ($true != $true) | ~spl4_2),
  inference(superposition,[status(thm)],[f129,f90])).
thf(f90,plain,(
  ((s__instance @ sK1 @ s__PayCheck) = $true)),
  inference(cnf_transformation,[status(thm)],[f68])).
thf(f129,plain,(
  ((s__instance @ sK1 @ s__PayCheck) != $true) | ~spl4_2),
  inference(trivial_inequality_removal,[status(thm)],[f128])).
thf(f128,plain,(
  ($true != $true) | ((s__instance @ sK1 @ s__PayCheck) != $true) | ~spl4_2),
  inference(superposition,[status(thm)],[f126,f87])).
thf(f87,plain,(
  ((s__patient @ sK2 @ sK1 @ cw) = $true)),
  inference(cnf_transformation,[status(thm)],[f68])).
thf(f126,plain,(
  ( ! [X0 : $i] : (((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true)) ) | ~spl4_2),
  inference(avatar_component_clause,[status(thm)],[f125])).
thf(f127,plain,(
  spl4_1 | spl4_2),
  inference(avatar_split_clause,[status(thm)],[f119,f125,f121])).
thf(f119,plain,(
  ( ! [X0 : $i] : (((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Tax) = $true) | ((s__instance @ X0 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f118])).
thf(f118,plain,(
  ( ! [X0 : $i] : (((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Tax) = $true) | ((s__instance @ X0 @ s__PayCheck) != $true) | ($true != $true)) )),
  inference(forward_demodulation,[status(thm)],[f117,f89])).
thf(f117,plain,(
  ( ! [X0 : $i] : (((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Tax) = $true) | ((s__instance @ sK2 @ s__Giving) != $true) | ((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f116])).
thf(f116,plain,(
  ( ! [X0 : $i] : (($true != $true) | ((s__instance @ sK2 @ s__Giving) != $true) | ((s__patient @ sK2 @ X0 @ cw) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true) | ((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ cw)) @ s__Tax) = $true)) )),
  inference(superposition,[status(thm)],[f113,f88])).
thf(f113,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ X0)) @ s__Tax) = $true) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__patient @ X1 @ X2 @ X0) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f112])).
thf(f112,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != $true) | ((s__instance @ (sK0 @ (sK3 @ s__USG @ s__Bill @ X0)) @ s__Tax) = $true) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__patient @ X1 @ X2 @ X0) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f85,f109])).
thf(f85,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3) != $true) | ((s__patient @ X0 @ X1 @ X2) != $true) | ((s__instance @ (sK0 @ X3) @ s__Tax) = $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__destination @ X0 @ s__Bill @ X2) != $true) | ((s__instance @ X0 @ s__Giving) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f66])).
% SZS output end Proof for SOT_BE6gJR
% (771440)------------------------------
% (771440)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (771440)Termination reason: Refutation

% (771440)Memory used [KB]: 5756
% (771440)Time elapsed: 0.012 s
% (771440)------------------------------
% (771440)------------------------------
% (771433)Success in time 0.023 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_BE6gJR - Vampire---5.0.1 says Theorem - CPU = 0.01 WC = 0.06 
% OUTPUT: SOT_BE6gJR - Vampire---5.0.1 says Refutation - CPU = 0.01 WC = 0.06 

