% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 1i_Obligation_singleO
% SZS output start Proof for 1i_Obligation_singleO
23. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [input(axiom)]
24. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [input(axiom)]
25. ! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [input(conjecture)]
26. ~! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [negated conjecture 25]
29. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [rectify 23]
30. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [fool elimination 29]
31. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [rectify 24]
32. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) [fool elimination 31]
33. ~! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [rectify 26]
34. ~! [X0 : w] : ((((accreln1 @ s__Obligation @ cw @ X0)) = $true) => ? [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) = $true) & (((s__patient @ X1 @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__PayCheck)) = $true) & (((s__instance @ X1 @ s__Giving)) = $true))) [fool elimination 33]
56. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [flattening 30]
57. ! [X0 : w,X1 : w,X2,X3] : ((? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true))) [ennf transformation 56]
58. ! [X0 : w,X1 : w,X2,X3] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [flattening 57]
59. ? [X0 : w] : (! [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) [ennf transformation 34]
60. ! [X1 : w] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) => (($true = ((s__destination @ (sK0 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK0 @ X1) @ (sK1 @ X1) @ X1))) & ($true = ((s__instance @ (sK1 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK0 @ X1) @ s__Giving))))) [skolem symbol introduction]
61. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK0 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK0 @ X1) @ (sK1 @ X1) @ X1))) & ($true = ((s__instance @ (sK1 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK0 @ X1) @ s__Giving)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [skolemisation 58,60]
62. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) => (($true = ((s__destination @ sK2 @ sK3 @ cw))) & ($true = ((s__agent @ sK2 @ s__Bill @ cw))) & ($true = ((s__instance @ sK3 @ s__Moon))) & ($true = ((s__instance @ sK2 @ s__Flying)))) [skolem symbol introduction]
63. ($true = ((s__destination @ sK2 @ sK3 @ cw))) & ($true = ((s__agent @ sK2 @ s__Bill @ cw))) & ($true = ((s__instance @ sK3 @ s__Moon))) & ($true = ((s__instance @ sK2 @ s__Flying))) [skolemisation 32,62]
64. ? [X0 : w] : (! [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) => (! [X2,X1] : (($true != ((s__destination @ X1 @ s__Bill @ sK4))) | ($true != ((s__patient @ X1 @ X2 @ sK4))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & ($true = ((accreln1 @ s__Obligation @ cw @ sK4)))) [skolem symbol introduction]
65. ! [X1,X2] : (($true != ((s__destination @ X1 @ s__Bill @ sK4))) | ($true != ((s__patient @ X1 @ X2 @ sK4))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & ($true = ((accreln1 @ s__Obligation @ cw @ sK4))) [skolemisation 59,64]
88. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK0 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 61]
89. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 61]
90. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__patient @ (sK0 @ X1) @ (sK1 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 61]
91. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK0 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 61]
92. ($true = ((s__instance @ sK2 @ s__Flying))) [cnf transformation 63]
93. ($true = ((s__instance @ sK3 @ s__Moon))) [cnf transformation 63]
94. ($true = ((s__agent @ sK2 @ s__Bill @ cw))) [cnf transformation 63]
95. ($true = ((s__destination @ sK2 @ sK3 @ cw))) [cnf transformation 63]
96. ($true = ((accreln1 @ s__Obligation @ cw @ sK4))) [cnf transformation 65]
97. ($true != ((s__destination @ X1 @ s__Bill @ sK4))) | ($true != ((s__patient @ X1 @ X2 @ sK4))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) [cnf transformation 65]
99. ($true = X0) | ($false = X0) [fool exhaustiveness axiom]
103. ($true != $true) | ($true != ((s__patient @ X0 @ X1 @ sK4))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK4))) [superposition 97,99]
104. ($true != ((s__patient @ X0 @ X1 @ sK4))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK4))) [trivial inequality removal 103]
105. ($true != $true) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK4))) | ($false = ((s__patient @ X0 @ X1 @ sK4))) [superposition 104,99]
106. ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK4))) | ($false = ((s__patient @ X0 @ X1 @ sK4))) [trivial inequality removal 105]
213. ($true != $true) | ($true = ((s__instance @ (sK0 @ sK4) @ s__Giving))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 88,96]
220. ($true = ((s__instance @ (sK0 @ sK4) @ s__Giving))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 213]
222. 1 <=> ! [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true)) [avatar definition]
223. (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) <- (1) [avatar component clause 222]
225. 2 <=> ($true = ((s__instance @ (sK0 @ sK4) @ s__Giving))) [avatar definition]
227. ($true = ((s__instance @ (sK0 @ sK4) @ s__Giving))) <- (2) [avatar component clause 225]
228. 1 | 2 [avatar split clause 220,225,222]
241. 4 <=> ($false = ((s__destination @ (sK0 @ sK4) @ s__Bill @ sK4))) [avatar definition]
243. ($false = ((s__destination @ (sK0 @ sK4) @ s__Bill @ sK4))) <- (4) [avatar component clause 241]
247. ($true != $true) | ($true = ((s__instance @ (sK1 @ sK4) @ s__PayCheck))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 89,96]
254. ($true = ((s__instance @ (sK1 @ sK4) @ s__PayCheck))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 247]
256. 5 <=> ($true = ((s__instance @ (sK1 @ sK4) @ s__PayCheck))) [avatar definition]
258. ($true = ((s__instance @ (sK1 @ sK4) @ s__PayCheck))) <- (5) [avatar component clause 256]
259. 1 | 5 [avatar split clause 254,256,222]
264. ($true != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK4))) | ($false = ((s__patient @ X0 @ (sK1 @ sK4) @ sK4))) <- (5) [superposition 106,258]
265. ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK4))) | ($false = ((s__patient @ X0 @ (sK1 @ sK4) @ sK4))) <- (5) [trivial inequality removal 264]
270. ($true != $true) | ($true = ((s__destination @ (sK0 @ sK4) @ s__Bill @ sK4))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 91,96]
277. ($true = ((s__destination @ (sK0 @ sK4) @ s__Bill @ sK4))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 270]
279. 6 <=> ($true = ((s__destination @ (sK0 @ sK4) @ s__Bill @ sK4))) [avatar definition]
281. ($true = ((s__destination @ (sK0 @ sK4) @ s__Bill @ sK4))) <- (6) [avatar component clause 279]
282. 1 | 6 [avatar split clause 277,279,222]
290. ($true != $true) | ($true = ((s__patient @ (sK0 @ sK4) @ (sK1 @ sK4) @ sK4))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 90,96]
297. ($true = ((s__patient @ (sK0 @ sK4) @ (sK1 @ sK4) @ sK4))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 290]
298. ($true != $true) | ($false = ((s__destination @ (sK0 @ sK4) @ s__Bill @ sK4))) | ($false = ((s__patient @ (sK0 @ sK4) @ (sK1 @ sK4) @ sK4))) <- (2, 5) [superposition 265,227]
305. ($false = ((s__destination @ (sK0 @ sK4) @ s__Bill @ sK4))) | ($false = ((s__patient @ (sK0 @ sK4) @ (sK1 @ sK4) @ sK4))) <- (2, 5) [trivial inequality removal 298]
307. 7 <=> ($false = ((s__patient @ (sK0 @ sK4) @ (sK1 @ sK4) @ sK4))) [avatar definition]
309. ($false = ((s__patient @ (sK0 @ sK4) @ (sK1 @ sK4) @ sK4))) <- (7) [avatar component clause 307]
310. 7 | 4 | ~2 | ~5 [avatar split clause 305,256,225,241,307]
325. ($true != $true) | ($true != ((s__instance @ sK2 @ s__Flying))) | ($true != ((s__instance @ sK3 @ s__Moon))) | ($true != ((s__agent @ sK2 @ s__Bill @ cw))) <- (1) [superposition 223,95]
332. ($true != ((s__instance @ sK2 @ s__Flying))) | ($true != ((s__instance @ sK3 @ s__Moon))) | ($true != ((s__agent @ sK2 @ s__Bill @ cw))) <- (1) [trivial inequality removal 325]
333. ($true != ((s__instance @ sK3 @ s__Moon))) | ($true != ((s__agent @ sK2 @ s__Bill @ cw))) <- (1) [forward subsumption resolution 332,92]
334. ($true != ((s__agent @ sK2 @ s__Bill @ cw))) <- (1) [forward subsumption resolution 333,93]
335. $false <- (1) [forward subsumption resolution 334,94]
336. ~1 [avatar contradiction clause 335]
337. ($true = $false) <- (4, 6) [forward demodulation 281,243]
338. $false <- (4, 6) [trivial inequality removal 337]
339. ~4 | ~6 [avatar contradiction clause 338]
341. 8 <=> ($true = ((s__patient @ (sK0 @ sK4) @ (sK1 @ sK4) @ sK4))) [avatar definition]
343. ($true = ((s__patient @ (sK0 @ sK4) @ (sK1 @ sK4) @ sK4))) <- (8) [avatar component clause 341]
344. 1 | 8 [avatar split clause 297,341,222]
345. ($true = $false) <- (7, 8) [forward demodulation 343,309]
346. $false <- (7, 8) [trivial inequality removal 345]
347. ~7 | ~8 [avatar contradiction clause 346]
s1. 1 | 2 [sat_conversion 228]
s3. 1 | 5 [sat_conversion 259]
s4. 1 | 6 [sat_conversion 282]
s5. ~5 | ~2 | 4 | 7 [sat_conversion 310]
s6. ~1 [sat_conversion 336]
s7. ~6 | ~4 [sat_conversion 339]
s8. 1 | 8 [sat_conversion 344]
s9. ~8 | ~7 [sat_conversion 347]
s10. 8 [rat s8,s6]
s11. ~7 [rat s9,s10]
s12. ~5 | ~2 | 4 [rat s5,s11]
s13. 6 [rat s4,s6]
s14. ~4 [rat s7,s13]
s15. 5 [rat s3,s6]
s16. ~2 [rat s12,s14,s15]
s17. # [rat s1,s16,s6]
348. $false [avatar sat refutation s17]
% SZS output end Proof for 1i_Obligation_singleO
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.189 s
% Peak memory usage: 12 MB
% Instructions burned: 26 (million)
% ------------------------------
% ------------------------------
