 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/pYkZqv5H_O/SOT_2chjh2
% (2153221)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_2chjh2 for (1ds/27Mi)
% (2153223)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_2chjh2 for (1ds/2Mi)
% (2153225)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_2chjh2 for (1ds/18Mi)
% (2153219)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_2chjh2 for (1ds/183Mi)
% (2153222)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_2chjh2 for (1ds/2Mi)
% (2153220)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_2chjh2 for (1ds/4Mi)
% (2153224)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_2chjh2 for (1ds/275Mi)
% (2153223)Refutation not found, incomplete strategy
% (2153223)------------------------------
% (2153223)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2153223)Termination reason: Refutation not found, incomplete strategy


% (2153223)Memory used [KB]: 5628
% (2153223)Time elapsed: 0.005 s
% (2153223)------------------------------
% (2153223)------------------------------
% (2153226)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_2chjh2 for (1ds/3Mi)
% (2153221)First to succeed.
% (2153221)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_2chjh2
% SZS output start Proof for SOT_2chjh2
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m > w > w > $o).
thf(func_def_5, type, accreln2: m > $i > w > w > $o).
thf(func_def_6, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_8, type, s__instance: $i > $i > $o).
thf(func_def_9, type, s__agent: $i > $i > w > $o).
thf(func_def_10, type, s__patient: $i > $i > w > $o).
thf(func_def_11, type, s__destination: $i > $i > w > $o).
thf(func_def_12, type, s__located: $i > $i > w > $o).
thf(func_def_22, type, vEPSILON: !>[X0: $tType]:((X0 > $o) > X0)).
thf(func_def_25, type, sK0: w > $i).
thf(func_def_26, type, sK1: w > $i).
thf(func_def_27, type, sK2: w > $i).
thf(func_def_28, type, sK3: w > $i).
thf(func_def_30, type, sK5: w > w).
thf(func_def_31, type, sK6: w > w).
thf(f397,plain,(
  $false),
  inference(avatar_sat_refutation,[status(thm)],[f260,f268,f276,f284,f299,f396])).
thf(f396,plain,(
  ~spl7_1 | ~spl7_3 | ~spl7_4 | ~spl7_5),
  inference(avatar_contradiction_clause,[status(thm)],[f395])).
thf(f395,plain,(
  $false | (~spl7_1 | ~spl7_3 | ~spl7_4 | ~spl7_5)),
  inference(trivial_inequality_removal,[status(thm)],[f394])).
thf(f394,plain,(
  ($true != $true) | (~spl7_1 | ~spl7_3 | ~spl7_4 | ~spl7_5)),
  inference(superposition,[status(thm)],[f393,f347])).
thf(f347,plain,(
  ($true = (s__agent @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ~spl7_4),
  inference(trivial_inequality_removal,[status(thm)],[f336])).
thf(f336,plain,(
  ($true = (s__agent @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != $true) | ~spl7_4),
  inference(superposition,[status(thm)],[f275,f119])).
thf(f119,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ X0 @ (sK5 @ X0)) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f78])).
thf(f78,plain,(
  ! [X0 : w] : ((accreln1 @ s__Obligation @ X0 @ (sK5 @ X0)) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f46,f77])).
thf(f77,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ((accreln1 @ s__Obligation @ X0 @ (sK5 @ X0)) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f46,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f45])).
thf(f45,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  file('/tmp/pYkZqv5H_O/SOT_2chjh2',serial_obligation)).
thf(f275,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0))) ) | ~spl7_4),
  inference(avatar_component_clause,[status(thm)],[f274])).
thf(f274,definition,(
  spl7_4 <=> ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0)))),
  introduced(definition,[new_symbols(naming,[spl7_4])],[avatar_definition])).
thf(f393,plain,(
  ($true != (s__agent @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | (~spl7_1 | ~spl7_3 | ~spl7_5)),
  inference(trivial_inequality_removal,[status(thm)],[f392])).
thf(f392,plain,(
  ($true != (s__agent @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != $true) | (~spl7_1 | ~spl7_3 | ~spl7_5)),
  inference(forward_demodulation,[status(thm)],[f391,f310])).
thf(f310,plain,(
  ((s__instance @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Driving) = $true) | ~spl7_1),
  inference(trivial_inequality_removal,[status(thm)],[f300])).
thf(f300,plain,(
  ($true != $true) | ((s__instance @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Driving) = $true) | ~spl7_1),
  inference(superposition,[status(thm)],[f255,f119])).
thf(f255,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true)) ) | ~spl7_1),
  inference(avatar_component_clause,[status(thm)],[f254])).
thf(f254,definition,(
  spl7_1 <=> ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true))),
  introduced(definition,[new_symbols(naming,[spl7_1])],[avatar_definition])).
thf(f391,plain,(
  ($true != (s__agent @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ((s__instance @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Driving) != $true) | (~spl7_3 | ~spl7_5)),
  inference(trivial_inequality_removal,[status(thm)],[f390])).
thf(f390,plain,(
  ($true != $true) | ($true != (s__agent @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ((s__instance @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Driving) != $true) | (~spl7_3 | ~spl7_5)),
  inference(forward_demodulation,[status(thm)],[f389,f318])).
thf(f318,plain,(
  ($true = (s__instance @ (sK3 @ (sK5 @ (sK5 @ cw))) @ s__OfficeBuilding)) | ~spl7_3),
  inference(trivial_inequality_removal,[status(thm)],[f312])).
thf(f312,plain,(
  ($true = (s__instance @ (sK3 @ (sK5 @ (sK5 @ cw))) @ s__OfficeBuilding)) | ($true != $true) | ~spl7_3),
  inference(superposition,[status(thm)],[f267,f119])).
thf(f267,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding))) ) | ~spl7_3),
  inference(avatar_component_clause,[status(thm)],[f266])).
thf(f266,definition,(
  spl7_3 <=> ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding)))),
  introduced(definition,[new_symbols(naming,[spl7_3])],[avatar_definition])).
thf(f389,plain,(
  ($true != (s__agent @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != (s__instance @ (sK3 @ (sK5 @ (sK5 @ cw))) @ s__OfficeBuilding)) | ((s__instance @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Driving) != $true) | ~spl7_5),
  inference(trivial_inequality_removal,[status(thm)],[f388])).
thf(f388,plain,(
  ($true != (s__instance @ (sK3 @ (sK5 @ (sK5 @ cw))) @ s__OfficeBuilding)) | ((s__instance @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Driving) != $true) | ($true != (s__agent @ (sK2 @ (sK5 @ (sK5 @ cw))) @ s__Bill @ (sK5 @ (sK5 @ cw)))) | ($true != $true) | ~spl7_5),
  inference(superposition,[status(thm)],[f293,f359])).
thf(f359,plain,(
  ($true = (s__destination @ (sK2 @ (sK5 @ (sK5 @ cw))) @ (sK3 @ (sK5 @ (sK5 @ cw))) @ (sK5 @ (sK5 @ cw)))) | ~spl7_5),
  inference(trivial_inequality_removal,[status(thm)],[f348])).
thf(f348,plain,(
  ($true != $true) | ($true = (s__destination @ (sK2 @ (sK5 @ (sK5 @ cw))) @ (sK3 @ (sK5 @ (sK5 @ cw))) @ (sK5 @ (sK5 @ cw)))) | ~spl7_5),
  inference(superposition,[status(thm)],[f283,f119])).
thf(f283,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true)) ) | ~spl7_5),
  inference(avatar_component_clause,[status(thm)],[f282])).
thf(f282,definition,(
  spl7_5 <=> ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true))),
  introduced(definition,[new_symbols(naming,[spl7_5])],[avatar_definition])).
thf(f293,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__destination @ X0 @ X1 @ (sK5 @ (sK5 @ cw)))) | ((s__agent @ X0 @ s__Bill @ (sK5 @ (sK5 @ cw))) != $true) | ($true != (s__instance @ X0 @ s__Driving)) | ((s__instance @ X1 @ s__OfficeBuilding) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f286])).
thf(f286,plain,(
  ( ! [X0 : $i,X1 : $i] : (((s__agent @ X0 @ s__Bill @ (sK5 @ (sK5 @ cw))) != $true) | ($true != (s__destination @ X0 @ X1 @ (sK5 @ (sK5 @ cw)))) | ($true != $true) | ((s__instance @ X1 @ s__OfficeBuilding) != $true) | ($true != (s__instance @ X0 @ s__Driving))) )),
  inference(superposition,[status(thm)],[f189,f130])).
thf(f130,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Permission @ X0 @ (sK5 @ X0)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f129])).
thf(f129,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Permission @ X0 @ (sK5 @ X0))) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f118,f119])).
thf(f118,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((accreln1 @ s__Permission @ X0 @ X1) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f66])).
thf(f66,plain,(
  ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((accreln1 @ s__Permission @ X0 @ X1) = $true))),
  inference(ennf_transformation,[status(thm)],[f58])).
thf(f58,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ((accreln1 @ s__Permission @ X0 @ X1) = $true))),
  inference(fool_elimination,[status(thm)],[f57])).
thf(f57,plain,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  inference(rectify,[status(thm)],[f4])).
thf(f4,axiom,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  file('/tmp/pYkZqv5H_O/SOT_2chjh2',bridge_obligation_permission)).
thf(f189,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != (accreln1 @ s__Permission @ (sK5 @ cw) @ X2)) | ($true != (s__agent @ X0 @ s__Bill @ X2)) | ((s__destination @ X0 @ X1 @ X2) != $true) | ($true != (s__instance @ X0 @ s__Driving)) | ((s__instance @ X1 @ s__OfficeBuilding) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f185])).
thf(f185,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (($true != (s__instance @ X0 @ s__Driving)) | ($true != (accreln1 @ s__Permission @ (sK5 @ cw) @ X2)) | ($true != (s__agent @ X0 @ s__Bill @ X2)) | ($true != $true) | ((s__instance @ X1 @ s__OfficeBuilding) != $true) | ((s__destination @ X0 @ X1 @ X2) != $true)) )),
  inference(superposition,[status(thm)],[f117,f119])).
thf(f117,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X1)) | ($true != (s__destination @ X3 @ X2 @ X0)) | ((accreln1 @ s__Permission @ X1 @ X0) != $true) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__agent @ X3 @ s__Bill @ X0)) | ($true != (s__instance @ X3 @ s__Driving))) )),
  inference(cnf_transformation,[status(thm)],[f63])).
thf(f63,plain,(
  ! [X0 : w,X1 : w] : (! [X2 : $i,X3 : $i] : (($true != (s__destination @ X3 @ X2 @ X0)) | ($true != (s__instance @ X3 @ s__Driving)) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__agent @ X3 @ s__Bill @ X0))) | ($true != (accreln1 @ s__Obligation @ cw @ X1)) | ((accreln1 @ s__Permission @ X1 @ X0) != $true))),
  inference(ennf_transformation,[status(thm)],[f52])).
thf(f52,plain,(
  ~? [X0 : w,X1 : w] : (($true = (accreln1 @ s__Obligation @ cw @ X1)) & ((accreln1 @ s__Permission @ X1 @ X0) = $true) & ? [X3 : $i,X2 : $i] : (($true = (s__instance @ X3 @ s__Driving)) & ($true = (s__agent @ X3 @ s__Bill @ X0)) & ($true = (s__instance @ X2 @ s__OfficeBuilding)) & ($true = (s__destination @ X3 @ X2 @ X0))))),
  inference(fool_elimination,[status(thm)],[f51])).
thf(f51,plain,(
  ~? [X0 : w,X1 : w] : ((accreln1 @ s__Permission @ X1 @ X0) & (accreln1 @ s__Obligation @ cw @ X1) & ? [X2 : $i,X3 : $i] : ((s__instance @ X2 @ s__OfficeBuilding) & (s__instance @ X3 @ s__Driving) & (s__agent @ X3 @ s__Bill @ X0) & (s__destination @ X3 @ X2 @ X0)))),
  inference(rectify,[status(thm)],[f41])).
thf(f41,negated_conjecture,(
  ~? [X9 : w,X8 : w] : ((accreln1 @ s__Permission @ X8 @ X9) & (accreln1 @ s__Obligation @ cw @ X8) & ? [X7 : $i,X6 : $i] : ((s__instance @ X7 @ s__OfficeBuilding) & (s__instance @ X6 @ s__Driving) & (s__agent @ X6 @ s__Bill @ X9) & (s__destination @ X6 @ X7 @ X9)))),
  inference(negated_conjecture,[status(cth)],[f40])).
thf(f40,conjecture,(
  ? [X9 : w,X8 : w] : ((accreln1 @ s__Permission @ X8 @ X9) & (accreln1 @ s__Obligation @ cw @ X8) & ? [X7 : $i,X6 : $i] : ((s__instance @ X7 @ s__OfficeBuilding) & (s__instance @ X6 @ s__Driving) & (s__agent @ X6 @ s__Bill @ X9) & (s__destination @ X6 @ X7 @ X9)))),
  file('/tmp/pYkZqv5H_O/SOT_2chjh2',conjecture_is_permitted)).
thf(f299,plain,(
  spl7_2),
  inference(avatar_contradiction_clause,[status(thm)],[f298])).
thf(f298,plain,(
  $false | spl7_2),
  inference(trivial_inequality_removal,[status(thm)],[f297])).
thf(f297,plain,(
  ($true != $true) | spl7_2),
  inference(superposition,[status(thm)],[f259,f236])).
thf(f236,plain,(
  ($true = (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f235])).
thf(f235,plain,(
  ($true != $true) | ($true = (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw)))),
  inference(forward_demodulation,[status(thm)],[f234,f112])).
thf(f112,plain,(
  ($true = (s__instance @ sK4 @ s__PayCheck))),
  inference(cnf_transformation,[status(thm)],[f76])).
thf(f76,plain,(
  ((s__destination @ sK4 @ s__Bill @ cw) = $true) & ($true = (s__instance @ sK4 @ s__PayCheck))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f56,f75])).
thf(f75,plain,(
  ? [X0 : $i] : (((s__destination @ X0 @ s__Bill @ cw) = $true) & ((s__instance @ X0 @ s__PayCheck) = $true)) => (((s__destination @ sK4 @ s__Bill @ cw) = $true) & ($true = (s__instance @ sK4 @ s__PayCheck)))),
  introduced(definition,[],[choice_axiom])).
thf(f56,plain,(
  ? [X0 : $i] : (((s__destination @ X0 @ s__Bill @ cw) = $true) & ((s__instance @ X0 @ s__PayCheck) = $true))),
  inference(fool_elimination,[status(thm)],[f55])).
thf(f55,plain,(
  ? [X0 : $i] : ((s__destination @ X0 @ s__Bill @ cw) & (s__instance @ X0 @ s__PayCheck))),
  inference(rectify,[status(thm)],[f39])).
thf(f39,axiom,(
  ? [X3 : $i] : ((s__destination @ X3 @ s__Bill @ cw) & (s__instance @ X3 @ s__PayCheck))),
  file('/tmp/pYkZqv5H_O/SOT_2chjh2',optionA_bill_gets_paycheck)).
thf(f234,plain,(
  ($true = (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true != (s__instance @ sK4 @ s__PayCheck))),
  inference(trivial_inequality_removal,[status(thm)],[f233])).
thf(f233,plain,(
  ($true != $true) | ($true = (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true != (s__instance @ sK4 @ s__PayCheck))),
  inference(superposition,[status(thm)],[f152,f113])).
thf(f113,plain,(
  ((s__destination @ sK4 @ s__Bill @ cw) = $true)),
  inference(cnf_transformation,[status(thm)],[f76])).
thf(f152,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true = (s__agent @ (sK1 @ (sK5 @ X0)) @ s__Bill @ (sK5 @ X0))) | ($true != (s__instance @ X1 @ s__PayCheck))) )),
  inference(trivial_inequality_removal,[status(thm)],[f149])).
thf(f149,plain,(
  ( ! [X0 : w,X1 : $i] : (($true = (s__agent @ (sK1 @ (sK5 @ X0)) @ s__Bill @ (sK5 @ X0))) | ((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true != $true) | ($true != (s__instance @ X1 @ s__PayCheck))) )),
  inference(superposition,[status(thm)],[f85,f119])).
thf(f85,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ((s__destination @ X1 @ s__Bill @ X2) != $true) | ((s__agent @ (sK1 @ X0) @ s__Bill @ X0) = $true) | ($true != (s__instance @ X1 @ s__PayCheck))) )),
  inference(cnf_transformation,[status(thm)],[f71])).
thf(f71,plain,(
  ! [X0 : w,X1 : $i,X2 : w] : ((((s__instance @ (sK0 @ X0) @ s__Tax) = $true) & ((s__agent @ (sK1 @ X0) @ s__Bill @ X0) = $true) & ((s__instance @ (sK1 @ X0) @ s__Giving) = $true) & ($true = (s__patient @ (sK1 @ X0) @ (sK0 @ X0) @ X0))) | ((s__destination @ X1 @ s__Bill @ X2) != $true) | ((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f69,f70])).
thf(f70,plain,(
  ! [X0 : w] : (? [X3 : $i,X4 : $i] : (((s__instance @ X3 @ s__Tax) = $true) & ((s__agent @ X4 @ s__Bill @ X0) = $true) & ((s__instance @ X4 @ s__Giving) = $true) & ($true = (s__patient @ X4 @ X3 @ X0))) => (((s__instance @ (sK0 @ X0) @ s__Tax) = $true) & ((s__agent @ (sK1 @ X0) @ s__Bill @ X0) = $true) & ((s__instance @ (sK1 @ X0) @ s__Giving) = $true) & ($true = (s__patient @ (sK1 @ X0) @ (sK0 @ X0) @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f69,plain,(
  ! [X0 : w,X1 : $i,X2 : w] : (? [X3 : $i,X4 : $i] : (((s__instance @ X3 @ s__Tax) = $true) & ((s__agent @ X4 @ s__Bill @ X0) = $true) & ((s__instance @ X4 @ s__Giving) = $true) & ($true = (s__patient @ X4 @ X3 @ X0))) | ((s__destination @ X1 @ s__Bill @ X2) != $true) | ((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)))),
  inference(rectify,[status(thm)],[f62])).
thf(f62,plain,(
  ! [X2 : w,X1 : $i,X0 : w] : (? [X3 : $i,X4 : $i] : (((s__instance @ X3 @ s__Tax) = $true) & ((s__agent @ X4 @ s__Bill @ X2) = $true) & ((s__instance @ X4 @ s__Giving) = $true) & ($true = (s__patient @ X4 @ X3 @ X2))) | ((s__destination @ X1 @ s__Bill @ X0) != $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)))),
  inference(flattening,[status(thm)],[f61])).
thf(f61,plain,(
  ! [X2 : w,X0 : w,X1 : $i] : ((? [X3 : $i,X4 : $i] : (((s__instance @ X3 @ s__Tax) = $true) & ((s__agent @ X4 @ s__Bill @ X2) = $true) & ((s__instance @ X4 @ s__Giving) = $true) & ($true = (s__patient @ X4 @ X3 @ X2))) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)) | (((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X1 @ s__PayCheck))))),
  inference(ennf_transformation,[status(thm)],[f50])).
thf(f50,plain,(
  ! [X2 : w,X0 : w,X1 : $i] : ((((s__destination @ X1 @ s__Bill @ X0) = $true) & ($true = (s__instance @ X1 @ s__PayCheck))) => (((accreln1 @ s__Obligation @ X0 @ X2) = $true) => ? [X3 : $i,X4 : $i] : (((s__instance @ X3 @ s__Tax) = $true) & ((s__agent @ X4 @ s__Bill @ X2) = $true) & ((s__instance @ X4 @ s__Giving) = $true) & ($true = (s__patient @ X4 @ X3 @ X2)))))),
  inference(fool_elimination,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X0 : w,X1 : $i,X2 : w] : (((s__destination @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__PayCheck)) => ((accreln1 @ s__Obligation @ X0 @ X2) => ? [X3 : $i,X4 : $i] : ((s__agent @ X4 @ s__Bill @ X2) & (s__instance @ X3 @ s__Tax) & (s__instance @ X4 @ s__Giving) & (s__patient @ X4 @ X3 @ X2))))),
  inference(rectify,[status(thm)],[f37])).
thf(f37,axiom,(
  ! [X0 : w,X3 : $i,X1 : w] : (((s__destination @ X3 @ s__Bill @ X0) & (s__instance @ X3 @ s__PayCheck)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X5 : $i,X4 : $i] : ((s__agent @ X4 @ s__Bill @ X1) & (s__instance @ X5 @ s__Tax) & (s__instance @ X4 @ s__Giving) & (s__patient @ X4 @ X5 @ X1))))),
  file('/tmp/pYkZqv5H_O/SOT_2chjh2',obligation1_pay_taxes)).
thf(f259,plain,(
  ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | spl7_2),
  inference(avatar_component_clause,[status(thm)],[f257])).
thf(f257,definition,(
  spl7_2 <=> ($true = (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw)))),
  introduced(definition,[new_symbols(naming,[spl7_2])],[avatar_definition])).
thf(f284,plain,(
  spl7_5 | ~spl7_2),
  inference(avatar_split_clause,[status(thm)],[f280,f257,f282])).
thf(f280,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f279])).
thf(f279,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true) | ($true != $true) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true)) )),
  inference(forward_demodulation,[status(thm)],[f278,f232])).
thf(f232,plain,(
  ($true = (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax))),
  inference(trivial_inequality_removal,[status(thm)],[f231])).
thf(f231,plain,(
  ($true = (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ($true != $true)),
  inference(forward_demodulation,[status(thm)],[f230,f112])).
thf(f230,plain,(
  ($true = (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ($true != (s__instance @ sK4 @ s__PayCheck))),
  inference(trivial_inequality_removal,[status(thm)],[f229])).
thf(f229,plain,(
  ($true != $true) | ($true = (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ($true != (s__instance @ sK4 @ s__PayCheck))),
  inference(superposition,[status(thm)],[f148,f113])).
thf(f148,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true = (s__instance @ (sK0 @ (sK5 @ X0)) @ s__Tax)) | ($true != (s__instance @ X1 @ s__PayCheck))) )),
  inference(trivial_inequality_removal,[status(thm)],[f145])).
thf(f145,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true != $true) | ($true = (s__instance @ (sK0 @ (sK5 @ X0)) @ s__Tax)) | ($true != (s__instance @ X1 @ s__PayCheck))) )),
  inference(superposition,[status(thm)],[f86,f119])).
thf(f86,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ((s__destination @ X1 @ s__Bill @ X2) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)) | ((s__instance @ (sK0 @ X0) @ s__Tax) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f71])).
thf(f278,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f277])).
thf(f277,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax))) )),
  inference(forward_demodulation,[status(thm)],[f245,f228])).
thf(f228,plain,(
  ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f227])).
thf(f227,plain,(
  ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) = $true) | ($true != $true)),
  inference(forward_demodulation,[status(thm)],[f226,f112])).
thf(f226,plain,(
  ($true != (s__instance @ sK4 @ s__PayCheck)) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f225])).
thf(f225,plain,(
  ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) = $true) | ($true != $true) | ($true != (s__instance @ sK4 @ s__PayCheck))),
  inference(superposition,[status(thm)],[f134,f113])).
thf(f134,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true = (s__instance @ (sK1 @ (sK5 @ X0)) @ s__Giving)) | ($true != (s__instance @ X1 @ s__PayCheck))) )),
  inference(trivial_inequality_removal,[status(thm)],[f133])).
thf(f133,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__instance @ X1 @ s__PayCheck)) | ((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true != $true) | ($true = (s__instance @ (sK1 @ (sK5 @ X0)) @ s__Giving))) )),
  inference(superposition,[status(thm)],[f84,f119])).
thf(f84,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)) | ((s__instance @ (sK1 @ X0) @ s__Giving) = $true) | ((s__destination @ X1 @ s__Bill @ X2) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f71])).
thf(f245,plain,(
  ( ! [X0 : w] : (((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f241])).
thf(f241,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true != $true) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true)) )),
  inference(superposition,[status(thm)],[f90,f240])).
thf(f240,plain,(
  ((s__patient @ (sK1 @ (sK5 @ cw)) @ (sK0 @ (sK5 @ cw)) @ (sK5 @ cw)) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f239])).
thf(f239,plain,(
  ((s__patient @ (sK1 @ (sK5 @ cw)) @ (sK0 @ (sK5 @ cw)) @ (sK5 @ cw)) = $true) | ($true != $true)),
  inference(forward_demodulation,[status(thm)],[f238,f112])).
thf(f238,plain,(
  ($true != (s__instance @ sK4 @ s__PayCheck)) | ((s__patient @ (sK1 @ (sK5 @ cw)) @ (sK0 @ (sK5 @ cw)) @ (sK5 @ cw)) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f237])).
thf(f237,plain,(
  ((s__patient @ (sK1 @ (sK5 @ cw)) @ (sK0 @ (sK5 @ cw)) @ (sK5 @ cw)) = $true) | ($true != $true) | ($true != (s__instance @ sK4 @ s__PayCheck))),
  inference(superposition,[status(thm)],[f169,f113])).
thf(f169,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)) | ($true = (s__patient @ (sK1 @ (sK5 @ X0)) @ (sK0 @ (sK5 @ X0)) @ (sK5 @ X0)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f165])).
thf(f165,plain,(
  ( ! [X0 : w,X1 : $i] : (($true = (s__patient @ (sK1 @ (sK5 @ X0)) @ (sK0 @ (sK5 @ X0)) @ (sK5 @ X0))) | ($true != (s__instance @ X1 @ s__PayCheck)) | ($true != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f83,f119])).
thf(f83,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ X2 @ X0) != $true) | ((s__destination @ X1 @ s__Bill @ X2) != $true) | ($true != (s__instance @ X1 @ s__PayCheck)) | ($true = (s__patient @ (sK1 @ X0) @ (sK0 @ X0) @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f71])).
thf(f90,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X2 @ X3)) | ((s__agent @ X1 @ s__Bill @ X3) != $true) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true) | ($true != (s__instance @ X1 @ s__Giving)) | ((accreln1 @ s__Obligation @ X3 @ X0) != $true) | ($true != (s__instance @ X2 @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f74,plain,(
  ! [X0 : w,X1 : $i,X2 : $i,X3 : w] : (((accreln1 @ s__Obligation @ X3 @ X0) != $true) | (($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding)) & ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) & ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true)) | ($true != (s__instance @ X2 @ s__Tax)) | ((s__agent @ X1 @ s__Bill @ X3) != $true) | ($true != (s__patient @ X1 @ X2 @ X3)) | ($true != (s__instance @ X1 @ s__Giving)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f72,f73])).
thf(f73,plain,(
  ! [X0 : w] : (? [X4 : $i,X5 : $i] : (((s__agent @ X4 @ s__Bill @ X0) = $true) & ((s__instance @ X5 @ s__OfficeBuilding) = $true) & ((s__instance @ X4 @ s__Driving) = $true) & ((s__destination @ X4 @ X5 @ X0) = $true)) => (($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding)) & ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) & ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f72,plain,(
  ! [X0 : w,X1 : $i,X2 : $i,X3 : w] : (((accreln1 @ s__Obligation @ X3 @ X0) != $true) | ? [X4 : $i,X5 : $i] : (((s__agent @ X4 @ s__Bill @ X0) = $true) & ((s__instance @ X5 @ s__OfficeBuilding) = $true) & ((s__instance @ X4 @ s__Driving) = $true) & ((s__destination @ X4 @ X5 @ X0) = $true)) | ($true != (s__instance @ X2 @ s__Tax)) | ((s__agent @ X1 @ s__Bill @ X3) != $true) | ($true != (s__patient @ X1 @ X2 @ X3)) | ($true != (s__instance @ X1 @ s__Giving)))),
  inference(rectify,[status(thm)],[f65])).
thf(f65,plain,(
  ! [X0 : w,X3 : $i,X2 : $i,X1 : w] : (($true != (accreln1 @ s__Obligation @ X1 @ X0)) | ? [X5 : $i,X4 : $i] : (($true = (s__agent @ X5 @ s__Bill @ X0)) & ($true = (s__instance @ X4 @ s__OfficeBuilding)) & ((s__instance @ X5 @ s__Driving) = $true) & ($true = (s__destination @ X5 @ X4 @ X0))) | ($true != (s__instance @ X2 @ s__Tax)) | ((s__agent @ X3 @ s__Bill @ X1) != $true) | ((s__patient @ X3 @ X2 @ X1) != $true) | ((s__instance @ X3 @ s__Giving) != $true))),
  inference(flattening,[status(thm)],[f64])).
thf(f64,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : ((? [X5 : $i,X4 : $i] : (($true = (s__agent @ X5 @ s__Bill @ X0)) & ($true = (s__instance @ X4 @ s__OfficeBuilding)) & ((s__instance @ X5 @ s__Driving) = $true) & ($true = (s__destination @ X5 @ X4 @ X0))) | ($true != (accreln1 @ s__Obligation @ X1 @ X0))) | (($true != (s__instance @ X2 @ s__Tax)) | ((s__instance @ X3 @ s__Giving) != $true) | ((s__agent @ X3 @ s__Bill @ X1) != $true) | ((s__patient @ X3 @ X2 @ X1) != $true)))),
  inference(ennf_transformation,[status(thm)],[f48])).
thf(f48,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : ((($true = (s__instance @ X2 @ s__Tax)) & ((s__instance @ X3 @ s__Giving) = $true) & ((s__agent @ X3 @ s__Bill @ X1) = $true) & ((s__patient @ X3 @ X2 @ X1) = $true)) => (($true = (accreln1 @ s__Obligation @ X1 @ X0)) => ? [X5 : $i,X4 : $i] : (($true = (s__agent @ X5 @ s__Bill @ X0)) & ($true = (s__instance @ X4 @ s__OfficeBuilding)) & ((s__instance @ X5 @ s__Driving) = $true) & ($true = (s__destination @ X5 @ X4 @ X0)))))),
  inference(fool_elimination,[status(thm)],[f47])).
thf(f47,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (((s__instance @ X2 @ s__Tax) & (s__instance @ X3 @ s__Giving) & (s__agent @ X3 @ s__Bill @ X1) & (s__patient @ X3 @ X2 @ X1)) => ((accreln1 @ s__Obligation @ X1 @ X0) => ? [X4 : $i,X5 : $i] : ((s__destination @ X5 @ X4 @ X0) & (s__instance @ X4 @ s__OfficeBuilding) & (s__instance @ X5 @ s__Driving) & (s__agent @ X5 @ s__Bill @ X0))))),
  inference(rectify,[status(thm)],[f38])).
thf(f38,axiom,(
  ! [X1 : w,X0 : w,X5 : $i,X4 : $i] : (((s__instance @ X5 @ s__Tax) & (s__instance @ X4 @ s__Giving) & (s__agent @ X4 @ s__Bill @ X0) & (s__patient @ X4 @ X5 @ X0)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X7 : $i,X6 : $i] : ((s__destination @ X6 @ X7 @ X1) & (s__instance @ X7 @ s__OfficeBuilding) & (s__instance @ X6 @ s__Driving) & (s__agent @ X6 @ s__Bill @ X1))))),
  file('/tmp/pYkZqv5H_O/SOT_2chjh2',obligation2_drive_to_office)).
thf(f276,plain,(
  spl7_4 | ~spl7_2),
  inference(avatar_split_clause,[status(thm)],[f272,f257,f274])).
thf(f272,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f271])).
thf(f271,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0))) )),
  inference(forward_demodulation,[status(thm)],[f270,f232])).
thf(f270,plain,(
  ( ! [X0 : w] : (($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0)) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f269])).
thf(f269,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0)) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ($true != $true)) )),
  inference(forward_demodulation,[status(thm)],[f246,f228])).
thf(f246,plain,(
  ( ! [X0 : w] : (((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0)) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax))) )),
  inference(trivial_inequality_removal,[status(thm)],[f242])).
thf(f242,plain,(
  ( ! [X0 : w] : (((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0)) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax))) )),
  inference(superposition,[status(thm)],[f93,f240])).
thf(f93,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X2 @ X3)) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0)) | ($true != (s__instance @ X1 @ s__Giving)) | ((s__agent @ X1 @ s__Bill @ X3) != $true) | ((accreln1 @ s__Obligation @ X3 @ X0) != $true) | ($true != (s__instance @ X2 @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f268,plain,(
  spl7_3 | ~spl7_2),
  inference(avatar_split_clause,[status(thm)],[f264,f257,f266])).
thf(f264,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding))) )),
  inference(trivial_inequality_removal,[status(thm)],[f263])).
thf(f263,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding)) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw)))) )),
  inference(forward_demodulation,[status(thm)],[f262,f228])).
thf(f262,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding)) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f261])).
thf(f261,plain,(
  ( ! [X0 : w] : (((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding))) )),
  inference(forward_demodulation,[status(thm)],[f247,f232])).
thf(f247,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding)) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f243])).
thf(f243,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding)) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true) | ($true != $true) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f92,f240])).
thf(f92,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X2 @ X3)) | ($true != (s__instance @ X1 @ s__Giving)) | ($true = (s__instance @ (sK3 @ X0) @ s__OfficeBuilding)) | ((s__agent @ X1 @ s__Bill @ X3) != $true) | ($true != (s__instance @ X2 @ s__Tax)) | ((accreln1 @ s__Obligation @ X3 @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f260,plain,(
  spl7_1 | ~spl7_2),
  inference(avatar_split_clause,[status(thm)],[f252,f257,f254])).
thf(f252,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f251])).
thf(f251,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true != $true) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true)) )),
  inference(forward_demodulation,[status(thm)],[f250,f228])).
thf(f250,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true) | ($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f249])).
thf(f249,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true != $true) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true)) )),
  inference(forward_demodulation,[status(thm)],[f248,f232])).
thf(f248,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f244])).
thf(f244,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK5 @ cw)) @ s__Bill @ (sK5 @ cw))) | ($true != $true) | ((s__instance @ (sK1 @ (sK5 @ cw)) @ s__Giving) != $true) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) | ($true != (s__instance @ (sK0 @ (sK5 @ cw)) @ s__Tax)) | ((accreln1 @ s__Obligation @ (sK5 @ cw) @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f91,f240])).
thf(f91,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X2 @ X3)) | ((s__agent @ X1 @ s__Bill @ X3) != $true) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) | ($true != (s__instance @ X1 @ s__Giving)) | ((accreln1 @ s__Obligation @ X3 @ X0) != $true) | ($true != (s__instance @ X2 @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f74])).
% SZS output end Proof for SOT_2chjh2
% (2153221)------------------------------
% (2153221)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2153221)Termination reason: Refutation

% (2153221)Memory used [KB]: 5884
% (2153221)Time elapsed: 0.036 s
% (2153221)------------------------------
% (2153221)------------------------------
% (2153215)Success in time 0.047 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_2chjh2 - Vampire---5.0.1 says Theorem - CPU = 0.00 WC = 0.09 
% OUTPUT: SOT_2chjh2 - Vampire---5.0.1 says Refutation - CPU = 0.00 WC = 0.09 

