 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/w7nKvtM0BM/SOT_6fVAoP
% (2148410)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_6fVAoP for (1ds/4Mi)
% (2148413)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_6fVAoP for (1ds/2Mi)
% (2148412)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_6fVAoP for (1ds/2Mi)
% (2148415)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_6fVAoP for (1ds/18Mi)
% (2148411)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_6fVAoP for (1ds/27Mi)
% (2148413)Refutation not found, incomplete strategy
% (2148413)------------------------------
% (2148413)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2148413)Termination reason: Refutation not found, incomplete strategy


% (2148413)Memory used [KB]: 5628
% (2148413)Time elapsed: 0.005 s
% (2148413)------------------------------
% (2148413)------------------------------
% (2148409)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_6fVAoP for (1ds/183Mi)
% (2148414)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_6fVAoP for (1ds/275Mi)
% (2148416)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_6fVAoP for (1ds/3Mi)
% (2148412)First to succeed.
% (2148412)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_6fVAoP
% SZS output start Proof for SOT_6fVAoP
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m > w > w > $o).
thf(func_def_5, type, accreln2: m > $i > w > w > $o).
thf(func_def_6, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_8, type, s__instance: $i > $i > $o).
thf(func_def_9, type, s__agent: $i > $i > w > $o).
thf(func_def_10, type, s__patient: $i > $i > w > $o).
thf(func_def_11, type, s__destination: $i > $i > w > $o).
thf(func_def_12, type, s__located: $i > $i > w > $o).
thf(func_def_24, type, sK0: w > w).
thf(func_def_26, type, sK2: w > $i).
thf(func_def_27, type, sK3: w > $i).
thf(func_def_28, type, sK4: w > $i).
thf(func_def_29, type, sK5: w > $i).
thf(func_def_30, type, sK6: w > $i).
thf(func_def_31, type, sK7: w > $i).
thf(f814,plain,(
  $false),
  inference(subsumption_resolution,[status(thm)],[f813,f441])).
thf(f441,plain,(
  ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f440,f374])).
thf(f374,plain,(
  ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) = $true)),
  inference(subsumption_resolution,[status(thm)],[f373,f363])).
thf(f363,plain,(
  ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) = $true)),
  inference(subsumption_resolution,[status(thm)],[f362,f101])).
thf(f101,plain,(
  ((s__instance @ sK1 @ s__PayCheck) = $true)),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f74,plain,(
  ((s__instance @ sK1 @ s__PayCheck) = $true) & ((s__destination @ sK1 @ s__Bill @ cw) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f48,f73])).
thf(f73,plain,(
  ? [X0 : $i] : (((s__instance @ X0 @ s__PayCheck) = $true) & ((s__destination @ X0 @ s__Bill @ cw) = $true)) => (((s__instance @ sK1 @ s__PayCheck) = $true) & ((s__destination @ sK1 @ s__Bill @ cw) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f48,plain,(
  ? [X0 : $i] : (((s__instance @ X0 @ s__PayCheck) = $true) & ((s__destination @ X0 @ s__Bill @ cw) = $true))),
  inference(fool_elimination,[status(thm)],[f47])).
thf(f47,plain,(
  ? [X0 : $i] : ((s__instance @ X0 @ s__PayCheck) & (s__destination @ X0 @ s__Bill @ cw))),
  inference(rectify,[status(thm)],[f39])).
thf(f39,axiom,(
  ? [X3 : $i] : ((s__instance @ X3 @ s__PayCheck) & (s__destination @ X3 @ s__Bill @ cw))),
  file('/tmp/w7nKvtM0BM/SOT_6fVAoP',optionA_bill_gets_paycheck)).
thf(f362,plain,(
  ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) = $true) | ((s__instance @ sK1 @ s__PayCheck) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f361])).
thf(f361,plain,(
  ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) = $true) | ((s__instance @ sK1 @ s__PayCheck) != $true) | ($true != $true)),
  inference(superposition,[status(thm)],[f143,f100])).
thf(f100,plain,(
  ((s__destination @ sK1 @ s__Bill @ cw) = $true)),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f143,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__instance @ (sK7 @ (sK0 @ X0)) @ s__Tax) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f142])).
thf(f142,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__instance @ (sK7 @ (sK0 @ X0)) @ s__Tax) = $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f121,f84])).
thf(f84,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f72])).
thf(f72,plain,(
  ! [X0 : w] : ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f46,f71])).
thf(f71,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f46,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f45])).
thf(f45,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  file('/tmp/w7nKvtM0BM/SOT_6fVAoP',serial_obligation)).
thf(f121,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__destination @ X2 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true) | ((s__instance @ (sK7 @ X1) @ s__Tax) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f83])).
thf(f83,plain,(
  ! [X0 : w,X1 : w,X2 : $i] : (((s__destination @ X2 @ s__Bill @ X0) != $true) | (((s__instance @ (sK7 @ X1) @ s__Tax) = $true) & ((s__instance @ (sK6 @ X1) @ s__Giving) = $true) & ((s__patient @ (sK6 @ X1) @ (sK7 @ X1) @ X1) = $true) & ((s__agent @ (sK6 @ X1) @ s__Bill @ X1) = $true)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f81,f82])).
thf(f82,plain,(
  ! [X1 : w] : (? [X3 : $i,X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__instance @ X3 @ s__Giving) = $true) & ((s__patient @ X3 @ X4 @ X1) = $true) & ((s__agent @ X3 @ s__Bill @ X1) = $true)) => (((s__instance @ (sK7 @ X1) @ s__Tax) = $true) & ((s__instance @ (sK6 @ X1) @ s__Giving) = $true) & ((s__patient @ (sK6 @ X1) @ (sK7 @ X1) @ X1) = $true) & ((s__agent @ (sK6 @ X1) @ s__Bill @ X1) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f81,plain,(
  ! [X0 : w,X1 : w,X2 : $i] : (((s__destination @ X2 @ s__Bill @ X0) != $true) | ? [X3 : $i,X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__instance @ X3 @ s__Giving) = $true) & ((s__patient @ X3 @ X4 @ X1) = $true) & ((s__agent @ X3 @ s__Bill @ X1) = $true)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true))),
  inference(rectify,[status(thm)],[f64])).
thf(f64,plain,(
  ! [X0 : w,X2 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ? [X3 : $i,X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__instance @ X3 @ s__Giving) = $true) & ($true = (s__patient @ X3 @ X4 @ X2)) & ((s__agent @ X3 @ s__Bill @ X2) = $true)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true))),
  inference(flattening,[status(thm)],[f63])).
thf(f63,plain,(
  ! [X2 : w,X0 : w,X1 : $i] : ((? [X3 : $i,X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__instance @ X3 @ s__Giving) = $true) & ($true = (s__patient @ X3 @ X4 @ X2)) & ((s__agent @ X3 @ s__Bill @ X2) = $true)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)) | (((s__instance @ X1 @ s__PayCheck) != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true)))),
  inference(ennf_transformation,[status(thm)],[f60])).
thf(f60,plain,(
  ! [X2 : w,X0 : w,X1 : $i] : ((((s__instance @ X1 @ s__PayCheck) = $true) & ((s__destination @ X1 @ s__Bill @ X0) = $true)) => (((accreln1 @ s__Obligation @ X0 @ X2) = $true) => ? [X3 : $i,X4 : $i] : (((s__instance @ X4 @ s__Tax) = $true) & ((s__instance @ X3 @ s__Giving) = $true) & ($true = (s__patient @ X3 @ X4 @ X2)) & ((s__agent @ X3 @ s__Bill @ X2) = $true))))),
  inference(fool_elimination,[status(thm)],[f59])).
thf(f59,plain,(
  ! [X0 : w,X1 : $i,X2 : w] : (((s__destination @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__PayCheck)) => ((accreln1 @ s__Obligation @ X0 @ X2) => ? [X3 : $i,X4 : $i] : ((s__instance @ X3 @ s__Giving) & (s__patient @ X3 @ X4 @ X2) & (s__instance @ X4 @ s__Tax) & (s__agent @ X3 @ s__Bill @ X2))))),
  inference(rectify,[status(thm)],[f36])).
thf(f36,axiom,(
  ! [X0 : w,X3 : $i,X1 : w] : (((s__destination @ X3 @ s__Bill @ X0) & (s__instance @ X3 @ s__PayCheck)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4 : $i,X5 : $i] : ((s__instance @ X4 @ s__Giving) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__Tax) & (s__agent @ X4 @ s__Bill @ X1))))),
  file('/tmp/w7nKvtM0BM/SOT_6fVAoP',obligation1_pay_taxes)).
thf(f373,plain,(
  ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) = $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true)),
  inference(subsumption_resolution,[status(thm)],[f372,f366])).
thf(f366,plain,(
  ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) = $true)),
  inference(subsumption_resolution,[status(thm)],[f365,f101])).
thf(f365,plain,(
  ((s__instance @ sK1 @ s__PayCheck) != $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f364])).
thf(f364,plain,(
  ((s__instance @ sK1 @ s__PayCheck) != $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) = $true) | ($true != $true)),
  inference(superposition,[status(thm)],[f156,f100])).
thf(f156,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__agent @ (sK6 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f154])).
thf(f154,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__agent @ (sK6 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)) = $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f118,f84])).
thf(f118,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__destination @ X2 @ s__Bill @ X0) != $true) | ((s__agent @ (sK6 @ X1) @ s__Bill @ X1) = $true) | ((s__instance @ X2 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f83])).
thf(f372,plain,(
  ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) = $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true)),
  inference(subsumption_resolution,[status(thm)],[f371,f360])).
thf(f360,plain,(
  ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) = $true)),
  inference(subsumption_resolution,[status(thm)],[f359,f101])).
thf(f359,plain,(
  ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) = $true) | ((s__instance @ sK1 @ s__PayCheck) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f358])).
thf(f358,plain,(
  ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) = $true) | ($true != $true) | ((s__instance @ sK1 @ s__PayCheck) != $true)),
  inference(superposition,[status(thm)],[f139,f100])).
thf(f139,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__instance @ (sK6 @ (sK0 @ X0)) @ s__Giving) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f138])).
thf(f138,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__instance @ (sK6 @ (sK0 @ X0)) @ s__Giving) = $true)) )),
  inference(superposition,[status(thm)],[f120,f84])).
thf(f120,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__instance @ X2 @ s__PayCheck) != $true) | ((s__instance @ (sK6 @ X1) @ s__Giving) = $true) | ((s__destination @ X2 @ s__Bill @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f83])).
thf(f371,plain,(
  ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) = $true) | ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f370])).
thf(f370,plain,(
  ($true != $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) = $true) | ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true)),
  inference(superposition,[status(thm)],[f187,f369])).
thf(f369,plain,(
  ((s__patient @ (sK6 @ (sK0 @ cw)) @ (sK7 @ (sK0 @ cw)) @ (sK0 @ cw)) = $true)),
  inference(subsumption_resolution,[status(thm)],[f368,f101])).
thf(f368,plain,(
  ((s__patient @ (sK6 @ (sK0 @ cw)) @ (sK7 @ (sK0 @ cw)) @ (sK0 @ cw)) = $true) | ((s__instance @ sK1 @ s__PayCheck) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f367])).
thf(f367,plain,(
  ($true != $true) | ((s__patient @ (sK6 @ (sK0 @ cw)) @ (sK7 @ (sK0 @ cw)) @ (sK0 @ cw)) = $true) | ((s__instance @ sK1 @ s__PayCheck) != $true)),
  inference(superposition,[status(thm)],[f161,f100])).
thf(f161,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__patient @ (sK6 @ (sK0 @ X0)) @ (sK7 @ (sK0 @ X0)) @ (sK0 @ X0)) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f158])).
thf(f158,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__PayCheck) != $true) | ((s__patient @ (sK6 @ (sK0 @ X0)) @ (sK7 @ (sK0 @ X0)) @ (sK0 @ X0)) = $true) | ($true != $true) | ((s__destination @ X1 @ s__Bill @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f119,f84])).
thf(f119,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__patient @ (sK6 @ X1) @ (sK7 @ X1) @ X1) = $true) | ((s__instance @ X2 @ s__PayCheck) != $true) | ((s__destination @ X2 @ s__Bill @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f83])).
thf(f187,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__patient @ X1 @ X2 @ X0) != $true) | ((s__instance @ X2 @ s__Tax) != $true) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Driving) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f182])).
thf(f182,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__instance @ X1 @ s__Giving) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__patient @ X1 @ X2 @ X0) != $true) | ($true != $true) | ((s__instance @ X2 @ s__Tax) != $true) | ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Driving) = $true)) )),
  inference(superposition,[status(thm)],[f104,f84])).
thf(f104,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ X3 @ X0) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) | ((s__patient @ X2 @ X1 @ X3) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f77])).
thf(f77,plain,(
  ! [X0 : w,X1 : $i,X2 : $i,X3 : w] : ((((s__instance @ (sK3 @ X0) @ s__OfficeBuilding) = $true) & ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true) & ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) & ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0))) | ((s__patient @ X2 @ X1 @ X3) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((accreln1 @ s__Obligation @ X3 @ X0) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f75,f76])).
thf(f76,plain,(
  ! [X0 : w] : (? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__OfficeBuilding) = $true) & ((s__destination @ X4 @ X5 @ X0) = $true) & ($true = (s__instance @ X4 @ s__Driving)) & ((s__agent @ X4 @ s__Bill @ X0) = $true)) => (((s__instance @ (sK3 @ X0) @ s__OfficeBuilding) = $true) & ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true) & ((s__instance @ (sK2 @ X0) @ s__Driving) = $true) & ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f75,plain,(
  ! [X0 : w,X1 : $i,X2 : $i,X3 : w] : (? [X4 : $i,X5 : $i] : (((s__instance @ X5 @ s__OfficeBuilding) = $true) & ((s__destination @ X4 @ X5 @ X0) = $true) & ($true = (s__instance @ X4 @ s__Driving)) & ((s__agent @ X4 @ s__Bill @ X0) = $true)) | ((s__patient @ X2 @ X1 @ X3) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((accreln1 @ s__Obligation @ X3 @ X0) != $true))),
  inference(rectify,[status(thm)],[f66])).
thf(f66,plain,(
  ! [X3 : w,X2 : $i,X1 : $i,X0 : w] : (? [X5 : $i,X4 : $i] : (((s__instance @ X4 @ s__OfficeBuilding) = $true) & ((s__destination @ X5 @ X4 @ X3) = $true) & ((s__instance @ X5 @ s__Driving) = $true) & ((s__agent @ X5 @ s__Bill @ X3) = $true)) | ((s__patient @ X1 @ X2 @ X0) != $true) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__instance @ X2 @ s__Tax) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((accreln1 @ s__Obligation @ X0 @ X3) != $true))),
  inference(flattening,[status(thm)],[f65])).
thf(f65,plain,(
  ! [X2 : $i,X1 : $i,X3 : w,X0 : w] : ((? [X5 : $i,X4 : $i] : (((s__instance @ X4 @ s__OfficeBuilding) = $true) & ((s__destination @ X5 @ X4 @ X3) = $true) & ((s__instance @ X5 @ s__Driving) = $true) & ((s__agent @ X5 @ s__Bill @ X3) = $true)) | ((accreln1 @ s__Obligation @ X0 @ X3) != $true)) | (((s__patient @ X1 @ X2 @ X0) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__Tax) != $true) | ((s__instance @ X1 @ s__Giving) != $true)))),
  inference(ennf_transformation,[status(thm)],[f50])).
thf(f50,plain,(
  ! [X2 : $i,X1 : $i,X3 : w,X0 : w] : ((((s__patient @ X1 @ X2 @ X0) = $true) & ((s__agent @ X1 @ s__Bill @ X0) = $true) & ((s__instance @ X2 @ s__Tax) = $true) & ((s__instance @ X1 @ s__Giving) = $true)) => (((accreln1 @ s__Obligation @ X0 @ X3) = $true) => ? [X5 : $i,X4 : $i] : (((s__instance @ X4 @ s__OfficeBuilding) = $true) & ((s__destination @ X5 @ X4 @ X3) = $true) & ((s__instance @ X5 @ s__Driving) = $true) & ((s__agent @ X5 @ s__Bill @ X3) = $true))))),
  inference(fool_elimination,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X0 : w,X1 : $i,X2 : $i,X3 : w] : (((s__instance @ X1 @ s__Giving) & (s__agent @ X1 @ s__Bill @ X0) & (s__instance @ X2 @ s__Tax) & (s__patient @ X1 @ X2 @ X0)) => ((accreln1 @ s__Obligation @ X0 @ X3) => ? [X4 : $i,X5 : $i] : ((s__destination @ X5 @ X4 @ X3) & (s__instance @ X5 @ s__Driving) & (s__instance @ X4 @ s__OfficeBuilding) & (s__agent @ X5 @ s__Bill @ X3))))),
  inference(rectify,[status(thm)],[f37])).
thf(f37,axiom,(
  ! [X0 : w,X4 : $i,X5 : $i,X1 : w] : (((s__instance @ X4 @ s__Giving) & (s__agent @ X4 @ s__Bill @ X0) & (s__instance @ X5 @ s__Tax) & (s__patient @ X4 @ X5 @ X0)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X7 : $i,X6 : $i] : ((s__destination @ X6 @ X7 @ X1) & (s__instance @ X6 @ s__Driving) & (s__instance @ X7 @ s__OfficeBuilding) & (s__agent @ X6 @ s__Bill @ X1))))),
  file('/tmp/w7nKvtM0BM/SOT_6fVAoP',obligation2_drive_to_office)).
thf(f440,plain,(
  ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f439,f379])).
thf(f379,plain,(
  ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) = $true)),
  inference(subsumption_resolution,[status(thm)],[f378,f366])).
thf(f378,plain,(
  ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) = $true)),
  inference(subsumption_resolution,[status(thm)],[f377,f360])).
thf(f377,plain,(
  ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) = $true)),
  inference(subsumption_resolution,[status(thm)],[f376,f363])).
thf(f376,plain,(
  ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f375])).
thf(f375,plain,(
  ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true) | ($true != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) = $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true)),
  inference(superposition,[status(thm)],[f210,f369])).
thf(f210,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__patient @ X2 @ X1 @ X0) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__agent @ X2 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__OfficeBuilding) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f204])).
thf(f204,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__agent @ X2 @ s__Bill @ X0) != $true) | ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__OfficeBuilding) = $true) | ((s__patient @ X2 @ X1 @ X0) != $true) | ($true != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__instance @ X1 @ s__Tax) != $true)) )),
  inference(superposition,[status(thm)],[f106,f84])).
thf(f106,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ X3 @ X0) != $true) | ((s__instance @ (sK3 @ X0) @ s__OfficeBuilding) = $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__patient @ X2 @ X1 @ X3) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f77])).
thf(f439,plain,(
  ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true)),
  inference(subsumption_resolution,[status(thm)],[f437,f384])).
thf(f384,plain,(
  ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f383,f363])).
thf(f383,plain,(
  ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f382,f360])).
thf(f382,plain,(
  ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f381,f366])).
thf(f381,plain,(
  ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) = $true) | ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f380])).
thf(f380,plain,(
  ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) = $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ($true != $true)),
  inference(superposition,[status(thm)],[f263,f369])).
thf(f263,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__patient @ X1 @ X2 @ X0) != $true) | ((s__instance @ X2 @ s__Tax) != $true) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__agent @ (sK2 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)) = $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f254])).
thf(f254,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__instance @ X1 @ s__Giving) != $true) | ((s__agent @ (sK2 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)) = $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__Tax) != $true) | ((s__patient @ X1 @ X2 @ X0) != $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f103,f84])).
thf(f103,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ X3 @ X0) != $true) | ((s__patient @ X2 @ X1 @ X3) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__instance @ X2 @ s__Giving) != $true) | ($true = (s__agent @ (sK2 @ X0) @ s__Bill @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f77])).
thf(f437,plain,(
  ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f436])).
thf(f436,plain,(
  ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true) | ($true != $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true)),
  inference(superposition,[status(thm)],[f291,f389])).
thf(f389,plain,(
  ((s__destination @ (sK2 @ (sK0 @ (sK0 @ cw))) @ (sK3 @ (sK0 @ (sK0 @ cw))) @ (sK0 @ (sK0 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f388,f363])).
thf(f388,plain,(
  ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__destination @ (sK2 @ (sK0 @ (sK0 @ cw))) @ (sK3 @ (sK0 @ (sK0 @ cw))) @ (sK0 @ (sK0 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f387,f366])).
thf(f387,plain,(
  ((s__destination @ (sK2 @ (sK0 @ (sK0 @ cw))) @ (sK3 @ (sK0 @ (sK0 @ cw))) @ (sK0 @ (sK0 @ cw))) = $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true)),
  inference(subsumption_resolution,[status(thm)],[f386,f360])).
thf(f386,plain,(
  ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ((s__destination @ (sK2 @ (sK0 @ (sK0 @ cw))) @ (sK3 @ (sK0 @ (sK0 @ cw))) @ (sK0 @ (sK0 @ cw))) = $true) | ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f385])).
thf(f385,plain,(
  ((s__agent @ (sK6 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)) != $true) | ((s__instance @ (sK7 @ (sK0 @ cw)) @ s__Tax) != $true) | ($true != $true) | ((s__destination @ (sK2 @ (sK0 @ (sK0 @ cw))) @ (sK3 @ (sK0 @ (sK0 @ cw))) @ (sK0 @ (sK0 @ cw))) = $true) | ((s__instance @ (sK6 @ (sK0 @ cw)) @ s__Giving) != $true)),
  inference(superposition,[status(thm)],[f309,f369])).
thf(f309,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__patient @ X1 @ X2 @ X0) != $true) | ((s__instance @ X2 @ s__Tax) != $true) | ($true = (s__destination @ (sK2 @ (sK0 @ X0)) @ (sK3 @ (sK0 @ X0)) @ (sK0 @ X0))) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f302])).
thf(f302,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__patient @ X1 @ X2 @ X0) != $true) | ($true != $true) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__instance @ X2 @ s__Tax) != $true) | ($true = (s__destination @ (sK2 @ (sK0 @ X0)) @ (sK3 @ (sK0 @ X0)) @ (sK0 @ X0)))) )),
  inference(superposition,[status(thm)],[f105,f84])).
thf(f105,plain,(
  ( ! [X2 : $i,X3 : w,X0 : w,X1 : $i] : (((accreln1 @ s__Obligation @ X3 @ X0) != $true) | ((s__patient @ X2 @ X1 @ X3) != $true) | ((s__instance @ X1 @ s__Tax) != $true) | ((s__destination @ (sK2 @ X0) @ (sK3 @ X0) @ X0) = $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X2 @ s__Giving) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f77])).
thf(f291,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__destination @ X1 @ X2 @ X0) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__instance @ X1 @ s__Driving) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__agent @ (sK4 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f284])).
thf(f284,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__destination @ X1 @ X2 @ X0) != $true) | ($true != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__Driving) != $true) | ((s__agent @ (sK4 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)) = $true)) )),
  inference(superposition,[status(thm)],[f115,f135])).
thf(f135,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Permission @ X0 @ (sK0 @ X0)) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f134])).
thf(f134,plain,(
  ( ! [X0 : w] : (($true != $true) | ((accreln1 @ s__Permission @ X0 @ (sK0 @ X0)) = $true)) )),
  inference(superposition,[status(thm)],[f85,f84])).
thf(f85,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ((accreln1 @ s__Permission @ X1 @ X0) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f70])).
thf(f70,plain,(
  ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ((accreln1 @ s__Permission @ X1 @ X0) = $true))),
  inference(ennf_transformation,[status(thm)],[f44])).
thf(f44,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X1 @ X0) = $true) => ((accreln1 @ s__Permission @ X1 @ X0) = $true))),
  inference(fool_elimination,[status(thm)],[f43])).
thf(f43,plain,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X1 @ X0) => (accreln1 @ s__Permission @ X1 @ X0))),
  inference(rectify,[status(thm)],[f3])).
thf(f3,axiom,(
  ! [X1 : w,X0 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  file('/tmp/w7nKvtM0BM/SOT_6fVAoP',bridge_obligation_permission)).
thf(f115,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (((accreln1 @ s__Permission @ X3 @ X1) != $true) | ((s__destination @ X0 @ X2 @ X3) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ((s__agent @ (sK4 @ X1) @ s__Bill @ X1) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f80])).
thf(f80,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ((accreln1 @ s__Permission @ X3 @ X1) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | (((s__agent @ (sK4 @ X1) @ s__Bill @ X1) = $true) & ((s__located @ (sK4 @ X1) @ (sK5 @ X1) @ X1) = $true) & ((s__instance @ (sK4 @ X1) @ s__Keeping) = $true) & ((s__instance @ (sK5 @ X1) @ s__ParkingLot) = $true)) | ((s__destination @ X0 @ X2 @ X3) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f78,f79])).
thf(f79,plain,(
  ! [X1 : w] : (? [X4 : $i,X5 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ((s__located @ X4 @ X5 @ X1) = $true) & ((s__instance @ X4 @ s__Keeping) = $true) & ((s__instance @ X5 @ s__ParkingLot) = $true)) => (((s__agent @ (sK4 @ X1) @ s__Bill @ X1) = $true) & ((s__located @ (sK4 @ X1) @ (sK5 @ X1) @ X1) = $true) & ((s__instance @ (sK4 @ X1) @ s__Keeping) = $true) & ((s__instance @ (sK5 @ X1) @ s__ParkingLot) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f78,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ((accreln1 @ s__Permission @ X3 @ X1) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ? [X4 : $i,X5 : $i] : (((s__agent @ X4 @ s__Bill @ X1) = $true) & ((s__located @ X4 @ X5 @ X1) = $true) & ((s__instance @ X4 @ s__Keeping) = $true) & ((s__instance @ X5 @ s__ParkingLot) = $true)) | ((s__destination @ X0 @ X2 @ X3) != $true))),
  inference(rectify,[status(thm)],[f69])).
thf(f69,plain,(
  ! [X2 : $i,X0 : w,X1 : $i,X3 : w] : (((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X2 @ s__Driving) != $true) | ((accreln1 @ s__Permission @ X3 @ X0) != $true) | ((s__instance @ X1 @ s__OfficeBuilding) != $true) | ? [X5 : $i,X4 : $i] : (((s__agent @ X5 @ s__Bill @ X0) = $true) & ((s__located @ X5 @ X4 @ X0) = $true) & ((s__instance @ X5 @ s__Keeping) = $true) & ((s__instance @ X4 @ s__ParkingLot) = $true)) | ((s__destination @ X2 @ X1 @ X3) != $true))),
  inference(flattening,[status(thm)],[f68])).
thf(f68,plain,(
  ! [X0 : w,X3 : w,X2 : $i,X1 : $i] : ((? [X5 : $i,X4 : $i] : (((s__agent @ X5 @ s__Bill @ X0) = $true) & ((s__located @ X5 @ X4 @ X0) = $true) & ((s__instance @ X5 @ s__Keeping) = $true) & ((s__instance @ X4 @ s__ParkingLot) = $true)) | ((accreln1 @ s__Permission @ X3 @ X0) != $true)) | (((s__instance @ X2 @ s__Driving) != $true) | ((s__agent @ X2 @ s__Bill @ X3) != $true) | ((s__instance @ X1 @ s__OfficeBuilding) != $true) | ((s__destination @ X2 @ X1 @ X3) != $true)))),
  inference(ennf_transformation,[status(thm)],[f54])).
thf(f54,plain,(
  ! [X0 : w,X3 : w,X2 : $i,X1 : $i] : ((((s__instance @ X2 @ s__Driving) = $true) & ((s__agent @ X2 @ s__Bill @ X3) = $true) & ((s__instance @ X1 @ s__OfficeBuilding) = $true) & ((s__destination @ X2 @ X1 @ X3) = $true)) => (((accreln1 @ s__Permission @ X3 @ X0) = $true) => ? [X5 : $i,X4 : $i] : (((s__agent @ X5 @ s__Bill @ X0) = $true) & ((s__located @ X5 @ X4 @ X0) = $true) & ((s__instance @ X5 @ s__Keeping) = $true) & ((s__instance @ X4 @ s__ParkingLot) = $true))))),
  inference(fool_elimination,[status(thm)],[f53])).
thf(f53,plain,(
  ! [X0 : w,X1 : $i,X2 : $i,X3 : w] : (((s__destination @ X2 @ X1 @ X3) & (s__instance @ X1 @ s__OfficeBuilding) & (s__instance @ X2 @ s__Driving) & (s__agent @ X2 @ s__Bill @ X3)) => ((accreln1 @ s__Permission @ X3 @ X0) => ? [X4 : $i,X5 : $i] : ((s__located @ X5 @ X4 @ X0) & (s__instance @ X5 @ s__Keeping) & (s__instance @ X4 @ s__ParkingLot) & (s__agent @ X5 @ s__Bill @ X0))))),
  inference(rectify,[status(thm)],[f38])).
thf(f38,axiom,(
  ! [X1 : w,X7 : $i,X6 : $i,X0 : w] : (((s__destination @ X6 @ X7 @ X0) & (s__instance @ X7 @ s__OfficeBuilding) & (s__instance @ X6 @ s__Driving) & (s__agent @ X6 @ s__Bill @ X0)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X9 : $i,X8 : $i] : ((s__located @ X8 @ X9 @ X1) & (s__instance @ X8 @ s__Keeping) & (s__instance @ X9 @ s__ParkingLot) & (s__agent @ X8 @ s__Bill @ X1))))),
  file('/tmp/w7nKvtM0BM/SOT_6fVAoP',permission_park_in_lot)).
thf(f813,plain,(
  ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) != $true)),
  inference(subsumption_resolution,[status(thm)],[f812,f434])).
thf(f434,plain,(
  ((s__instance @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Keeping) = $true)),
  inference(subsumption_resolution,[status(thm)],[f433,f379])).
thf(f433,plain,(
  ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__instance @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Keeping) = $true)),
  inference(subsumption_resolution,[status(thm)],[f432,f384])).
thf(f432,plain,(
  ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__instance @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Keeping) = $true)),
  inference(subsumption_resolution,[status(thm)],[f431,f374])).
thf(f431,plain,(
  ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true) | ((s__instance @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Keeping) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f429])).
thf(f429,plain,(
  ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__instance @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Keeping) = $true) | ($true != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true)),
  inference(superposition,[status(thm)],[f246,f389])).
thf(f246,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__destination @ X1 @ X2 @ X0) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__instance @ X1 @ s__Driving) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ (sK4 @ (sK0 @ X0)) @ s__Keeping) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f240])).
thf(f240,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__destination @ X1 @ X2 @ X0) != $true) | ($true != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ (sK4 @ (sK0 @ X0)) @ s__Keeping) = $true) | ((s__instance @ X1 @ s__Driving) != $true)) )),
  inference(superposition,[status(thm)],[f113,f135])).
thf(f113,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (((accreln1 @ s__Permission @ X3 @ X1) != $true) | ((s__instance @ (sK4 @ X1) @ s__Keeping) = $true) | ((s__destination @ X0 @ X2 @ X3) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ((s__agent @ X0 @ s__Bill @ X3) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f80])).
thf(f812,plain,(
  ((s__instance @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Keeping) != $true) | ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) != $true)),
  inference(subsumption_resolution,[status(thm)],[f811,f410])).
thf(f410,plain,(
  ((s__instance @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__ParkingLot) = $true)),
  inference(subsumption_resolution,[status(thm)],[f409,f379])).
thf(f409,plain,(
  ((s__instance @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__ParkingLot) = $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true)),
  inference(subsumption_resolution,[status(thm)],[f408,f384])).
thf(f408,plain,(
  ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true) | ((s__instance @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__ParkingLot) = $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true)),
  inference(subsumption_resolution,[status(thm)],[f407,f374])).
thf(f407,plain,(
  ((s__instance @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__ParkingLot) = $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f405])).
thf(f405,plain,(
  ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true) | ($true != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__instance @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__ParkingLot) = $true)),
  inference(superposition,[status(thm)],[f219,f389])).
thf(f219,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__destination @ X1 @ X2 @ X0) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__instance @ (sK5 @ (sK0 @ X0)) @ s__ParkingLot) = $true) | ((s__instance @ X1 @ s__Driving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f214])).
thf(f214,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != $true) | ((s__destination @ X1 @ X2 @ X0) != $true) | ((s__instance @ X1 @ s__Driving) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ (sK5 @ (sK0 @ X0)) @ s__ParkingLot) = $true)) )),
  inference(superposition,[status(thm)],[f112,f135])).
thf(f112,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (((accreln1 @ s__Permission @ X3 @ X1) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__destination @ X0 @ X2 @ X3) != $true) | ((s__instance @ (sK5 @ X1) @ s__ParkingLot) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f80])).
thf(f811,plain,(
  ((s__instance @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__ParkingLot) != $true) | ((s__instance @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Keeping) != $true) | ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f810])).
thf(f810,plain,(
  ((s__agent @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) != $true) | ($true != $true) | ((s__instance @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__ParkingLot) != $true) | ((s__instance @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ s__Keeping) != $true)),
  inference(superposition,[status(thm)],[f181,f448])).
thf(f448,plain,(
  ((s__located @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f447,f379])).
thf(f447,plain,(
  ((s__located @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true)),
  inference(subsumption_resolution,[status(thm)],[f446,f374])).
thf(f446,plain,(
  ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__located @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f445,f384])).
thf(f445,plain,(
  ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true) | ((s__located @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f443])).
thf(f443,plain,(
  ($true != $true) | ((s__instance @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Driving) != $true) | ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__OfficeBuilding) != $true) | ((s__located @ (sK4 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK5 @ (sK0 @ (sK0 @ (sK0 @ cw)))) @ (sK0 @ (sK0 @ (sK0 @ cw)))) = $true) | ((s__agent @ (sK2 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))) != $true)),
  inference(superposition,[status(thm)],[f344,f389])).
thf(f344,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__destination @ X1 @ X2 @ X0) != $true) | ((s__instance @ X1 @ s__Driving) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__located @ (sK4 @ (sK0 @ X0)) @ (sK5 @ (sK0 @ X0)) @ (sK0 @ X0)) = $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f336])).
thf(f336,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__located @ (sK4 @ (sK0 @ X0)) @ (sK5 @ (sK0 @ X0)) @ (sK0 @ X0)) = $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__destination @ X1 @ X2 @ X0) != $true) | ($true != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__instance @ X1 @ s__Driving) != $true)) )),
  inference(superposition,[status(thm)],[f114,f135])).
thf(f114,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (((accreln1 @ s__Permission @ X3 @ X1) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ((s__located @ (sK4 @ X1) @ (sK5 @ X1) @ X1) = $true) | ((s__agent @ X0 @ s__Bill @ X3) != $true) | ((s__destination @ X0 @ X2 @ X3) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f80])).
thf(f181,plain,(
  ( ! [X0 : $i,X1 : $i] : (((s__located @ X0 @ X1 @ (sK0 @ (sK0 @ (sK0 @ cw)))) != $true) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ((s__instance @ X0 @ s__Keeping) != $true) | ((s__agent @ X0 @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f180])).
thf(f180,plain,(
  ( ! [X0 : $i,X1 : $i] : (((s__agent @ X0 @ s__Bill @ (sK0 @ (sK0 @ (sK0 @ cw)))) != $true) | ((s__instance @ X0 @ s__Keeping) != $true) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ((s__located @ X0 @ X1 @ (sK0 @ (sK0 @ (sK0 @ cw)))) != $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f86,f171])).
thf(f171,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Permission @ X0 @ (sK0 @ (sK0 @ (sK0 @ X0)))) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f162])).
thf(f162,plain,(
  ( ! [X0 : w] : (($true != $true) | ((accreln1 @ s__Permission @ X0 @ (sK0 @ (sK0 @ (sK0 @ X0)))) = $true)) )),
  inference(superposition,[status(thm)],[f85,f150])).
thf(f150,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ X0 @ (sK0 @ (sK0 @ (sK0 @ X0)))) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f145])).
thf(f145,plain,(
  ( ! [X0 : w] : (($true != $true) | ((accreln1 @ s__Obligation @ X0 @ (sK0 @ (sK0 @ (sK0 @ X0)))) = $true)) )),
  inference(superposition,[status(thm)],[f137,f141])).
thf(f141,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Obligation @ X0 @ (sK0 @ (sK0 @ X0))) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f140])).
thf(f140,plain,(
  ( ! [X0 : w] : (($true != $true) | ((accreln1 @ s__Obligation @ X0 @ (sK0 @ (sK0 @ X0))) = $true)) )),
  inference(superposition,[status(thm)],[f137,f84])).
thf(f137,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X1 @ X0) != $true) | ((accreln1 @ s__Obligation @ X1 @ (sK0 @ X0)) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f136])).
thf(f136,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X1 @ (sK0 @ X0)) = $true) | ($true != $true) | ((accreln1 @ s__Obligation @ X1 @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f133,f84])).
thf(f133,plain,(
  ( ! [X2 : w,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X2 @ X1) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) = $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f62])).
thf(f62,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) | ((accreln1 @ s__Obligation @ X2 @ X1) != $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true))),
  inference(flattening,[status(thm)],[f61])).
thf(f61,plain,(
  ! [X2 : w,X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) | (((accreln1 @ s__Obligation @ X2 @ X1) != $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)))),
  inference(ennf_transformation,[status(thm)],[f58])).
thf(f58,plain,(
  ! [X2 : w,X1 : w,X0 : w] : ((((accreln1 @ s__Obligation @ X2 @ X1) = $true) & ((accreln1 @ s__Obligation @ X0 @ X2) = $true)) => ((accreln1 @ s__Obligation @ X0 @ X1) = $true))),
  inference(fool_elimination,[status(thm)],[f57])).
thf(f57,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X0 @ X2) & (accreln1 @ s__Obligation @ X2 @ X1)) => (accreln1 @ s__Obligation @ X0 @ X1))),
  inference(rectify,[status(thm)],[f2])).
thf(f2,axiom,(
  ! [X0 : w,X2 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) & (accreln1 @ s__Obligation @ X1 @ X2)) => (accreln1 @ s__Obligation @ X0 @ X2))),
  file('/tmp/w7nKvtM0BM/SOT_6fVAoP',transitive_obligation)).
thf(f86,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((accreln1 @ s__Permission @ cw @ X0) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__located @ X1 @ X2 @ X0) != $true) | ((s__instance @ X1 @ s__Keeping) != $true) | ((s__instance @ X2 @ s__ParkingLot) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f67])).
thf(f67,plain,(
  ! [X0 : w] : (! [X1 : $i,X2 : $i] : (((s__instance @ X2 @ s__ParkingLot) != $true) | ((s__instance @ X1 @ s__Keeping) != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__located @ X1 @ X2 @ X0) != $true)) | ((accreln1 @ s__Permission @ cw @ X0) != $true))),
  inference(ennf_transformation,[status(thm)],[f56])).
thf(f56,plain,(
  ~? [X0 : w] : (((accreln1 @ s__Permission @ cw @ X0) = $true) & ? [X2 : $i,X1 : $i] : (((s__agent @ X1 @ s__Bill @ X0) = $true) & ((s__instance @ X1 @ s__Keeping) = $true) & ((s__located @ X1 @ X2 @ X0) = $true) & ((s__instance @ X2 @ s__ParkingLot) = $true)))),
  inference(fool_elimination,[status(thm)],[f55])).
thf(f55,plain,(
  ~? [X0 : w] : ((accreln1 @ s__Permission @ cw @ X0) & ? [X1 : $i,X2 : $i] : ((s__agent @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Keeping) & (s__instance @ X2 @ s__ParkingLot) & (s__located @ X1 @ X2 @ X0)))),
  inference(rectify,[status(thm)],[f41])).
thf(f41,negated_conjecture,(
  ~? [X10 : w] : ((accreln1 @ s__Permission @ cw @ X10) & ? [X8 : $i,X9 : $i] : ((s__agent @ X8 @ s__Bill @ X10) & (s__instance @ X8 @ s__Keeping) & (s__instance @ X9 @ s__ParkingLot) & (s__located @ X8 @ X9 @ X10)))),
  inference(negated_conjecture,[status(cth)],[f40])).
thf(f40,conjecture,(
  ? [X10 : w] : ((accreln1 @ s__Permission @ cw @ X10) & ? [X8 : $i,X9 : $i] : ((s__agent @ X8 @ s__Bill @ X10) & (s__instance @ X8 @ s__Keeping) & (s__instance @ X9 @ s__ParkingLot) & (s__located @ X8 @ X9 @ X10)))),
  file('/tmp/w7nKvtM0BM/SOT_6fVAoP',conjecture_is_permitted)).
% SZS output end Proof for SOT_6fVAoP
% (2148412)------------------------------
% (2148412)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2148412)Termination reason: Refutation

% (2148412)Memory used [KB]: 6524
% (2148412)Time elapsed: 0.107 s
% (2148412)------------------------------
% (2148412)------------------------------
% (2148405)Success in time 0.119 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_6fVAoP - Vampire---5.0.1 says Theorem - CPU = 0.03 WC = 0.16 
% OUTPUT: SOT_6fVAoP - Vampire---5.0.1 says Refutation - CPU = 0.03 WC = 0.17 

