% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 2b_Permission_chained_existConj
% SZS output start Proof for 2b_Permission_chained_existConj
1. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Permission @ X0 @ X1) [input(axiom)]
2. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Permission @ X1 @ X2) & (accreln1 @ s__Permission @ X0 @ X1)) => (accreln1 @ s__Permission @ X0 @ X2)) [input(axiom)]
19. ! [X0 : w,X1 : w] : ! [X2] : (((s__possesses @ s__Bill @ X2 @ X0) & (s__instance @ X2 @ s__DriversLicense)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X3,X4] : ((s__located @ X3 @ X4 @ X1) & (s__agent @ X3 @ s__Bill @ X1) & (s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking)))) [input(axiom)]
20. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__located @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__ParkingLot) & (s__instance @ X2 @ s__Parking)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X4] : ((s__destination @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Reimbursement)))) [input(axiom)]
21. ? [X0] : ((s__possesses @ s__Bill @ X0 @ cw) & (s__instance @ X0 @ s__DriversLicense)) [input(axiom)]
22. ? [X0 : w] : (? [X1] : ((s__destination @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Reimbursement)) & (accreln1 @ s__Permission @ cw @ X0)) [input(conjecture)]
23. ~? [X0 : w] : (? [X1] : ((s__destination @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Reimbursement)) & (accreln1 @ s__Permission @ cw @ X0)) [negated conjecture 22]
24. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Permission @ X0 @ X1) [rectify 1]
25. ! [X0 : w] : ? [X1 : w] : (((accreln1 @ s__Permission @ X0 @ X1)) = $true) [fool elimination 24]
26. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Permission @ X1 @ X2) & (accreln1 @ s__Permission @ X0 @ X1)) => (accreln1 @ s__Permission @ X0 @ X2)) [rectify 2]
27. ! [X0 : w,X1 : w,X2 : w] : (((((accreln1 @ s__Permission @ X1 @ X2)) = $true) & (((accreln1 @ s__Permission @ X0 @ X1)) = $true)) => (((accreln1 @ s__Permission @ X0 @ X2)) = $true)) [fool elimination 26]
30. ! [X0 : w,X1 : w] : ! [X2] : (((s__possesses @ s__Bill @ X2 @ X0) & (s__instance @ X2 @ s__DriversLicense)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X3,X4] : ((s__located @ X3 @ X4 @ X1) & (s__agent @ X3 @ s__Bill @ X1) & (s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking)))) [rectify 19]
31. ! [X0 : w,X1 : w] : ! [X2] : (((((s__possesses @ s__Bill @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__DriversLicense)) = $true)) => ((((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)))) [fool elimination 30]
32. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__located @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__ParkingLot) & (s__instance @ X2 @ s__Parking)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X4] : ((s__destination @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Reimbursement)))) [rectify 20]
33. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__located @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__ParkingLot)) = $true) & (((s__instance @ X2 @ s__Parking)) = $true)) => ((((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)))) [fool elimination 32]
34. ? [X0] : ((s__possesses @ s__Bill @ X0 @ cw) & (s__instance @ X0 @ s__DriversLicense)) [rectify 21]
35. ? [X0] : ((((s__possesses @ s__Bill @ X0 @ cw)) = $true) & (((s__instance @ X0 @ s__DriversLicense)) = $true)) [fool elimination 34]
36. ~? [X0 : w] : (? [X1] : ((s__destination @ X1 @ s__Bill @ X0) & (s__instance @ X1 @ s__Reimbursement)) & (accreln1 @ s__Permission @ cw @ X0)) [rectify 23]
37. ~? [X0 : w] : (? [X1] : ((((s__destination @ X1 @ s__Bill @ X0)) = $true) & (((s__instance @ X1 @ s__Reimbursement)) = $true)) & (((accreln1 @ s__Permission @ cw @ X0)) = $true)) [fool elimination 36]
53. ! [X0 : w,X1 : w,X2] : (((((s__possesses @ s__Bill @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__DriversLicense)) = $true)) => ((((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)))) [flattening 31]
54. ! [X0 : w,X1 : w,X2,X3] : (((((s__located @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__ParkingLot)) = $true) & (((s__instance @ X2 @ s__Parking)) = $true)) => ((((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)))) [flattening 33]
55. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Permission @ X0 @ X2)) = $true) | ((((accreln1 @ s__Permission @ X1 @ X2)) != $true) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true))) [ennf transformation 27]
56. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Permission @ X0 @ X2)) = $true) | (((accreln1 @ s__Permission @ X1 @ X2)) != $true) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true)) [flattening 55]
57. ! [X0 : w,X1 : w,X2] : ((? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true)) | ((((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true))) [ennf transformation 53]
58. ! [X0 : w,X1 : w,X2] : (? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true)) [flattening 57]
59. ! [X0 : w,X1 : w,X2,X3] : ((? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true)) | ((((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true))) [ennf transformation 54]
60. ! [X0 : w,X1 : w,X2,X3] : (? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true)) [flattening 59]
61. ! [X0 : w] : (! [X1] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__instance @ X1 @ s__Reimbursement)) != $true)) | (((accreln1 @ s__Permission @ cw @ X0)) != $true)) [ennf transformation 37]
62. ! [X0 : w] : (? [X1 : w] : (((accreln1 @ s__Permission @ X0 @ X1)) = $true) => ($true = ((accreln1 @ s__Permission @ X0 @ (sK0 @ X0))))) [skolem symbol introduction]
63. ! [X0 : w] : ($true = ((accreln1 @ s__Permission @ X0 @ (sK0 @ X0)))) [skolemisation 25,62]
64. ! [X1 : w] : (? [X3,X4] : ((((s__located @ X3 @ X4 @ X1)) = $true) & (((s__agent @ X3 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__ParkingLot)) = $true) & (((s__instance @ X3 @ s__Parking)) = $true)) => (($true = ((s__located @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__ParkingLot))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Parking))))) [skolem symbol introduction]
65. ! [X0 : w,X1 : w,X2] : ((($true = ((s__located @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__ParkingLot))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Parking)))) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true)) [skolemisation 58,64]
66. ! [X1 : w] : (? [X4] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__instance @ X4 @ s__Reimbursement)) = $true)) => (($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Reimbursement))))) [skolem symbol introduction]
67. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Reimbursement)))) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true)) [skolemisation 60,66]
68. ? [X0] : ((((s__possesses @ s__Bill @ X0 @ cw)) = $true) & (((s__instance @ X0 @ s__DriversLicense)) = $true)) => (($true = ((s__possesses @ s__Bill @ sK4 @ cw))) & ($true = ((s__instance @ sK4 @ s__DriversLicense)))) [skolem symbol introduction]
69. ($true = ((s__possesses @ s__Bill @ sK4 @ cw))) & ($true = ((s__instance @ sK4 @ s__DriversLicense))) [skolemisation 35,68]
70. ($true = ((accreln1 @ s__Permission @ X0 @ (sK0 @ X0)))) [cnf transformation 63]
71. (((accreln1 @ s__Permission @ X1 @ X2)) != $true) | (((accreln1 @ s__Permission @ X0 @ X2)) = $true) | (((accreln1 @ s__Permission @ X0 @ X1)) != $true) [cnf transformation 56]
88. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__Parking))) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true) [cnf transformation 65]
89. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK2 @ X1) @ s__ParkingLot))) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true) [cnf transformation 65]
90. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true) [cnf transformation 65]
91. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__located @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) | (((s__possesses @ s__Bill @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__DriversLicense)) != $true) [cnf transformation 65]
92. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Reimbursement))) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true) [cnf transformation 67]
93. (((accreln1 @ s__Permission @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__located @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__ParkingLot)) != $true) | (((s__instance @ X2 @ s__Parking)) != $true) [cnf transformation 67]
94. ($true = ((s__instance @ sK4 @ s__DriversLicense))) [cnf transformation 69]
95. ($true = ((s__possesses @ s__Bill @ sK4 @ cw))) [cnf transformation 69]
96. (((accreln1 @ s__Permission @ cw @ X0)) != $true) | (((s__instance @ X1 @ s__Reimbursement)) != $true) | (((s__destination @ X1 @ s__Bill @ X0)) != $true) [cnf transformation 61]
210. ($true != $true) | ($true = ((accreln1 @ s__Permission @ X1 @ (sK0 @ X0)))) | ($true != ((accreln1 @ s__Permission @ X1 @ X0))) [superposition 71,70]
218. ($true != ((accreln1 @ s__Permission @ X1 @ X0))) | ($true = ((accreln1 @ s__Permission @ X1 @ (sK0 @ X0)))) [trivial inequality removal 210]
225. ($true != $true) | ($true = ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Parking))) | ($true != ((s__possesses @ s__Bill @ X1 @ X0))) | ($true != ((s__instance @ X1 @ s__DriversLicense))) [superposition 88,70]
233. ($true != ((s__possesses @ s__Bill @ X1 @ X0))) | ($true = ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Parking))) | ($true != ((s__instance @ X1 @ s__DriversLicense))) [trivial inequality removal 225]
244. ($true != $true) | ($true = ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__ParkingLot))) | ($true != ((s__possesses @ s__Bill @ X1 @ X0))) | ($true != ((s__instance @ X1 @ s__DriversLicense))) [superposition 89,70]
252. ($true != ((s__possesses @ s__Bill @ X1 @ X0))) | ($true = ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__ParkingLot))) | ($true != ((s__instance @ X1 @ s__DriversLicense))) [trivial inequality removal 244]
600. ($true != $true) | ($true = ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__possesses @ s__Bill @ X1 @ X0))) | ($true != ((s__instance @ X1 @ s__DriversLicense))) [superposition 90,70]
617. ($true != ((s__possesses @ s__Bill @ X1 @ X0))) | ($true = ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ X1 @ s__DriversLicense))) [trivial inequality removal 600]
668. ($true != $true) | ($true = ((s__located @ (sK1 @ (sK0 @ X0)) @ (sK2 @ (sK0 @ X0)) @ (sK0 @ X0)))) | ($true != ((s__possesses @ s__Bill @ X1 @ X0))) | ($true != ((s__instance @ X1 @ s__DriversLicense))) [superposition 91,70]
685. ($true != ((s__possesses @ s__Bill @ X1 @ X0))) | ($true = ((s__located @ (sK1 @ (sK0 @ X0)) @ (sK2 @ (sK0 @ X0)) @ (sK0 @ X0)))) | ($true != ((s__instance @ X1 @ s__DriversLicense))) [trivial inequality removal 668]
1051. ($true != $true) | ($true = ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__Reimbursement))) | ($true != ((s__located @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__ParkingLot))) | ($true != ((s__instance @ X1 @ s__Parking))) [superposition 92,70]
1074. ($true != ((s__located @ X1 @ X2 @ X0))) | ($true = ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__Reimbursement))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__ParkingLot))) | ($true != ((s__instance @ X1 @ s__Parking))) [trivial inequality removal 1051]
1134. ($true != $true) | ($true = ((s__destination @ (sK3 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__located @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__ParkingLot))) | ($true != ((s__instance @ X1 @ s__Parking))) [superposition 93,70]
1157. ($true != ((s__located @ X1 @ X2 @ X0))) | ($true = ((s__destination @ (sK3 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__ParkingLot))) | ($true != ((s__instance @ X1 @ s__Parking))) [trivial inequality removal 1134]
1462. ($true != $true) | ($true = ((accreln1 @ s__Permission @ X0 @ (sK0 @ (sK0 @ X0))))) [superposition 218,70]
1489. ($true = ((accreln1 @ s__Permission @ X0 @ (sK0 @ (sK0 @ X0))))) [trivial inequality removal 1462]
1663. ($true != $true) | ($true = ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) [superposition 233,95]
1692. ($true = ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) [trivial inequality removal 1663]
1703. ($true = ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [forward subsumption resolution 1692,94]
1747. ($true != $true) | ($true = ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) [superposition 252,95]
1776. ($true = ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) [trivial inequality removal 1747]
1787. ($true = ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) [forward subsumption resolution 1776,94]
1904. ($true != $true) | ($true = ((s__agent @ (sK1 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) [superposition 617,95]
1933. ($true = ((s__agent @ (sK1 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) [trivial inequality removal 1904]
1944. ($true = ((s__agent @ (sK1 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward subsumption resolution 1933,94]
2382. ($true != $true) | ($true = ((s__located @ (sK1 @ (sK0 @ cw)) @ (sK2 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) [superposition 685,95]
2411. ($true = ((s__located @ (sK1 @ (sK0 @ cw)) @ (sK2 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__instance @ sK4 @ s__DriversLicense))) [trivial inequality removal 2382]
2423. ($true = ((s__located @ (sK1 @ (sK0 @ cw)) @ (sK2 @ (sK0 @ cw)) @ (sK0 @ cw)))) [forward subsumption resolution 2411,94]
8542. ($true != $true) | ($true != ((s__instance @ X0 @ s__Reimbursement))) | ($true != ((s__destination @ X0 @ s__Bill @ (sK0 @ (sK0 @ cw))))) [superposition 96,1489]
8561. ($true != ((s__destination @ X0 @ s__Bill @ (sK0 @ (sK0 @ cw))))) | ($true != ((s__instance @ X0 @ s__Reimbursement))) [trivial inequality removal 8542]
10002. ($true != $true) | ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))))) | ($true != ((s__agent @ (sK1 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [superposition 1157,2423]
10003. ($true != $true) | ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Reimbursement))) | ($true != ((s__agent @ (sK1 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [superposition 1074,2423]
10004. ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Reimbursement))) | ($true != ((s__agent @ (sK1 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [trivial inequality removal 10003]
10005. ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))))) | ($true != ((s__agent @ (sK1 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [trivial inequality removal 10002]
10017. ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Reimbursement))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [forward subsumption resolution 10004,1944]
10018. ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))))) | ($true != ((s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) | ($true != ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [forward subsumption resolution 10005,1944]
10047. ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Reimbursement))) | ($true != ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [forward subsumption resolution 10017,1787]
10048. ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))))) | ($true != ((s__instance @ (sK1 @ (sK0 @ cw)) @ s__Parking))) [forward subsumption resolution 10018,1787]
10063. ($true = ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Reimbursement))) [forward subsumption resolution 10047,1703]
10064. ($true = ((s__destination @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Bill @ (sK0 @ (sK0 @ cw))))) [forward subsumption resolution 10048,1703]
54328. ($true != $true) | ($true != ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Reimbursement))) [superposition 8561,10064]
54343. ($true != ((s__instance @ (sK3 @ (sK0 @ (sK0 @ cw))) @ s__Reimbursement))) [trivial inequality removal 54328]
54382. $false [forward subsumption resolution 54343,10063]
% SZS output end Proof for 2b_Permission_chained_existConj
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 15.126 s
% Peak memory usage: 27 MB
% Instructions burned: 117390 (million)
% ------------------------------
% ------------------------------
