 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/vPpZ9gEjRi/SOT_5gjQCS
% (3736539)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_5gjQCS for (1ds/183Mi)
% (3736542)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_5gjQCS for (1ds/27Mi)
% (3736546)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_5gjQCS for (1ds/18Mi)
% (3736543)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_5gjQCS for (1ds/2Mi)
% (3736541)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_5gjQCS for (1ds/4Mi)
% (3736545)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_5gjQCS for (1ds/275Mi)
% (3736544)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_5gjQCS for (1ds/2Mi)
% (3736539)First to succeed.
% (3736546)Also succeeded, but the first one will report.
% (3736543)Also succeeded, but the first one will report.
% (3736544)Refutation not found, incomplete strategy
% (3736544)------------------------------
% (3736544)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (3736544)Termination reason: Refutation not found, incomplete strategy


% (3736544)Memory used [KB]: 5500
% (3736544)Time elapsed: 0.005 s
% (3736544)------------------------------
% (3736544)------------------------------
% (3736539)Refutation found. Thanks to Tanya!
% SZS status ContradictoryAxioms for SOT_5gjQCS
% SZS output start Proof for SOT_5gjQCS
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m > w > w > $o).
thf(func_def_5, type, accreln2: m > $i > w > w > $o).
thf(func_def_6, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_8, type, s__instance: $i > $i > $o).
thf(func_def_9, type, s__agent: $i > $i > w > $o).
thf(func_def_10, type, s__patient: $i > $i > w > $o).
thf(func_def_11, type, s__destination: $i > $i > w > $o).
thf(func_def_12, type, s__located: $i > $i > w > $o).
thf(func_def_13, type, s__possesses: $i > $i > w > $o).
thf(func_def_23, type, sK0: w > $i).
thf(func_def_24, type, sK1: w > $i).
thf(func_def_25, type, sK2: w > $i).
thf(func_def_26, type, sK3: w > w).
thf(f80,plain,(
  $false),
  inference(trivial_inequality_removal,[status(thm)],[f79])).
thf(f79,plain,(
  ($true != $true)),
  inference(superposition,[status(thm)],[f78,f66])).
thf(f66,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Permission @ X0 @ (sK3 @ X0)))) )),
  inference(cnf_transformation,[status(thm)],[f43])).
thf(f43,plain,(
  ! [X0 : w] : ($true = (accreln1 @ s__Permission @ X0 @ (sK3 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f28,f42])).
thf(f42,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Permission @ X0 @ X1) = $true) => ($true = (accreln1 @ s__Permission @ X0 @ (sK3 @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f28,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Permission @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f27])).
thf(f27,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Permission @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Permission @ X0 @ X1)),
  file('/tmp/vPpZ9gEjRi/SOT_5gjQCS',serial_permission)).
thf(f78,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Permission @ cw @ X0) != $true)) )),
  inference(subsumption_resolution,[status(thm)],[f77,f75])).
thf(f75,plain,(
  ( ! [X0 : w] : (($true = (s__instance @ (sK0 @ X0) @ s__DriversLicense))) )),
  inference(trivial_inequality_removal,[status(thm)],[f74])).
thf(f74,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true = (s__instance @ (sK0 @ X0) @ s__DriversLicense))) )),
  inference(superposition,[status(thm)],[f60,f66])).
thf(f60,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Permission @ X1 @ X0) != $true) | ($true = (s__instance @ (sK0 @ X1) @ s__DriversLicense))) )),
  inference(cnf_transformation,[status(thm)],[f40])).
thf(f40,plain,(
  ! [X0 : w,X1 : w] : ((((($true = (s__instance @ (sK0 @ X1) @ s__DriversLicense)) & ($true = (s__possesses @ s__Bill @ (sK0 @ X1) @ X1))) & (((s__located @ (sK1 @ X0) @ (sK2 @ X0) @ X0) = $true) & ($true = (s__agent @ (sK1 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK1 @ X0) @ s__Parking)) & ($true = (s__instance @ (sK2 @ X0) @ s__ParkingLot)))) | ((accreln1 @ s__Permission @ X1 @ X0) != $true)) & (((accreln1 @ s__Permission @ X1 @ X0) = $true) | ! [X5 : $i] : (($true != (s__instance @ X5 @ s__DriversLicense)) | ($true != (s__possesses @ s__Bill @ X5 @ X1))) | ! [X6 : $i,X7 : $i] : (($true != (s__located @ X6 @ X7 @ X0)) | ((s__agent @ X6 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X6 @ s__Parking)) | ((s__instance @ X7 @ s__ParkingLot) != $true))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f37,f39,f38])).
thf(f38,plain,(
  ! [X1 : w] : (? [X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) = $true) & ($true = (s__possesses @ s__Bill @ X2 @ X1))) => (($true = (s__instance @ (sK0 @ X1) @ s__DriversLicense)) & ($true = (s__possesses @ s__Bill @ (sK0 @ X1) @ X1))))),
  introduced(definition,[],[choice_axiom])).
thf(f39,plain,(
  ! [X0 : w] : (? [X3 : $i,X4 : $i] : (($true = (s__located @ X3 @ X4 @ X0)) & ((s__agent @ X3 @ s__Bill @ X0) = $true) & ((s__instance @ X3 @ s__Parking) = $true) & ((s__instance @ X4 @ s__ParkingLot) = $true)) => (((s__located @ (sK1 @ X0) @ (sK2 @ X0) @ X0) = $true) & ($true = (s__agent @ (sK1 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK1 @ X0) @ s__Parking)) & ($true = (s__instance @ (sK2 @ X0) @ s__ParkingLot))))),
  introduced(definition,[],[choice_axiom])).
thf(f37,plain,(
  ! [X0 : w,X1 : w] : (((? [X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) = $true) & ($true = (s__possesses @ s__Bill @ X2 @ X1))) & ? [X3 : $i,X4 : $i] : (($true = (s__located @ X3 @ X4 @ X0)) & ((s__agent @ X3 @ s__Bill @ X0) = $true) & ((s__instance @ X3 @ s__Parking) = $true) & ((s__instance @ X4 @ s__ParkingLot) = $true))) | ((accreln1 @ s__Permission @ X1 @ X0) != $true)) & (((accreln1 @ s__Permission @ X1 @ X0) = $true) | ! [X5 : $i] : (($true != (s__instance @ X5 @ s__DriversLicense)) | ($true != (s__possesses @ s__Bill @ X5 @ X1))) | ! [X6 : $i,X7 : $i] : (($true != (s__located @ X6 @ X7 @ X0)) | ((s__agent @ X6 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X6 @ s__Parking)) | ((s__instance @ X7 @ s__ParkingLot) != $true))))),
  inference(rectify,[status(thm)],[f36])).
thf(f36,plain,(
  ! [X0 : w,X1 : w] : (((? [X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) = $true) & ($true = (s__possesses @ s__Bill @ X2 @ X1))) & ? [X4 : $i,X3 : $i] : (((s__located @ X4 @ X3 @ X0) = $true) & ($true = (s__agent @ X4 @ s__Bill @ X0)) & ((s__instance @ X4 @ s__Parking) = $true) & ($true = (s__instance @ X3 @ s__ParkingLot)))) | ((accreln1 @ s__Permission @ X1 @ X0) != $true)) & (((accreln1 @ s__Permission @ X1 @ X0) = $true) | ! [X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) != $true) | ($true != (s__possesses @ s__Bill @ X2 @ X1))) | ! [X4 : $i,X3 : $i] : (((s__located @ X4 @ X3 @ X0) != $true) | ($true != (s__agent @ X4 @ s__Bill @ X0)) | ((s__instance @ X4 @ s__Parking) != $true) | ($true != (s__instance @ X3 @ s__ParkingLot)))))),
  inference(flattening,[status(thm)],[f35])).
thf(f35,plain,(
  ! [X0 : w,X1 : w] : (((? [X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) = $true) & ($true = (s__possesses @ s__Bill @ X2 @ X1))) & ? [X4 : $i,X3 : $i] : (((s__located @ X4 @ X3 @ X0) = $true) & ($true = (s__agent @ X4 @ s__Bill @ X0)) & ((s__instance @ X4 @ s__Parking) = $true) & ($true = (s__instance @ X3 @ s__ParkingLot)))) | ((accreln1 @ s__Permission @ X1 @ X0) != $true)) & (((accreln1 @ s__Permission @ X1 @ X0) = $true) | (! [X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) != $true) | ($true != (s__possesses @ s__Bill @ X2 @ X1))) | ! [X4 : $i,X3 : $i] : (((s__located @ X4 @ X3 @ X0) != $true) | ($true != (s__agent @ X4 @ s__Bill @ X0)) | ((s__instance @ X4 @ s__Parking) != $true) | ($true != (s__instance @ X3 @ s__ParkingLot))))))),
  inference(nnf_transformation,[status(thm)],[f24])).
thf(f24,plain,(
  ! [X0 : w,X1 : w] : ((? [X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) = $true) & ($true = (s__possesses @ s__Bill @ X2 @ X1))) & ? [X4 : $i,X3 : $i] : (((s__located @ X4 @ X3 @ X0) = $true) & ($true = (s__agent @ X4 @ s__Bill @ X0)) & ((s__instance @ X4 @ s__Parking) = $true) & ($true = (s__instance @ X3 @ s__ParkingLot)))) <=> ((accreln1 @ s__Permission @ X1 @ X0) = $true))),
  inference(fool_elimination,[status(thm)],[f23])).
thf(f23,plain,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Permission @ X1 @ X0) <=> (? [X2 : $i] : ((s__possesses @ s__Bill @ X2 @ X1) & (s__instance @ X2 @ s__DriversLicense)) & ? [X3 : $i,X4 : $i] : ((s__located @ X4 @ X3 @ X0) & (s__agent @ X4 @ s__Bill @ X0) & (s__instance @ X3 @ s__ParkingLot) & (s__instance @ X4 @ s__Parking))))),
  inference(rectify,[status(thm)],[f18])).
thf(f18,axiom,(
  ! [X1 : w,X0 : w] : ((accreln1 @ s__Permission @ X0 @ X1) <=> (? [X2 : $i] : ((s__possesses @ s__Bill @ X2 @ X0) & (s__instance @ X2 @ s__DriversLicense)) & ? [X4 : $i,X3 : $i] : ((s__located @ X3 @ X4 @ X1) & (s__agent @ X3 @ s__Bill @ X1) & (s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking))))),
  file('/tmp/vPpZ9gEjRi/SOT_5gjQCS',permission1_to_park)).
thf(f77,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Permission @ cw @ X0) != $true) | ($true != (s__instance @ (sK0 @ cw) @ s__DriversLicense))) )),
  inference(trivial_inequality_removal,[status(thm)],[f76])).
thf(f76,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK0 @ cw) @ s__DriversLicense)) | ((accreln1 @ s__Permission @ cw @ X0) != $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f64,f59])).
thf(f59,plain,(
  ( ! [X0 : w,X1 : w] : (($true = (s__possesses @ s__Bill @ (sK0 @ X1) @ X1)) | ((accreln1 @ s__Permission @ X1 @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f40])).
thf(f64,plain,(
  ( ! [X0 : $i] : (($true != (s__possesses @ s__Bill @ X0 @ cw)) | ($true != (s__instance @ X0 @ s__DriversLicense))) )),
  inference(cnf_transformation,[status(thm)],[f33])).
thf(f33,plain,(
  ! [X0 : $i] : (($true != (s__instance @ X0 @ s__DriversLicense)) | ($true != (s__possesses @ s__Bill @ X0 @ cw)))),
  inference(ennf_transformation,[status(thm)],[f32])).
thf(f32,plain,(
  ~? [X0 : $i] : (($true = (s__instance @ X0 @ s__DriversLicense)) & ($true = (s__possesses @ s__Bill @ X0 @ cw)))),
  inference(fool_elimination,[status(thm)],[f31])).
thf(f31,plain,(
  ~? [X0 : $i] : ((s__instance @ X0 @ s__DriversLicense) & (s__possesses @ s__Bill @ X0 @ cw))),
  inference(rectify,[status(thm)],[f19])).
thf(f19,axiom,(
  ~? [X2 : $i] : ((s__instance @ X2 @ s__DriversLicense) & (s__possesses @ s__Bill @ X2 @ cw))),
  file('/tmp/vPpZ9gEjRi/SOT_5gjQCS',fact_no_license)).
% SZS output end Proof for SOT_5gjQCS
% (3736539)------------------------------
% (3736539)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (3736539)Termination reason: Refutation

% (3736539)Memory used [KB]: 5628
% (3736539)Time elapsed: 0.006 s
% (3736539)------------------------------
% (3736539)------------------------------
% (3736534)Success in time 0.019 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_5gjQCS - Vampire---5.0.1 says ContradictoryAxioms - CPU = 0.01 WC = 0.06 
% OUTPUT: SOT_5gjQCS - Vampire---5.0.1 says Refutation - CPU = 0.01 WC = 0.06 

