% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 1g_Obligation_UO1_existConj
% SZS output start Proof for 1g_Obligation_UO1_existConj
1. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1) [input(axiom)]
19. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [input(axiom)]
20. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [input(axiom)]
21. ? [X0 : w] : (? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving)) & (accreln1 @ s__Obligation @ cw @ X0)) [input(conjecture)]
22. ~? [X0 : w] : (? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving)) & (accreln1 @ s__Obligation @ cw @ X0)) [negated conjecture 21]
23. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1) [rectify 1]
24. ! [X0 : w] : ? [X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) [fool elimination 23]
29. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [rectify 19]
30. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [fool elimination 29]
31. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [rectify 20]
32. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) [fool elimination 31]
33. ~? [X0 : w] : (? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving)) & (accreln1 @ s__Obligation @ cw @ X0)) [rectify 22]
34. ~? [X0 : w] : (? [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) = $true) & (((s__patient @ X1 @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__PayCheck)) = $true) & (((s__instance @ X1 @ s__Giving)) = $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) [fool elimination 33]
50. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [flattening 30]
52. ! [X0 : w,X1 : w,X2,X3] : ((? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true))) [ennf transformation 50]
53. ! [X0 : w,X1 : w,X2,X3] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [flattening 52]
54. ! [X0 : w] : (! [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) | (((accreln1 @ s__Obligation @ cw @ X0)) != $true)) [ennf transformation 34]
55. ! [X0 : w] : (? [X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0))))) [skolem symbol introduction]
56. ! [X0 : w] : ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)))) [skolemisation 24,55]
59. ! [X1 : w] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) => (($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) & ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))))) [skolem symbol introduction]
60. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) & ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Giving)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [skolemisation 53,59]
61. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) => (($true = ((s__destination @ sK5 @ sK6 @ cw))) & ($true = ((s__agent @ sK5 @ s__Bill @ cw))) & ($true = ((s__instance @ sK6 @ s__Moon))) & ($true = ((s__instance @ sK5 @ s__Flying)))) [skolem symbol introduction]
62. ($true = ((s__destination @ sK5 @ sK6 @ cw))) & ($true = ((s__agent @ sK5 @ s__Bill @ cw))) & ($true = ((s__instance @ sK6 @ s__Moon))) & ($true = ((s__instance @ sK5 @ s__Flying))) [skolemisation 32,61]
63. ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)))) [cnf transformation 56]
84. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 60]
85. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 60]
86. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 60]
87. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 60]
88. ($true = ((s__instance @ sK5 @ s__Flying))) [cnf transformation 62]
89. ($true = ((s__instance @ sK6 @ s__Moon))) [cnf transformation 62]
90. ($true = ((s__agent @ sK5 @ s__Bill @ cw))) [cnf transformation 62]
91. ($true = ((s__destination @ sK5 @ sK6 @ cw))) [cnf transformation 62]
92. (((accreln1 @ s__Obligation @ cw @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | (((s__destination @ X1 @ s__Bill @ X0)) != $true) [cnf transformation 54]
94. ($true = X0) | ($false = X0) [fool exhaustiveness axiom]
95. (X0 = X1) | ($false = X1) | ($false = X0) [superposition 94,94]
98. ($true != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | (((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((accreln1 @ s__Obligation @ cw @ X0)) = $false) [superposition 92,94]
99. (((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((accreln1 @ s__Obligation @ cw @ X0)) = $false) [trivial inequality removal 98]
100. ($true != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | ($true != ((s__patient @ X0 @ X2 @ X1))) | ($false = ((accreln1 @ s__Obligation @ cw @ X1))) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) [superposition 99,94]
101. ($true != ((s__patient @ X0 @ X2 @ X1))) | ($true != ((s__instance @ X0 @ s__Giving))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | ($false = ((accreln1 @ s__Obligation @ cw @ X1))) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) [trivial inequality removal 100]
102. ($true != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($false = ((accreln1 @ s__Obligation @ cw @ X2))) | ($false = ((s__destination @ X0 @ s__Bill @ X2))) | ($false = ((s__patient @ X0 @ X1 @ X2))) [superposition 101,94]
103. ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((accreln1 @ s__Obligation @ cw @ X2))) | ($false = ((s__destination @ X0 @ s__Bill @ X2))) | ($false = ((s__patient @ X0 @ X1 @ X2))) [trivial inequality removal 102]
104. ($true != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | ($false = ((accreln1 @ s__Obligation @ cw @ X2))) | ($false = ((s__destination @ X1 @ s__Bill @ X2))) | ($false = ((s__patient @ X1 @ X0 @ X2))) | ($false = ((s__instance @ X0 @ s__PayCheck))) [superposition 103,94]
105. (((s__instance @ X1 @ s__Giving)) != $true) | ($false = ((accreln1 @ s__Obligation @ cw @ X2))) | ($false = ((s__destination @ X1 @ s__Bill @ X2))) | ($false = ((s__patient @ X1 @ X0 @ X2))) | ($false = ((s__instance @ X0 @ s__PayCheck))) [trivial inequality removal 104]
106. ($true != $true) | ($false = ((accreln1 @ s__Obligation @ cw @ X1))) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) | ($false = ((s__patient @ X0 @ X2 @ X1))) | (((s__instance @ X2 @ s__PayCheck)) = $false) | ($false = ((s__instance @ X0 @ s__Giving))) [superposition 105,94]
107. ($false = ((accreln1 @ s__Obligation @ cw @ X1))) | ($false = ((s__patient @ X0 @ X2 @ X1))) | ($false = ((s__destination @ X0 @ s__Bill @ X1))) | (((s__instance @ X2 @ s__PayCheck)) = $false) | ($false = ((s__instance @ X0 @ s__Giving))) [trivial inequality removal 106]
118. (X0 = X1) | (X1 = X2) | (X0 = X2) | (X0 = X3) | ($false = X3) [superposition 95,95]
218. ($true = $false) | ($false = ((s__patient @ X0 @ X1 @ (sK0 @ cw)))) | ($false = ((s__destination @ X0 @ s__Bill @ (sK0 @ cw)))) | ($false = ((s__instance @ X1 @ s__PayCheck))) | ($false = ((s__instance @ X0 @ s__Giving))) [superposition 107,63]
219. ($true != $true) | ($true != ((s__patient @ X0 @ X1 @ (sK0 @ cw)))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($true != ((s__destination @ X0 @ s__Bill @ (sK0 @ cw)))) [superposition 92,63]
220. ($true != ((s__destination @ X0 @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($true != ((s__patient @ X0 @ X1 @ (sK0 @ cw)))) [trivial inequality removal 219]
221. ($false = ((s__patient @ X0 @ X1 @ (sK0 @ cw)))) | ($false = ((s__destination @ X0 @ s__Bill @ (sK0 @ cw)))) | ($false = ((s__instance @ X1 @ s__PayCheck))) | ($false = ((s__instance @ X0 @ s__Giving))) [trivial inequality removal 218]
300. ($true != $true) | ($true = ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__Giving))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 84,63]
307. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__instance @ (sK3 @ (sK0 @ X0)) @ s__Giving))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 300]
414. ($false = X0) | ($false = ((s__destination @ X1 @ s__Bill @ (sK0 @ cw)))) | (((s__instance @ X2 @ s__PayCheck)) = $false) | (((s__instance @ X1 @ s__Giving)) = $false) | (X0 = X3) | (((s__patient @ X1 @ X2 @ (sK0 @ cw))) = X3) | (X3 = X4) | ($false = X4) [superposition 221,118]
576. ($false != X1) | (X1 = X2) | (X0 = X2) | (X0 = X2) | ($false = X0) [equality factoring 118]
623. ($false != X1) | (X1 = X2) | (X0 = X2) | ($false = X0) [duplicate literal removal 576]
713. (((s__patient @ X1 @ X2 @ (sK0 @ cw))) = X3) | ($false = X0) | (X3 = X4) | (X0 = X3) | (((s__instance @ X1 @ s__Giving)) = $false) | (((s__instance @ X2 @ s__PayCheck)) = $false) | ($false = ((s__destination @ X1 @ s__Bill @ (sK0 @ cw)))) [forward subsumption resolution 414,623]
749. ($true != $true) | ($true = ((s__instance @ (sK4 @ (sK0 @ X0)) @ s__PayCheck))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 85,63]
765. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__instance @ (sK4 @ (sK0 @ X0)) @ s__PayCheck))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 749]
829. ($true != $true) | ($true = ((s__destination @ (sK3 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 87,63]
845. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__destination @ (sK3 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 829]
982. ($true = X0) | (((s__patient @ X2 @ X3 @ (sK0 @ cw))) = X0) | ($false = ((accreln1 @ s__Obligation @ X1 @ (sK0 @ X1)))) | (X0 = X4) | ($false = ((s__instance @ X2 @ s__Giving))) | ($false = ((s__instance @ X3 @ s__PayCheck))) | ($false = ((s__destination @ X2 @ s__Bill @ (sK0 @ cw)))) [superposition 63,713]
1122. ($true = $false) | ($true = X0) | (((s__patient @ X2 @ X3 @ (sK0 @ cw))) = X0) | (X0 = X4) | ($false = ((s__instance @ X2 @ s__Giving))) | ($false = ((s__instance @ X3 @ s__PayCheck))) | ($false = ((s__destination @ X2 @ s__Bill @ (sK0 @ cw)))) [forward demodulation 982,63]
1123. (((s__patient @ X2 @ X3 @ (sK0 @ cw))) = X0) | (X0 = X4) | ($true = X0) | ($false = ((s__instance @ X2 @ s__Giving))) | ($false = ((s__instance @ X3 @ s__PayCheck))) | ($false = ((s__destination @ X2 @ s__Bill @ (sK0 @ cw)))) [trivial inequality removal 1122]
1317. ($true != $true) | ($true = ((s__patient @ (sK3 @ (sK0 @ X0)) @ (sK4 @ (sK0 @ X0)) @ (sK0 @ X0)))) | ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [superposition 86,63]
1339. ($true != ((s__destination @ X1 @ X2 @ X0))) | ($true = ((s__patient @ (sK3 @ (sK0 @ X0)) @ (sK4 @ (sK0 @ X0)) @ (sK0 @ X0)))) | ($true != ((s__agent @ X1 @ s__Bill @ X0))) | ($true != ((s__instance @ X2 @ s__Moon))) | ($true != ((s__instance @ X1 @ s__Flying))) [trivial inequality removal 1317]
2365. ($true != $true) | ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__Giving))) | ($true != ((s__agent @ sK5 @ s__Bill @ cw))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [superposition 307,91]
2394. ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__Giving))) | ($true != ((s__agent @ sK5 @ s__Bill @ cw))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [trivial inequality removal 2365]
2407. ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__Giving))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [forward subsumption resolution 2394,90]
2410. ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__Giving))) | ($true != ((s__instance @ sK5 @ s__Flying))) [forward subsumption resolution 2407,89]
2413. ($true = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__Giving))) [forward subsumption resolution 2410,88]
2479. ($true != $true) | ($true = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__agent @ sK5 @ s__Bill @ cw))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [superposition 765,91]
2508. ($true = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__agent @ sK5 @ s__Bill @ cw))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [trivial inequality removal 2479]
2531. ($true = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [forward subsumption resolution 2508,90]
2542. ($true = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) | ($true != ((s__instance @ sK5 @ s__Flying))) [forward subsumption resolution 2531,89]
2545. ($true = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) [forward subsumption resolution 2542,88]
2634. ($true != $true) | ($true = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__agent @ sK5 @ s__Bill @ cw))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [superposition 845,91]
2663. ($true = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__agent @ sK5 @ s__Bill @ cw))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [trivial inequality removal 2634]
2676. ($true = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [forward subsumption resolution 2663,90]
2679. ($true = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) | ($true != ((s__instance @ sK5 @ s__Flying))) [forward subsumption resolution 2676,89]
2682. ($true = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward subsumption resolution 2679,88]
3815. ($true != $true) | ($true = ((s__patient @ (sK3 @ (sK0 @ cw)) @ (sK4 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__agent @ sK5 @ s__Bill @ cw))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [superposition 1339,91]
3844. ($true = ((s__patient @ (sK3 @ (sK0 @ cw)) @ (sK4 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__agent @ sK5 @ s__Bill @ cw))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [trivial inequality removal 3815]
3858. ($true = ((s__patient @ (sK3 @ (sK0 @ cw)) @ (sK4 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__instance @ sK6 @ s__Moon))) | ($true != ((s__instance @ sK5 @ s__Flying))) [forward subsumption resolution 3844,90]
3861. ($true = ((s__patient @ (sK3 @ (sK0 @ cw)) @ (sK4 @ (sK0 @ cw)) @ (sK0 @ cw)))) | ($true != ((s__instance @ sK5 @ s__Flying))) [forward subsumption resolution 3858,89]
3864. ($true = ((s__patient @ (sK3 @ (sK0 @ cw)) @ (sK4 @ (sK0 @ cw)) @ (sK0 @ cw)))) [forward subsumption resolution 3861,88]
9031. ($true = X0) | (X0 = X1) | ($true = X0) | ($false = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__Giving))) | ($false = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) | ($false = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [superposition 1123,3864]
9038. ($true = X0) | (X0 = X1) | ($false = ((s__instance @ (sK3 @ (sK0 @ cw)) @ s__Giving))) | ($false = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) | ($false = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [duplicate literal removal 9031]
9058. ($true = $false) | ($true = X0) | (X0 = X1) | ($false = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) | ($false = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward demodulation 9038,2413]
9059. ($true = X0) | (X0 = X1) | ($false = ((s__instance @ (sK4 @ (sK0 @ cw)) @ s__PayCheck))) | ($false = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [trivial inequality removal 9058]
9095. ($true = $false) | ($true = X0) | (X0 = X1) | ($false = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [forward demodulation 9059,2545]
9096. ($true = X0) | (X0 = X1) | ($false = ((s__destination @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) [trivial inequality removal 9095]
9129. ($true = $false) | ($true = X0) | (X0 = X1) [forward demodulation 9096,2682]
9130. ($true = X0) | (X0 = X1) [trivial inequality removal 9129]
9231. ($true != X0) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | ($true != ((s__patient @ X1 @ X2 @ (sK0 @ cw)))) | ($true = ((s__destination @ X1 @ s__Bill @ (sK0 @ cw)))) [superposition 220,9130]
9343. ($true != X1) | (X0 = X1) [equality factoring 9130]
9345. (X0 != X0) | ($true = X0) [equality factoring 9130]
9346. ($true = X0) [trivial inequality removal 9345]
9418. (X0 = X1) [forward subsumption resolution 9343,9346]
9435. ($true != X0) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | ($true != ((s__patient @ X1 @ X2 @ (sK0 @ cw)))) [forward subsumption resolution 9231,220]
9493. (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | ($true != ((s__patient @ X1 @ X2 @ (sK0 @ cw)))) [forward subsumption resolution 9435,9346]
9495. (((s__instance @ X1 @ s__Giving)) != $true) | ($true != ((s__patient @ X1 @ X2 @ (sK0 @ cw)))) [forward subsumption resolution 9493,9418]
9497. ($true != ((s__patient @ X1 @ X2 @ (sK0 @ cw)))) [forward subsumption resolution 9495,9418]
9499. $false [forward subsumption resolution 9497,9346]
% SZS output end Proof for 1g_Obligation_UO1_existConj
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.757 s
% Peak memory usage: 16 MB
% Instructions burned: 4370 (million)
% ------------------------------
% ------------------------------
