% Checking upload ...
% Checker ran ...
% No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/S2ICK9GvZx/SOT_yfUSpD
% (2606084)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_yfUSpD for (1ds/27Mi)
% (2606088)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_yfUSpD for (1ds/18Mi)
% (2606087)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_yfUSpD for (1ds/275Mi)
% (2606083)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_yfUSpD for (1ds/4Mi)
% (2606086)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_yfUSpD for (1ds/2Mi)
% (2606082)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_yfUSpD for (1ds/183Mi)
% (2606085)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_yfUSpD for (1ds/2Mi)
% (2606086)Refutation not found, incomplete strategy
% (2606086)------------------------------
% (2606086)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2606086)Termination reason: Refutation not found, incomplete strategy


% (2606086)Memory used [KB]: 5628
% (2606086)Time elapsed: 0.006 s
% (2606086)------------------------------
% (2606086)------------------------------
% (2606089)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_yfUSpD for (1ds/3Mi)
% (2606085)First to succeed.
% (2606085)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_yfUSpD
% SZS output start Proof for SOT_yfUSpD
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m &gt; w &gt; w &gt; $o).
thf(func_def_5, type, accreln2: m &gt; $i &gt; w &gt; w &gt; $o).
thf(func_def_6, type, accreln3: m &gt; $i &gt; $i &gt; w &gt; w &gt; $o).
thf(func_def_8, type, s__instance: $i &gt; $i &gt; $o).
thf(func_def_9, type, s__agent: $i &gt; $i &gt; w &gt; $o).
thf(func_def_10, type, s__patient: $i &gt; $i &gt; w &gt; $o).
thf(func_def_11, type, s__destination: $i &gt; $i &gt; w &gt; $o).
thf(func_def_12, type, s__located: $i &gt; $i &gt; w &gt; $o).
thf(func_def_25, type, sK1: w &gt; $i).
thf(func_def_26, type, sK2: w &gt; $i).
thf(func_def_27, type, sK3: w &gt; w).
thf(func_def_28, type, sK4: w &gt; w).
thf(func_def_29, type, sK5: w &gt; $i).
thf(func_def_30, type, sK6: w &gt; $i).
thf(f335,plain,(
  $false),
  inference(subsumption_resolution,[status(thm)],[f334,f281])).
thf(f281,plain,(
  ($true = (s__instance @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Driving))),
  inference(trivial_inequality_removal,[status(thm)],[f270])).
thf(f270,plain,(
  ($true = (s__instance @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Driving)) | ($true != $true)),
  inference(superposition,[status(thm)],[f251,f103])).
thf(f103,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Obligation @ X0 @ (sK3 @ X0)))) )),
  inference(cnf_transformation,[status(thm)],[f76])).
thf(f76,plain,(
  ! [X0 : w] : ($true = (accreln1 @ s__Obligation @ X0 @ (sK3 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f44,f75])).
thf(f75,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true) =&gt; ($true = (accreln1 @ s__Obligation @ X0 @ (sK3 @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f44,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f43])).
thf(f43,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  file('/tmp/S2ICK9GvZx/SOT_yfUSpD',serial_obligation)).
thf(f251,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__instance @ (sK6 @ X0) @ s__Driving))) )),
  inference(subsumption_resolution,[status(thm)],[f250,f231])).
thf(f231,plain,(
  ($true = (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax))),
  inference(subsumption_resolution,[status(thm)],[f230,f83])).
thf(f83,plain,(
  ($true = (s__instance @ sK0 @ s__PayCheck))),
  inference(cnf_transformation,[status(thm)],[f70])).
thf(f70,plain,(
  ((s__destination @ sK0 @ s__Bill @ cw) = $true) & ($true = (s__instance @ sK0 @ s__PayCheck))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f60,f69])).
thf(f69,plain,(
  ? [X0 : $i] : (($true = (s__destination @ X0 @ s__Bill @ cw)) & ($true = (s__instance @ X0 @ s__PayCheck))) =&gt; (((s__destination @ sK0 @ s__Bill @ cw) = $true) & ($true = (s__instance @ sK0 @ s__PayCheck)))),
  introduced(definition,[],[choice_axiom])).
thf(f60,plain,(
  ? [X0 : $i] : (($true = (s__destination @ X0 @ s__Bill @ cw)) & ($true = (s__instance @ X0 @ s__PayCheck)))),
  inference(fool_elimination,[status(thm)],[f59])).
thf(f59,plain,(
  ? [X0 : $i] : ((s__destination @ X0 @ s__Bill @ cw) & (s__instance @ X0 @ s__PayCheck))),
  inference(rectify,[status(thm)],[f39])).
thf(f39,axiom,(
  ? [X3 : $i] : ((s__destination @ X3 @ s__Bill @ cw) & (s__instance @ X3 @ s__PayCheck))),
  file('/tmp/S2ICK9GvZx/SOT_yfUSpD',optionA_bill_gets_paycheck)).
thf(f230,plain,(
  ($true != (s__instance @ sK0 @ s__PayCheck)) | ($true = (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax))),
  inference(trivial_inequality_removal,[status(thm)],[f229])).
thf(f229,plain,(
  ($true != (s__instance @ sK0 @ s__PayCheck)) | ($true = (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != $true)),
  inference(superposition,[status(thm)],[f148,f84])).
thf(f84,plain,(
  ((s__destination @ sK0 @ s__Bill @ cw) = $true)),
  inference(cnf_transformation,[status(thm)],[f70])).
thf(f148,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true = (s__instance @ (sK2 @ (sK3 @ X0)) @ s__Tax)) | ((s__instance @ X1 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f146])).
thf(f146,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__PayCheck) != $true) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true != $true) | ($true = (s__instance @ (sK2 @ (sK3 @ X0)) @ s__Tax))) )),
  inference(superposition,[status(thm)],[f92,f103])).
thf(f92,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ($true = (s__instance @ (sK2 @ X1) @ s__Tax)) | ($true != (s__destination @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X0 @ s__PayCheck))) )),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f74,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (($true != (accreln1 @ s__Obligation @ X2 @ X1)) | (((s__agent @ (sK1 @ X1) @ s__Bill @ X1) = $true) & ($true = (s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1)) & ($true = (s__instance @ (sK2 @ X1) @ s__Tax)) & ($true = (s__instance @ (sK1 @ X1) @ s__Giving))) | ($true != (s__destination @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X0 @ s__PayCheck)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f72,f73])).
thf(f73,plain,(
  ! [X1 : w] : (? [X3 : $i,X4 : $i] : (((s__agent @ X3 @ s__Bill @ X1) = $true) & ($true = (s__patient @ X3 @ X4 @ X1)) & ($true = (s__instance @ X4 @ s__Tax)) & ($true = (s__instance @ X3 @ s__Giving))) =&gt; (((s__agent @ (sK1 @ X1) @ s__Bill @ X1) = $true) & ($true = (s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1)) & ($true = (s__instance @ (sK2 @ X1) @ s__Tax)) & ($true = (s__instance @ (sK1 @ X1) @ s__Giving))))),
  introduced(definition,[],[choice_axiom])).
thf(f72,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ? [X3 : $i,X4 : $i] : (((s__agent @ X3 @ s__Bill @ X1) = $true) & ($true = (s__patient @ X3 @ X4 @ X1)) & ($true = (s__instance @ X4 @ s__Tax)) & ($true = (s__instance @ X3 @ s__Giving))) | ($true != (s__destination @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X0 @ s__PayCheck)))),
  inference(rectify,[status(thm)],[f62])).
thf(f62,plain,(
  ! [X1 : $i,X2 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ? [X3 : $i,X4 : $i] : (($true = (s__agent @ X3 @ s__Bill @ X2)) & ($true = (s__patient @ X3 @ X4 @ X2)) & ($true = (s__instance @ X4 @ s__Tax)) & ($true = (s__instance @ X3 @ s__Giving))) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__PayCheck) != $true))),
  inference(flattening,[status(thm)],[f61])).
thf(f61,plain,(
  ! [X1 : $i,X0 : w,X2 : w] : ((? [X3 : $i,X4 : $i] : (($true = (s__agent @ X3 @ s__Bill @ X2)) & ($true = (s__patient @ X3 @ X4 @ X2)) & ($true = (s__instance @ X4 @ s__Tax)) & ($true = (s__instance @ X3 @ s__Giving))) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)) | (((s__instance @ X1 @ s__PayCheck) != $true) | ($true != (s__destination @ X1 @ s__Bill @ X0))))),
  inference(ennf_transformation,[status(thm)],[f50])).
thf(f50,plain,(
  ! [X1 : $i,X0 : w,X2 : w] : ((((s__instance @ X1 @ s__PayCheck) = $true) & ($true = (s__destination @ X1 @ s__Bill @ X0))) =&gt; (((accreln1 @ s__Obligation @ X0 @ X2) = $true) =&gt; ? [X3 : $i,X4 : $i] : (($true = (s__agent @ X3 @ s__Bill @ X2)) & ($true = (s__patient @ X3 @ X4 @ X2)) & ($true = (s__instance @ X4 @ s__Tax)) & ($true = (s__instance @ X3 @ s__Giving)))))),
  inference(fool_elimination,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X0 : w,X1 : $i,X2 : w] : (((s__instance @ X1 @ s__PayCheck) & (s__destination @ X1 @ s__Bill @ X0)) =&gt; ((accreln1 @ s__Obligation @ X0 @ X2) =&gt; ? [X3 : $i,X4 : $i] : ((s__instance @ X3 @ s__Giving) & (s__patient @ X3 @ X4 @ X2) & (s__instance @ X4 @ s__Tax) & (s__agent @ X3 @ s__Bill @ X2))))),
  inference(rectify,[status(thm)],[f37])).
thf(f37,axiom,(
  ! [X0 : w,X3 : $i,X1 : w] : (((s__instance @ X3 @ s__PayCheck) & (s__destination @ X3 @ s__Bill @ X0)) =&gt; ((accreln1 @ s__Obligation @ X0 @ X1) =&gt; ? [X4 : $i,X5 : $i] : ((s__instance @ X4 @ s__Giving) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__Tax) & (s__agent @ X4 @ s__Bill @ X1))))),
  file('/tmp/S2ICK9GvZx/SOT_yfUSpD',obligation1_pay_taxes)).
thf(f250,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__instance @ (sK6 @ X0) @ s__Driving))) )),
  inference(subsumption_resolution,[status(thm)],[f249,f228])).
thf(f228,plain,(
  ($true = (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving))),
  inference(subsumption_resolution,[status(thm)],[f227,f83])).
thf(f227,plain,(
  ($true != (s__instance @ sK0 @ s__PayCheck)) | ($true = (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f226])).
thf(f226,plain,(
  ($true != $true) | ($true = (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true != (s__instance @ sK0 @ s__PayCheck))),
  inference(superposition,[status(thm)],[f137,f84])).
thf(f137,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ (sK1 @ (sK3 @ X0)) @ s__Giving) = $true) | ((s__instance @ X1 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f136])).
thf(f136,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__PayCheck) != $true) | ($true != $true) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ (sK1 @ (sK3 @ X0)) @ s__Giving) = $true)) )),
  inference(superposition,[status(thm)],[f91,f103])).
thf(f91,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ($true = (s__instance @ (sK1 @ X1) @ s__Giving)) | ($true != (s__destination @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X0 @ s__PayCheck))) )),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f249,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__instance @ (sK6 @ X0) @ s__Driving)) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax))) )),
  inference(subsumption_resolution,[status(thm)],[f244,f234])).
thf(f234,plain,(
  ($true = (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw)))),
  inference(subsumption_resolution,[status(thm)],[f233,f83])).
thf(f233,plain,(
  ($true != (s__instance @ sK0 @ s__PayCheck)) | ($true = (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f232])).
thf(f232,plain,(
  ($true = (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != (s__instance @ sK0 @ s__PayCheck)) | ($true != $true)),
  inference(superposition,[status(thm)],[f153,f84])).
thf(f153,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true = (s__agent @ (sK1 @ (sK3 @ X0)) @ s__Bill @ (sK3 @ X0))) | ((s__instance @ X1 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f150])).
thf(f150,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != $true) | ($true = (s__agent @ (sK1 @ (sK3 @ X0)) @ s__Bill @ (sK3 @ X0))) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__PayCheck) != $true)) )),
  inference(superposition,[status(thm)],[f94,f103])).
thf(f94,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ($true != (s__instance @ X0 @ s__PayCheck)) | ((s__agent @ (sK1 @ X1) @ s__Bill @ X1) = $true) | ($true != (s__destination @ X0 @ s__Bill @ X2))) )),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f244,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__instance @ (sK6 @ X0) @ s__Driving)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving))) )),
  inference(trivial_inequality_removal,[status(thm)],[f241])).
thf(f241,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true != $true) | ($true = (s__instance @ (sK6 @ X0) @ s__Driving))) )),
  inference(superposition,[status(thm)],[f120,f237])).
thf(f237,plain,(
  ($true = (s__patient @ (sK1 @ (sK3 @ cw)) @ (sK2 @ (sK3 @ cw)) @ (sK3 @ cw)))),
  inference(subsumption_resolution,[status(thm)],[f236,f83])).
thf(f236,plain,(
  ($true = (s__patient @ (sK1 @ (sK3 @ cw)) @ (sK2 @ (sK3 @ cw)) @ (sK3 @ cw))) | ($true != (s__instance @ sK0 @ s__PayCheck))),
  inference(trivial_inequality_removal,[status(thm)],[f235])).
thf(f235,plain,(
  ($true != (s__instance @ sK0 @ s__PayCheck)) | ($true != $true) | ($true = (s__patient @ (sK1 @ (sK3 @ cw)) @ (sK2 @ (sK3 @ cw)) @ (sK3 @ cw)))),
  inference(superposition,[status(thm)],[f170,f84])).
thf(f170,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((s__patient @ (sK1 @ (sK3 @ X0)) @ (sK2 @ (sK3 @ X0)) @ (sK3 @ X0)) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f166])).
thf(f166,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__PayCheck) != $true) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__patient @ (sK1 @ (sK3 @ X0)) @ (sK2 @ (sK3 @ X0)) @ (sK3 @ X0)) = $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f93,f103])).
thf(f93,plain,(
  ( ! [X2 : w,X0 : $i,X1 : w] : (($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ($true = (s__patient @ (sK1 @ X1) @ (sK2 @ X1) @ X1)) | ($true != (s__destination @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X0 @ s__PayCheck))) )),
  inference(cnf_transformation,[status(thm)],[f74])).
thf(f120,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X3 @ X0)) | ((s__instance @ X1 @ s__Giving) != $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true = (s__instance @ (sK6 @ X2) @ s__Driving)) | ($true != (s__instance @ X3 @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f82])).
thf(f82,plain,(
  ! [X0 : w,X1 : $i,X2 : w,X3 : $i] : (($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X3 @ s__Tax)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ((s__instance @ X1 @ s__Giving) != $true) | (($true = (s__agent @ (sK6 @ X2) @ s__Bill @ X2)) & ($true = (s__instance @ (sK5 @ X2) @ s__OfficeBuilding)) & ((s__destination @ (sK6 @ X2) @ (sK5 @ X2) @ X2) = $true) & ($true = (s__instance @ (sK6 @ X2) @ s__Driving))) | ($true != (s__patient @ X1 @ X3 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f80,f81])).
thf(f81,plain,(
  ! [X2 : w] : (? [X4 : $i,X5 : $i] : (((s__agent @ X5 @ s__Bill @ X2) = $true) & ($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X2)) & ($true = (s__instance @ X5 @ s__Driving))) =&gt; (($true = (s__agent @ (sK6 @ X2) @ s__Bill @ X2)) & ($true = (s__instance @ (sK5 @ X2) @ s__OfficeBuilding)) & ((s__destination @ (sK6 @ X2) @ (sK5 @ X2) @ X2) = $true) & ($true = (s__instance @ (sK6 @ X2) @ s__Driving))))),
  introduced(definition,[],[choice_axiom])).
thf(f80,plain,(
  ! [X0 : w,X1 : $i,X2 : w,X3 : $i] : (($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X3 @ s__Tax)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ((s__instance @ X1 @ s__Giving) != $true) | ? [X4 : $i,X5 : $i] : (((s__agent @ X5 @ s__Bill @ X2) = $true) & ($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X2)) & ($true = (s__instance @ X5 @ s__Driving))) | ($true != (s__patient @ X1 @ X3 @ X0)))),
  inference(rectify,[status(thm)],[f68])).
thf(f68,plain,(
  ! [X3 : w,X2 : $i,X0 : w,X1 : $i] : (($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (s__instance @ X1 @ s__Tax)) | ($true != (accreln1 @ s__Obligation @ X3 @ X0)) | ($true != (s__instance @ X2 @ s__Giving)) | ? [X4 : $i,X5 : $i] : (((s__agent @ X5 @ s__Bill @ X0) = $true) & ($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X0)) & ($true = (s__instance @ X5 @ s__Driving))) | ($true != (s__patient @ X2 @ X1 @ X3)))),
  inference(flattening,[status(thm)],[f67])).
thf(f67,plain,(
  ! [X1 : $i,X3 : w,X2 : $i,X0 : w] : ((? [X4 : $i,X5 : $i] : (((s__agent @ X5 @ s__Bill @ X0) = $true) & ($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X0)) & ($true = (s__instance @ X5 @ s__Driving))) | ($true != (accreln1 @ s__Obligation @ X3 @ X0))) | (($true != (s__agent @ X2 @ s__Bill @ X3)) | ($true != (s__patient @ X2 @ X1 @ X3)) | ($true != (s__instance @ X2 @ s__Giving)) | ($true != (s__instance @ X1 @ s__Tax))))),
  inference(ennf_transformation,[status(thm)],[f58])).
thf(f58,plain,(
  ! [X1 : $i,X3 : w,X2 : $i,X0 : w] : ((($true = (s__agent @ X2 @ s__Bill @ X3)) & ($true = (s__patient @ X2 @ X1 @ X3)) & ($true = (s__instance @ X2 @ s__Giving)) & ($true = (s__instance @ X1 @ s__Tax))) =&gt; (($true = (accreln1 @ s__Obligation @ X3 @ X0)) =&gt; ? [X4 : $i,X5 : $i] : (((s__agent @ X5 @ s__Bill @ X0) = $true) & ($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X0)) & ($true = (s__instance @ X5 @ s__Driving)))))),
  inference(fool_elimination,[status(thm)],[f57])).
thf(f57,plain,(
  ! [X0 : w,X1 : $i,X2 : $i,X3 : w] : (((s__patient @ X2 @ X1 @ X3) & (s__agent @ X2 @ s__Bill @ X3) & (s__instance @ X1 @ s__Tax) & (s__instance @ X2 @ s__Giving)) =&gt; ((accreln1 @ s__Obligation @ X3 @ X0) =&gt; ? [X4 : $i,X5 : $i] : ((s__destination @ X5 @ X4 @ X0) & (s__agent @ X5 @ s__Bill @ X0) & (s__instance @ X4 @ s__OfficeBuilding) & (s__instance @ X5 @ s__Driving))))),
  inference(rectify,[status(thm)],[f38])).
thf(f38,axiom,(
  ! [X1 : w,X5 : $i,X4 : $i,X0 : w] : (((s__patient @ X4 @ X5 @ X0) & (s__agent @ X4 @ s__Bill @ X0) & (s__instance @ X5 @ s__Tax) & (s__instance @ X4 @ s__Giving)) =&gt; ((accreln1 @ s__Obligation @ X0 @ X1) =&gt; ? [X7 : $i,X6 : $i] : ((s__destination @ X6 @ X7 @ X1) & (s__agent @ X6 @ s__Bill @ X1) & (s__instance @ X7 @ s__OfficeBuilding) & (s__instance @ X6 @ s__Driving))))),
  file('/tmp/S2ICK9GvZx/SOT_yfUSpD',obligation2_drive_to_office)).
thf(f334,plain,(
  ($true != (s__instance @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Driving))),
  inference(subsumption_resolution,[status(thm)],[f333,f267])).
thf(f267,plain,(
  ((s__instance @ (sK5 @ (sK3 @ (sK3 @ cw))) @ s__OfficeBuilding) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f258])).
thf(f258,plain,(
  ($true != $true) | ((s__instance @ (sK5 @ (sK3 @ (sK3 @ cw))) @ s__OfficeBuilding) = $true)),
  inference(superposition,[status(thm)],[f248,f103])).
thf(f248,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__instance @ (sK5 @ X0) @ s__OfficeBuilding))) )),
  inference(subsumption_resolution,[status(thm)],[f247,f228])).
thf(f247,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true = (s__instance @ (sK5 @ X0) @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f246,f234])).
thf(f246,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true = (s__instance @ (sK5 @ X0) @ s__OfficeBuilding)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f245,f231])).
thf(f245,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true = (s__instance @ (sK5 @ X0) @ s__OfficeBuilding)) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving))) )),
  inference(trivial_inequality_removal,[status(thm)],[f240])).
thf(f240,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != $true) | ($true = (s__instance @ (sK5 @ X0) @ s__OfficeBuilding))) )),
  inference(superposition,[status(thm)],[f122,f237])).
thf(f122,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X3 @ X0)) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__Giving) != $true) | ($true = (s__instance @ (sK5 @ X2) @ s__OfficeBuilding)) | ($true != (s__instance @ X3 @ s__Tax)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f82])).
thf(f333,plain,(
  ((s__instance @ (sK5 @ (sK3 @ (sK3 @ cw))) @ s__OfficeBuilding) != $true) | ($true != (s__instance @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Driving))),
  inference(subsumption_resolution,[status(thm)],[f332,f305])).
thf(f305,plain,(
  ((s__agent @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Bill @ (sK3 @ (sK3 @ cw))) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f294])).
thf(f294,plain,(
  ((s__agent @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Bill @ (sK3 @ (sK3 @ cw))) = $true) | ($true != $true)),
  inference(superposition,[status(thm)],[f254,f103])).
thf(f254,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__agent @ (sK6 @ X0) @ s__Bill @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f253,f234])).
thf(f253,plain,(
  ( ! [X0 : w] : (($true = (s__agent @ (sK6 @ X0) @ s__Bill @ X0)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw)))) )),
  inference(subsumption_resolution,[status(thm)],[f252,f231])).
thf(f252,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true = (s__agent @ (sK6 @ X0) @ s__Bill @ X0)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f243,f228])).
thf(f243,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__agent @ (sK6 @ X0) @ s__Bill @ X0)) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax))) )),
  inference(trivial_inequality_removal,[status(thm)],[f239])).
thf(f239,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != $true) | ($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true = (s__agent @ (sK6 @ X0) @ s__Bill @ X0))) )),
  inference(superposition,[status(thm)],[f123,f237])).
thf(f123,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X3 @ X0)) | ((s__instance @ X1 @ s__Giving) != $true) | ($true != (s__instance @ X3 @ s__Tax)) | ($true = (s__agent @ (sK6 @ X2) @ s__Bill @ X2)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ($true != (s__agent @ X1 @ s__Bill @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f82])).
thf(f332,plain,(
  ((s__agent @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Bill @ (sK3 @ (sK3 @ cw))) != $true) | ((s__instance @ (sK5 @ (sK3 @ (sK3 @ cw))) @ s__OfficeBuilding) != $true) | ($true != (s__instance @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Driving))),
  inference(subsumption_resolution,[status(thm)],[f331,f103])).
thf(f331,plain,(
  ($true != (accreln1 @ s__Obligation @ cw @ (sK3 @ cw))) | ((s__instance @ (sK5 @ (sK3 @ (sK3 @ cw))) @ s__OfficeBuilding) != $true) | ((s__agent @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Bill @ (sK3 @ (sK3 @ cw))) != $true) | ($true != (s__instance @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Driving))),
  inference(trivial_inequality_removal,[status(thm)],[f330])).
thf(f330,plain,(
  ($true != (s__instance @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Driving)) | ((s__agent @ (sK6 @ (sK3 @ (sK3 @ cw))) @ s__Bill @ (sK3 @ (sK3 @ cw))) != $true) | ($true != $true) | ((s__instance @ (sK5 @ (sK3 @ (sK3 @ cw))) @ s__OfficeBuilding) != $true) | ($true != (accreln1 @ s__Obligation @ cw @ (sK3 @ cw)))),
  inference(superposition,[status(thm)],[f194,f317])).
thf(f317,plain,(
  ($true = (s__destination @ (sK6 @ (sK3 @ (sK3 @ cw))) @ (sK5 @ (sK3 @ (sK3 @ cw))) @ (sK3 @ (sK3 @ cw))))),
  inference(trivial_inequality_removal,[status(thm)],[f306])).
thf(f306,plain,(
  ($true = (s__destination @ (sK6 @ (sK3 @ (sK3 @ cw))) @ (sK5 @ (sK3 @ (sK3 @ cw))) @ (sK3 @ (sK3 @ cw)))) | ($true != $true)),
  inference(superposition,[status(thm)],[f257,f103])).
thf(f257,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__destination @ (sK6 @ X0) @ (sK5 @ X0) @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f256,f228])).
thf(f256,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__destination @ (sK6 @ X0) @ (sK5 @ X0) @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f255,f234])).
thf(f255,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true = (s__destination @ (sK6 @ X0) @ (sK5 @ X0) @ X0)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f242,f231])).
thf(f242,plain,(
  ( ! [X0 : w] : (($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true = (s__destination @ (sK6 @ X0) @ (sK5 @ X0) @ X0)) | ($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving))) )),
  inference(trivial_inequality_removal,[status(thm)],[f238])).
thf(f238,plain,(
  ( ! [X0 : w] : (($true != (s__agent @ (sK1 @ (sK3 @ cw)) @ s__Bill @ (sK3 @ cw))) | ($true != (s__instance @ (sK2 @ (sK3 @ cw)) @ s__Tax)) | ($true != (s__instance @ (sK1 @ (sK3 @ cw)) @ s__Giving)) | ($true = (s__destination @ (sK6 @ X0) @ (sK5 @ X0) @ X0)) | ($true != (accreln1 @ s__Obligation @ (sK3 @ cw) @ X0)) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f121,f237])).
thf(f121,plain,(
  ( ! [X2 : w,X3 : $i,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X3 @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ((s__destination @ (sK6 @ X2) @ (sK5 @ X2) @ X2) = $true) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__Giving) != $true) | ($true != (s__instance @ X3 @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f82])).
thf(f194,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ X2 @ (sK3 @ X0))) | ($true != (s__instance @ X1 @ s__Driving)) | ($true != (s__agent @ X1 @ s__Bill @ (sK3 @ X0))) | ((accreln1 @ s__Obligation @ cw @ X0) != $true) | ($true != (s__instance @ X2 @ s__OfficeBuilding))) )),
  inference(trivial_inequality_removal,[status(thm)],[f187])).
thf(f187,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ X2 @ (sK3 @ X0))) | ($true != $true) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__agent @ X1 @ s__Bill @ (sK3 @ X0))) | ($true != (s__instance @ X1 @ s__Driving)) | ((accreln1 @ s__Obligation @ cw @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f111,f131])).
thf(f131,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Permission @ X0 @ (sK3 @ X0)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f130])).
thf(f130,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true = (accreln1 @ s__Permission @ X0 @ (sK3 @ X0)))) )),
  inference(superposition,[status(thm)],[f126,f103])).
thf(f126,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((accreln1 @ s__Permission @ X0 @ X1) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f63])).
thf(f63,plain,(
  ! [X0 : w,X1 : w] : (((accreln1 @ s__Permission @ X0 @ X1) = $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true))),
  inference(ennf_transformation,[status(thm)],[f48])).
thf(f48,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) =&gt; ((accreln1 @ s__Permission @ X0 @ X1) = $true))),
  inference(fool_elimination,[status(thm)],[f47])).
thf(f47,plain,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) =&gt; (accreln1 @ s__Permission @ X0 @ X1))),
  inference(rectify,[status(thm)],[f4])).
thf(f4,axiom,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) =&gt; (accreln1 @ s__Permission @ X0 @ X1))),
  file('/tmp/S2ICK9GvZx/SOT_yfUSpD',bridge_obligation_permission)).
thf(f111,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (accreln1 @ s__Permission @ X1 @ X0)) | ($true != (s__instance @ X2 @ s__Driving)) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ($true != (s__instance @ X3 @ s__OfficeBuilding)) | ((s__destination @ X2 @ X3 @ X0) != $true) | ($true != (accreln1 @ s__Obligation @ cw @ X1))) )),
  inference(cnf_transformation,[status(thm)],[f79])).
thf(f79,plain,(
  ! [X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X1)) | ($true != (accreln1 @ s__Permission @ X1 @ X0)) | ! [X2 : $i,X3 : $i] : (((s__destination @ X2 @ X3 @ X0) != $true) | ($true != (s__agent @ X2 @ s__Bill @ X0)) | ($true != (s__instance @ X3 @ s__OfficeBuilding)) | ($true != (s__instance @ X2 @ s__Driving))))),
  inference(rectify,[status(thm)],[f66])).
thf(f66,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ cw @ X0) != $true) | ((accreln1 @ s__Permission @ X0 @ X1) != $true) | ! [X3 : $i,X2 : $i] : (($true != (s__destination @ X3 @ X2 @ X1)) | ((s__agent @ X3 @ s__Bill @ X1) != $true) | ($true != (s__instance @ X2 @ s__OfficeBuilding)) | ($true != (s__instance @ X3 @ s__Driving))))),
  inference(ennf_transformation,[status(thm)],[f56])).
thf(f56,plain,(
  ~? [X0 : w,X1 : w] : (((accreln1 @ s__Permission @ X0 @ X1) = $true) & ? [X3 : $i,X2 : $i] : (($true = (s__destination @ X3 @ X2 @ X1)) & ($true = (s__instance @ X2 @ s__OfficeBuilding)) & ($true = (s__instance @ X3 @ s__Driving)) & ((s__agent @ X3 @ s__Bill @ X1) = $true)) & ((accreln1 @ s__Obligation @ cw @ X0) = $true))),
  inference(fool_elimination,[status(thm)],[f55])).
thf(f55,plain,(
  ~? [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ cw @ X0) & (accreln1 @ s__Permission @ X0 @ X1) & ? [X2 : $i,X3 : $i] : ((s__destination @ X3 @ X2 @ X1) & (s__agent @ X3 @ s__Bill @ X1) & (s__instance @ X2 @ s__OfficeBuilding) & (s__instance @ X3 @ s__Driving)))),
  inference(rectify,[status(thm)],[f41])).
thf(f41,negated_conjecture,(
  ~? [X8 : w,X9 : w] : ((accreln1 @ s__Obligation @ cw @ X8) & (accreln1 @ s__Permission @ X8 @ X9) & ? [X7 : $i,X6 : $i] : ((s__destination @ X6 @ X7 @ X9) & (s__agent @ X6 @ s__Bill @ X9) & (s__instance @ X7 @ s__OfficeBuilding) & (s__instance @ X6 @ s__Driving)))),
  inference(negated_conjecture,[status(cth)],[f40])).
thf(f40,conjecture,(
  ? [X8 : w,X9 : w] : ((accreln1 @ s__Obligation @ cw @ X8) & (accreln1 @ s__Permission @ X8 @ X9) & ? [X7 : $i,X6 : $i] : ((s__destination @ X6 @ X7 @ X9) & (s__agent @ X6 @ s__Bill @ X9) & (s__instance @ X7 @ s__OfficeBuilding) & (s__instance @ X6 @ s__Driving)))),
  file('/tmp/S2ICK9GvZx/SOT_yfUSpD',conjecture_is_permitted)).
% SZS output end Proof for SOT_yfUSpD
% (2606085)------------------------------
% (2606085)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2606085)Termination reason: Refutation

% (2606085)Memory used [KB]: 6012
% (2606085)Time elapsed: 0.038 s
% (2606085)------------------------------
% (2606085)------------------------------
% (2606078)Success in time 0.052 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_yfUSpD - Vampire---5.0.1 says Theorem - CPU = 0.00 WC = 0.10 
% OUTPUT: SOT_yfUSpD - Vampire---5.0.1 says Refutation - CPU = 0.01 WC = 0.10 