% CF
% 2d_p_existconj_weakP output under sysK, optA

 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/b9RIdAuXMH/SOT_n2tfEp
% (3338267)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_n2tfEp for (1ds/275Mi)
% (3338263)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_n2tfEp for (1ds/4Mi)
% (3338265)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_n2tfEp for (1ds/2Mi)
% (3338262)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_n2tfEp for (1ds/183Mi)
% (3338266)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_n2tfEp for (1ds/2Mi)
% (3338264)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_n2tfEp for (1ds/27Mi)
% (3338268)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_n2tfEp for (1ds/18Mi)
% (3338267)First to succeed.
% (3338266)Refutation not found, incomplete strategy
% (3338266)------------------------------
% (3338266)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (3338266)Termination reason: Refutation not found, incomplete strategy


% (3338266)Memory used [KB]: 5500
% (3338266)Time elapsed: 0.005 s
% (3338266)------------------------------
% (3338266)------------------------------
% (3338265)Also succeeded, but the first one will report.
% (3338267)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_n2tfEp
% SZS output start Proof for SOT_n2tfEp
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m > w > w > $o).
thf(func_def_5, type, accreln2: m > $i > w > w > $o).
thf(func_def_6, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_7, type, s__instance: $i > $i > $o).
thf(func_def_8, type, s__agent: $i > $i > w > $o).
thf(func_def_9, type, s__patient: $i > $i > w > $o).
thf(func_def_10, type, s__destination: $i > $i > w > $o).
thf(func_def_11, type, s__located: $i > $i > w > $o).
thf(func_def_12, type, s__possesses: $i > $i > w > $o).
thf(func_def_24, type, sK1: w > w).
thf(func_def_25, type, sK2: w > $i).
thf(func_def_26, type, sK3: w > $i).
thf(f78,plain,(
  $false),
  inference(avatar_sat_refutation,[status(thm)],[f62,f66,f73,f77])).
thf(f77,plain,(
  spl4_2 | ~spl4_3),
  inference(avatar_split_clause,[status(thm)],[f76,f71,f60])).
thf(f60,definition,(
  spl4_2 <=> ! [X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) != $true) | ((s__possesses @ s__Bill @ X2 @ cw) != $true))),
  introduced(definition,[new_symbols(naming,[spl4_2])],[avatar_definition])).
thf(f71,definition,(
  spl4_3 <=> ! [X0 : $i] : (((s__located @ (sK3 @ cw) @ X0 @ (sK1 @ cw)) != $true) | ((s__instance @ X0 @ s__ParkingLot) != $true))),
  introduced(definition,[new_symbols(naming,[spl4_3])],[avatar_definition])).
thf(f76,plain,(
  ( ! [X0 : $i] : (((s__possesses @ s__Bill @ X0 @ cw) != $true) | ((s__instance @ X0 @ s__DriversLicense) != $true)) ) | ~spl4_3),
  inference(subsumption_resolution,[status(thm)],[f75,f46])).
thf(f46,plain,(
  ( ! [X0 : $i,X1 : w] : (((s__possesses @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__DriversLicense) != $true) | ((s__instance @ (sK2 @ X1) @ s__ParkingLot) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f38])).
thf(f38,plain,(
  ! [X0 : $i,X1 : w] : (((((s__instance @ (sK2 @ X1) @ s__ParkingLot) = $true) & ((s__agent @ (sK3 @ X1) @ s__Bill @ (sK1 @ X1)) = $true) & ((s__instance @ (sK3 @ X1) @ s__Parking) = $true) & ((s__located @ (sK3 @ X1) @ (sK2 @ X1) @ (sK1 @ X1)) = $true)) & ((accreln1 @ s__Permission @ X1 @ (sK1 @ X1)) = $true)) | ((s__possesses @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__DriversLicense) != $true))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3])],[f35,f37,f36])).
thf(f36,plain,(
  ! [X1 : w] : (? [X2 : w] : (? [X3 : $i,X4 : $i] : (((s__instance @ X3 @ s__ParkingLot) = $true) & ((s__agent @ X4 @ s__Bill @ X2) = $true) & ((s__instance @ X4 @ s__Parking) = $true) & ((s__located @ X4 @ X3 @ X2) = $true)) & ((accreln1 @ s__Permission @ X1 @ X2) = $true)) => (? [X4 : $i,X3 : $i] : (((s__instance @ X3 @ s__ParkingLot) = $true) & ((s__agent @ X4 @ s__Bill @ (sK1 @ X1)) = $true) & ((s__instance @ X4 @ s__Parking) = $true) & ((s__located @ X4 @ X3 @ (sK1 @ X1)) = $true)) & ((accreln1 @ s__Permission @ X1 @ (sK1 @ X1)) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f37,plain,(
  ! [X1 : w] : (? [X4 : $i,X3 : $i] : (((s__instance @ X3 @ s__ParkingLot) = $true) & ((s__agent @ X4 @ s__Bill @ (sK1 @ X1)) = $true) & ((s__instance @ X4 @ s__Parking) = $true) & ((s__located @ X4 @ X3 @ (sK1 @ X1)) = $true)) => (((s__instance @ (sK2 @ X1) @ s__ParkingLot) = $true) & ((s__agent @ (sK3 @ X1) @ s__Bill @ (sK1 @ X1)) = $true) & ((s__instance @ (sK3 @ X1) @ s__Parking) = $true) & ((s__located @ (sK3 @ X1) @ (sK2 @ X1) @ (sK1 @ X1)) = $true)))),
  introduced(definition,[],[choice_axiom])).
thf(f35,plain,(
  ! [X0 : $i,X1 : w] : (? [X2 : w] : (? [X3 : $i,X4 : $i] : (((s__instance @ X3 @ s__ParkingLot) = $true) & ((s__agent @ X4 @ s__Bill @ X2) = $true) & ((s__instance @ X4 @ s__Parking) = $true) & ((s__located @ X4 @ X3 @ X2) = $true)) & ((accreln1 @ s__Permission @ X1 @ X2) = $true)) | ((s__possesses @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__DriversLicense) != $true))),
  inference(rectify,[status(thm)],[f31])).
thf(f31,plain,(
  ! [X1 : $i,X0 : w] : (? [X2 : w] : (? [X4 : $i,X3 : $i] : (((s__instance @ X4 @ s__ParkingLot) = $true) & ((s__agent @ X3 @ s__Bill @ X2) = $true) & ((s__instance @ X3 @ s__Parking) = $true) & ((s__located @ X3 @ X4 @ X2) = $true)) & ((accreln1 @ s__Permission @ X0 @ X2) = $true)) | ((s__possesses @ s__Bill @ X1 @ X0) != $true) | ((s__instance @ X1 @ s__DriversLicense) != $true))),
  inference(flattening,[status(thm)],[f30])).
thf(f30,plain,(
  ! [X1 : $i,X0 : w] : (? [X2 : w] : (? [X4 : $i,X3 : $i] : (((s__instance @ X4 @ s__ParkingLot) = $true) & ((s__agent @ X3 @ s__Bill @ X2) = $true) & ((s__instance @ X3 @ s__Parking) = $true) & ((s__located @ X3 @ X4 @ X2) = $true)) & ((accreln1 @ s__Permission @ X0 @ X2) = $true)) | (((s__instance @ X1 @ s__DriversLicense) != $true) | ((s__possesses @ s__Bill @ X1 @ X0) != $true)))),
  inference(ennf_transformation,[status(thm)],[f29])).
thf(f29,plain,(
  ! [X1 : $i,X0 : w] : ((((s__instance @ X1 @ s__DriversLicense) = $true) & ((s__possesses @ s__Bill @ X1 @ X0) = $true)) => ? [X2 : w] : (? [X4 : $i,X3 : $i] : (((s__instance @ X4 @ s__ParkingLot) = $true) & ((s__agent @ X3 @ s__Bill @ X2) = $true) & ((s__instance @ X3 @ s__Parking) = $true) & ((s__located @ X3 @ X4 @ X2) = $true)) & ((accreln1 @ s__Permission @ X0 @ X2) = $true)))),
  inference(fool_elimination,[status(thm)],[f28])).
thf(f28,plain,(
  ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__DriversLicense) & (s__possesses @ s__Bill @ X1 @ X0)) => ? [X2 : w] : ((accreln1 @ s__Permission @ X0 @ X2) & ? [X3 : $i,X4 : $i] : ((s__agent @ X3 @ s__Bill @ X2) & (s__instance @ X3 @ s__Parking) & (s__located @ X3 @ X4 @ X2) & (s__instance @ X4 @ s__ParkingLot))))),
  inference(rectify,[status(thm)],[f17])).
thf(f17,axiom,(
  ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__DriversLicense) & (s__possesses @ s__Bill @ X1 @ X0)) => ? [X2 : w] : ((accreln1 @ s__Permission @ X0 @ X2) & ? [X3 : $i,X4 : $i] : ((s__agent @ X3 @ s__Bill @ X2) & (s__instance @ X3 @ s__Parking) & (s__located @ X3 @ X4 @ X2) & (s__instance @ X4 @ s__ParkingLot))))),
  file('/tmp/b9RIdAuXMH/SOT_n2tfEp',permission1_to_park)).
thf(f75,plain,(
  ( ! [X0 : $i] : (((s__instance @ (sK2 @ cw) @ s__ParkingLot) != $true) | ((s__possesses @ s__Bill @ X0 @ cw) != $true) | ((s__instance @ X0 @ s__DriversLicense) != $true)) ) | ~spl4_3),
  inference(trivial_inequality_removal,[status(thm)],[f74])).
thf(f74,plain,(
  ( ! [X0 : $i] : (((s__possesses @ s__Bill @ X0 @ cw) != $true) | ((s__instance @ (sK2 @ cw) @ s__ParkingLot) != $true) | ((s__instance @ X0 @ s__DriversLicense) != $true) | ($true != $true)) ) | ~spl4_3),
  inference(superposition,[status(thm)],[f72,f43])).
thf(f43,plain,(
  ( ! [X0 : $i,X1 : w] : (((s__located @ (sK3 @ X1) @ (sK2 @ X1) @ (sK1 @ X1)) = $true) | ((s__instance @ X0 @ s__DriversLicense) != $true) | ((s__possesses @ s__Bill @ X0 @ X1) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f38])).
thf(f72,plain,(
  ( ! [X0 : $i] : (((s__located @ (sK3 @ cw) @ X0 @ (sK1 @ cw)) != $true) | ((s__instance @ X0 @ s__ParkingLot) != $true)) ) | ~spl4_3),
  inference(avatar_component_clause,[status(thm)],[f71])).
thf(f73,plain,(
  spl4_3 | spl4_2 | ~spl4_1),
  inference(avatar_split_clause,[status(thm)],[f69,f57,f60,f71])).
thf(f57,definition,(
  spl4_1 <=> ! [X0 : $i,X1 : $i] : (((s__instance @ X0 @ s__Parking) != $true) | ((s__located @ X0 @ X1 @ (sK1 @ cw)) != $true) | ($true != (s__agent @ X0 @ s__Bill @ (sK1 @ cw))) | ((s__instance @ X1 @ s__ParkingLot) != $true))),
  introduced(definition,[new_symbols(naming,[spl4_1])],[avatar_definition])).
thf(f69,plain,(
  ( ! [X0 : $i,X1 : $i] : (((s__instance @ X1 @ s__DriversLicense) != $true) | ((s__possesses @ s__Bill @ X1 @ cw) != $true) | ((s__located @ (sK3 @ cw) @ X0 @ (sK1 @ cw)) != $true) | ((s__instance @ X0 @ s__ParkingLot) != $true)) ) | ~spl4_1),
  inference(subsumption_resolution,[status(thm)],[f68,f44])).
thf(f44,plain,(
  ( ! [X0 : $i,X1 : w] : (((s__possesses @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__DriversLicense) != $true) | ((s__instance @ (sK3 @ X1) @ s__Parking) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f38])).
thf(f68,plain,(
  ( ! [X0 : $i,X1 : $i] : (((s__possesses @ s__Bill @ X1 @ cw) != $true) | ((s__instance @ (sK3 @ cw) @ s__Parking) != $true) | ((s__located @ (sK3 @ cw) @ X0 @ (sK1 @ cw)) != $true) | ((s__instance @ X1 @ s__DriversLicense) != $true) | ((s__instance @ X0 @ s__ParkingLot) != $true)) ) | ~spl4_1),
  inference(trivial_inequality_removal,[status(thm)],[f67])).
thf(f67,plain,(
  ( ! [X0 : $i,X1 : $i] : (((s__possesses @ s__Bill @ X1 @ cw) != $true) | ((s__instance @ (sK3 @ cw) @ s__Parking) != $true) | ($true != $true) | ((s__instance @ X1 @ s__DriversLicense) != $true) | ((s__instance @ X0 @ s__ParkingLot) != $true) | ((s__located @ (sK3 @ cw) @ X0 @ (sK1 @ cw)) != $true)) ) | ~spl4_1),
  inference(superposition,[status(thm)],[f58,f45])).
thf(f45,plain,(
  ( ! [X0 : $i,X1 : w] : (((s__agent @ (sK3 @ X1) @ s__Bill @ (sK1 @ X1)) = $true) | ((s__possesses @ s__Bill @ X0 @ X1) != $true) | ((s__instance @ X0 @ s__DriversLicense) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f38])).
thf(f58,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__agent @ X0 @ s__Bill @ (sK1 @ cw))) | ((s__instance @ X0 @ s__Parking) != $true) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ((s__located @ X0 @ X1 @ (sK1 @ cw)) != $true)) ) | ~spl4_1),
  inference(avatar_component_clause,[status(thm)],[f57])).
thf(f66,plain,(
  ~spl4_2),
  inference(avatar_contradiction_clause,[status(thm)],[f65])).
thf(f65,plain,(
  $false | ~spl4_2),
  inference(subsumption_resolution,[status(thm)],[f64,f41])).
thf(f41,plain,(
  ((s__instance @ sK0 @ s__DriversLicense) = $true)),
  inference(cnf_transformation,[status(thm)],[f34])).
thf(f34,plain,(
  ((s__instance @ sK0 @ s__DriversLicense) = $true) & ((s__possesses @ s__Bill @ sK0 @ cw) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f25,f33])).
thf(f33,plain,(
  ? [X0 : $i] : (((s__instance @ X0 @ s__DriversLicense) = $true) & ((s__possesses @ s__Bill @ X0 @ cw) = $true)) => (((s__instance @ sK0 @ s__DriversLicense) = $true) & ((s__possesses @ s__Bill @ sK0 @ cw) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f25,plain,(
  ? [X0 : $i] : (((s__instance @ X0 @ s__DriversLicense) = $true) & ((s__possesses @ s__Bill @ X0 @ cw) = $true))),
  inference(fool_elimination,[status(thm)],[f24])).
thf(f24,plain,(
  ? [X0 : $i] : ((s__possesses @ s__Bill @ X0 @ cw) & (s__instance @ X0 @ s__DriversLicense))),
  inference(rectify,[status(thm)],[f18])).
thf(f18,axiom,(
  ? [X1 : $i] : ((s__possesses @ s__Bill @ X1 @ cw) & (s__instance @ X1 @ s__DriversLicense))),
  file('/tmp/b9RIdAuXMH/SOT_n2tfEp',fact_has_license)).
thf(f64,plain,(
  ((s__instance @ sK0 @ s__DriversLicense) != $true) | ~spl4_2),
  inference(trivial_inequality_removal,[status(thm)],[f63])).
thf(f63,plain,(
  ($true != $true) | ((s__instance @ sK0 @ s__DriversLicense) != $true) | ~spl4_2),
  inference(superposition,[status(thm)],[f61,f40])).
thf(f40,plain,(
  ((s__possesses @ s__Bill @ sK0 @ cw) = $true)),
  inference(cnf_transformation,[status(thm)],[f34])).
thf(f61,plain,(
  ( ! [X2 : $i] : (((s__possesses @ s__Bill @ X2 @ cw) != $true) | ((s__instance @ X2 @ s__DriversLicense) != $true)) ) | ~spl4_2),
  inference(avatar_component_clause,[status(thm)],[f60])).
thf(f62,plain,(
  spl4_1 | spl4_2),
  inference(avatar_split_clause,[status(thm)],[f55,f60,f57])).
thf(f55,plain,(
  ( ! [X2 : $i,X0 : $i,X1 : $i] : (((s__instance @ X2 @ s__DriversLicense) != $true) | ((s__possesses @ s__Bill @ X2 @ cw) != $true) | ((s__instance @ X0 @ s__Parking) != $true) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ($true != (s__agent @ X0 @ s__Bill @ (sK1 @ cw))) | ((s__located @ X0 @ X1 @ (sK1 @ cw)) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f54])).
thf(f54,plain,(
  ( ! [X2 : $i,X0 : $i,X1 : $i] : (($true != (s__agent @ X0 @ s__Bill @ (sK1 @ cw))) | ((s__possesses @ s__Bill @ X2 @ cw) != $true) | ((s__located @ X0 @ X1 @ (sK1 @ cw)) != $true) | ((s__instance @ X2 @ s__DriversLicense) != $true) | ($true != $true) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ((s__instance @ X0 @ s__Parking) != $true)) )),
  inference(superposition,[status(thm)],[f49,f42])).
thf(f42,plain,(
  ( ! [X0 : $i,X1 : w] : (((accreln1 @ s__Permission @ X1 @ (sK1 @ X1)) = $true) | ((s__instance @ X0 @ s__DriversLicense) != $true) | ((s__possesses @ s__Bill @ X0 @ X1) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f38])).
thf(f49,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((accreln1 @ s__Permission @ cw @ X0) != $true) | ((s__instance @ X1 @ s__Parking) != $true) | ((s__located @ X1 @ X2 @ X0) != $true) | ($true != (s__instance @ X2 @ s__ParkingLot)) | ((s__agent @ X1 @ s__Bill @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f32])).
thf(f32,plain,(
  ! [X0 : w] : (((accreln1 @ s__Permission @ cw @ X0) != $true) | ! [X1 : $i,X2 : $i] : (((s__instance @ X1 @ s__Parking) != $true) | ($true != (s__instance @ X2 @ s__ParkingLot)) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ((s__located @ X1 @ X2 @ X0) != $true)))),
  inference(ennf_transformation,[status(thm)],[f23])).
thf(f23,plain,(
  ~? [X0 : w] : (? [X1 : $i,X2 : $i] : (((s__instance @ X1 @ s__Parking) = $true) & ((s__located @ X1 @ X2 @ X0) = $true) & ($true = (s__instance @ X2 @ s__ParkingLot)) & ((s__agent @ X1 @ s__Bill @ X0) = $true)) & ((accreln1 @ s__Permission @ cw @ X0) = $true))),
  inference(fool_elimination,[status(thm)],[f22])).
thf(f22,plain,(
  ~? [X0 : w] : ((accreln1 @ s__Permission @ cw @ X0) & ? [X1 : $i,X2 : $i] : ((s__instance @ X2 @ s__ParkingLot) & (s__instance @ X1 @ s__Parking) & (s__located @ X1 @ X2 @ X0) & (s__agent @ X1 @ s__Bill @ X0)))),
  inference(rectify,[status(thm)],[f20])).
thf(f20,negated_conjecture,(
  ~? [X5 : w] : ((accreln1 @ s__Permission @ cw @ X5) & ? [X3 : $i,X4 : $i] : ((s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking) & (s__located @ X3 @ X4 @ X5) & (s__agent @ X3 @ s__Bill @ X5)))),
  inference(negated_conjecture,[status(cth)],[f19])).
thf(f19,conjecture,(
  ? [X5 : w] : ((accreln1 @ s__Permission @ cw @ X5) & ? [X3 : $i,X4 : $i] : ((s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking) & (s__located @ X3 @ X4 @ X5) & (s__agent @ X3 @ s__Bill @ X5)))),
  file('/tmp/b9RIdAuXMH/SOT_n2tfEp',conjecture_is_permitted)).
% SZS output end Proof for SOT_n2tfEp
% (3338267)------------------------------
% (3338267)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (3338267)Termination reason: Refutation

% (3338267)Memory used [KB]: 5628
% (3338267)Time elapsed: 0.006 s
% (3338267)------------------------------
% (3338267)------------------------------
% (3338258)Success in time 0.02 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_n2tfEp - Vampire---5.0.1 says Theorem - CPU = 0.00 WC = 0.06 
% OUTPUT: SOT_n2tfEp - Vampire---5.0.1 says Refutation - CPU = 0.00 WC = 0.06 

