% Checking upload ...
% Checker ran ...
% No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/B7lmNlObdh/SOT_O6ja4t
% (2605249)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_O6ja4t for (1ds/4Mi)
% (2605253)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_O6ja4t for (1ds/275Mi)
% (2605250)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_O6ja4t for (1ds/27Mi)
% (2605254)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_O6ja4t for (1ds/18Mi)
% (2605248)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_O6ja4t for (1ds/183Mi)
% (2605252)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_O6ja4t for (1ds/2Mi)

% (2605251)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_O6ja4t for (1ds/2Mi)
% (2605253)First to succeed.
% (2605252)Refutation not found, incomplete strategy
% (2605252)------------------------------
% (2605252)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2605252)Termination reason: Refutation not found, incomplete strategy


% (2605252)Memory used [KB]: 5500
% (2605252)Time elapsed: 0.004 s
% (2605252)------------------------------
% (2605252)------------------------------
% (2605253)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_O6ja4t
% SZS output start Proof for SOT_O6ja4t
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m &gt; w &gt; w &gt; $o).
thf(func_def_5, type, accreln2: m &gt; $i &gt; w &gt; w &gt; $o).
thf(func_def_6, type, accreln3: m &gt; $i &gt; $i &gt; w &gt; w &gt; $o).
thf(func_def_8, type, s__instance: $i &gt; $i &gt; $o).
thf(func_def_9, type, s__agent: $i &gt; $i &gt; w &gt; $o).
thf(func_def_10, type, s__patient: $i &gt; $i &gt; w &gt; $o).
thf(func_def_11, type, s__destination: $i &gt; $i &gt; w &gt; $o).
thf(func_def_12, type, s__located: $i &gt; $i &gt; w &gt; $o).
thf(func_def_13, type, s__possesses: $i &gt; $i &gt; w &gt; $o).
thf(func_def_23, type, sK0: w &gt; w).
thf(func_def_25, type, sK2: w &gt; $i).
thf(func_def_26, type, sK3: w &gt; $i).
thf(f68,plain,(
  $false),
  inference(subsumption_resolution,[status(thm)],[f67,f48])).
thf(f48,plain,(
  ((s__instance @ sK1 @ s__DriversLicense) = $true)),
  inference(cnf_transformation,[status(thm)],[f39])).
thf(f39,plain,(
  ($true = (s__possesses @ s__Bill @ sK1 @ cw)) & ((s__instance @ sK1 @ s__DriversLicense) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f26,f38])).
thf(f38,plain,(
  ? [X0 : $i] : (((s__possesses @ s__Bill @ X0 @ cw) = $true) & ($true = (s__instance @ X0 @ s__DriversLicense))) =&gt; (($true = (s__possesses @ s__Bill @ sK1 @ cw)) & ((s__instance @ sK1 @ s__DriversLicense) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f26,plain,(
  ? [X0 : $i] : (((s__possesses @ s__Bill @ X0 @ cw) = $true) & ($true = (s__instance @ X0 @ s__DriversLicense)))),
  inference(fool_elimination,[status(thm)],[f25])).
thf(f25,plain,(
  ? [X0 : $i] : ((s__possesses @ s__Bill @ X0 @ cw) & (s__instance @ X0 @ s__DriversLicense))),
  inference(rectify,[status(thm)],[f19])).
thf(f19,axiom,(
  ? [X2 : $i] : ((s__possesses @ s__Bill @ X2 @ cw) & (s__instance @ X2 @ s__DriversLicense))),
  file('/tmp/B7lmNlObdh/SOT_O6ja4t',fact_has_license)).
thf(f67,plain,(
  ((s__instance @ sK1 @ s__DriversLicense) != $true)),
  inference(trivial_inequality_removal,[status(thm)],[f66])).
thf(f66,plain,(
  ($true != $true) | ((s__instance @ sK1 @ s__DriversLicense) != $true)),
  inference(superposition,[status(thm)],[f65,f49])).
thf(f49,plain,(
  ($true = (s__possesses @ s__Bill @ sK1 @ cw))),
  inference(cnf_transformation,[status(thm)],[f39])).
thf(f65,plain,(
  ( ! [X0 : $i] : (((s__possesses @ s__Bill @ X0 @ cw) != $true) | ($true != (s__instance @ X0 @ s__DriversLicense))) )),
  inference(trivial_inequality_removal,[status(thm)],[f64])).
thf(f64,plain,(
  ( ! [X0 : $i] : (($true != $true) | ($true != (s__instance @ X0 @ s__DriversLicense)) | ((s__possesses @ s__Bill @ X0 @ cw) != $true)) )),
  inference(superposition,[status(thm)],[f63,f42])).
thf(f42,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Permission @ X0 @ (sK0 @ X0)))) )),
  inference(cnf_transformation,[status(thm)],[f37])).
thf(f37,plain,(
  ! [X0 : w] : ($true = (accreln1 @ s__Permission @ X0 @ (sK0 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f28,f36])).
thf(f36,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Permission @ X0 @ X1) = $true) =&gt; ($true = (accreln1 @ s__Permission @ X0 @ (sK0 @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f28,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Permission @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f27])).
thf(f27,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Permission @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Permission @ X0 @ X1)),
  file('/tmp/B7lmNlObdh/SOT_O6ja4t',serial_permission)).
thf(f63,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (accreln1 @ s__Permission @ X0 @ (sK0 @ cw))) | ($true != (s__instance @ X1 @ s__DriversLicense)) | ($true != (s__possesses @ s__Bill @ X1 @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f62,f53])).
thf(f53,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (($true = (s__agent @ (sK3 @ X0) @ s__Bill @ X0)) | ($true != (s__possesses @ s__Bill @ X2 @ X1)) | ((s__instance @ X2 @ s__DriversLicense) != $true) | ($true != (accreln1 @ s__Permission @ X1 @ X0))) )),
  inference(cnf_transformation,[status(thm)],[f41])).
thf(f41,plain,(
  ! [X0 : w,X1 : w,X2 : $i] : (($true != (s__possesses @ s__Bill @ X2 @ X1)) | ((s__instance @ X2 @ s__DriversLicense) != $true) | (($true = (s__instance @ (sK3 @ X0) @ s__Parking)) & ((s__located @ (sK3 @ X0) @ (sK2 @ X0) @ X0) = $true) & ($true = (s__agent @ (sK3 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK2 @ X0) @ s__ParkingLot))) | ($true != (accreln1 @ s__Permission @ X1 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f34,f40])).
thf(f40,plain,(
  ! [X0 : w] : (? [X3 : $i,X4 : $i] : (($true = (s__instance @ X4 @ s__Parking)) & ((s__located @ X4 @ X3 @ X0) = $true) & ($true = (s__agent @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X3 @ s__ParkingLot))) =&gt; (($true = (s__instance @ (sK3 @ X0) @ s__Parking)) & ((s__located @ (sK3 @ X0) @ (sK2 @ X0) @ X0) = $true) & ($true = (s__agent @ (sK3 @ X0) @ s__Bill @ X0)) & ($true = (s__instance @ (sK2 @ X0) @ s__ParkingLot))))),
  introduced(definition,[],[choice_axiom])).
thf(f34,plain,(
  ! [X0 : w,X1 : w,X2 : $i] : (($true != (s__possesses @ s__Bill @ X2 @ X1)) | ((s__instance @ X2 @ s__DriversLicense) != $true) | ? [X3 : $i,X4 : $i] : (($true = (s__instance @ X4 @ s__Parking)) & ((s__located @ X4 @ X3 @ X0) = $true) & ($true = (s__agent @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X3 @ s__ParkingLot))) | ($true != (accreln1 @ s__Permission @ X1 @ X0)))),
  inference(flattening,[status(thm)],[f33])).
thf(f33,plain,(
  ! [X0 : w,X2 : $i,X1 : w] : ((? [X3 : $i,X4 : $i] : (($true = (s__instance @ X4 @ s__Parking)) & ((s__located @ X4 @ X3 @ X0) = $true) & ($true = (s__agent @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X3 @ s__ParkingLot))) | ($true != (accreln1 @ s__Permission @ X1 @ X0))) | (($true != (s__possesses @ s__Bill @ X2 @ X1)) | ((s__instance @ X2 @ s__DriversLicense) != $true)))),
  inference(ennf_transformation,[status(thm)],[f30])).
thf(f30,plain,(
  ! [X0 : w,X2 : $i,X1 : w] : ((($true = (s__possesses @ s__Bill @ X2 @ X1)) & ((s__instance @ X2 @ s__DriversLicense) = $true)) =&gt; (($true = (accreln1 @ s__Permission @ X1 @ X0)) =&gt; ? [X3 : $i,X4 : $i] : (($true = (s__instance @ X4 @ s__Parking)) & ((s__located @ X4 @ X3 @ X0) = $true) & ($true = (s__agent @ X4 @ s__Bill @ X0)) & ($true = (s__instance @ X3 @ s__ParkingLot)))))),
  inference(fool_elimination,[status(thm)],[f29])).
thf(f29,plain,(
  ! [X0 : w,X1 : w,X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) & (s__possesses @ s__Bill @ X2 @ X1)) =&gt; ((accreln1 @ s__Permission @ X1 @ X0) =&gt; ? [X3 : $i,X4 : $i] : ((s__instance @ X3 @ s__ParkingLot) & (s__located @ X4 @ X3 @ X0) & (s__instance @ X4 @ s__Parking) & (s__agent @ X4 @ s__Bill @ X0))))),
  inference(rectify,[status(thm)],[f18])).
thf(f18,axiom,(
  ! [X1 : w,X0 : w,X2 : $i] : (((s__instance @ X2 @ s__DriversLicense) & (s__possesses @ s__Bill @ X2 @ X0)) =&gt; ((accreln1 @ s__Permission @ X0 @ X1) =&gt; ? [X4 : $i,X3 : $i] : ((s__instance @ X4 @ s__ParkingLot) & (s__located @ X3 @ X4 @ X1) & (s__instance @ X3 @ s__Parking) & (s__agent @ X3 @ s__Bill @ X1))))),
  file('/tmp/B7lmNlObdh/SOT_O6ja4t',permission1_to_park)).
thf(f62,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (accreln1 @ s__Permission @ X0 @ (sK0 @ cw))) | ($true != (s__instance @ X1 @ s__DriversLicense)) | ($true != (s__possesses @ s__Bill @ X1 @ X0)) | ($true != (s__agent @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) )),
  inference(subsumption_resolution,[status(thm)],[f61,f52])).
thf(f52,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (($true != (accreln1 @ s__Permission @ X1 @ X0)) | ($true != (s__possesses @ s__Bill @ X2 @ X1)) | ($true = (s__instance @ (sK2 @ X0) @ s__ParkingLot)) | ((s__instance @ X2 @ s__DriversLicense) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f41])).
thf(f61,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (accreln1 @ s__Permission @ X0 @ (sK0 @ cw))) | ($true != (s__instance @ X1 @ s__DriversLicense)) | ($true != (s__possesses @ s__Bill @ X1 @ X0)) | ($true != (s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot)) | ($true != (s__agent @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw)))) )),
  inference(subsumption_resolution,[status(thm)],[f60,f55])).
thf(f55,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (($true != (accreln1 @ s__Permission @ X1 @ X0)) | ((s__instance @ X2 @ s__DriversLicense) != $true) | ($true = (s__instance @ (sK3 @ X0) @ s__Parking)) | ($true != (s__possesses @ s__Bill @ X2 @ X1))) )),
  inference(cnf_transformation,[status(thm)],[f41])).
thf(f60,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__agent @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw))) | ($true != (accreln1 @ s__Permission @ X0 @ (sK0 @ cw))) | ($true != (s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot)) | ($true != (s__possesses @ s__Bill @ X1 @ X0)) | ($true != (s__instance @ X1 @ s__DriversLicense)) | ($true != (s__instance @ (sK3 @ (sK0 @ cw)) @ s__Parking))) )),
  inference(trivial_inequality_removal,[status(thm)],[f59])).
thf(f59,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != (s__instance @ X1 @ s__DriversLicense)) | ($true != (s__agent @ (sK3 @ (sK0 @ cw)) @ s__Bill @ (sK0 @ cw))) | ($true != (accreln1 @ s__Permission @ X0 @ (sK0 @ cw))) | ($true != (s__instance @ (sK3 @ (sK0 @ cw)) @ s__Parking)) | ($true != $true) | ($true != (s__possesses @ s__Bill @ X1 @ X0)) | ($true != (s__instance @ (sK2 @ (sK0 @ cw)) @ s__ParkingLot))) )),
  inference(superposition,[status(thm)],[f58,f54])).
thf(f54,plain,(
  ( ! [X2 : $i,X0 : w,X1 : w] : (((s__located @ (sK3 @ X0) @ (sK2 @ X0) @ X0) = $true) | ($true != (accreln1 @ s__Permission @ X1 @ X0)) | ((s__instance @ X2 @ s__DriversLicense) != $true) | ($true != (s__possesses @ s__Bill @ X2 @ X1))) )),
  inference(cnf_transformation,[status(thm)],[f41])).
thf(f58,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__located @ X0 @ X1 @ (sK0 @ cw))) | ($true != (s__agent @ X0 @ s__Bill @ (sK0 @ cw))) | ($true != (s__instance @ X0 @ s__Parking)) | ($true != (s__instance @ X1 @ s__ParkingLot))) )),
  inference(trivial_inequality_removal,[status(thm)],[f57])).
thf(f57,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__located @ X0 @ X1 @ (sK0 @ cw))) | ($true != (s__instance @ X0 @ s__Parking)) | ($true != (s__agent @ X0 @ s__Bill @ (sK0 @ cw))) | ($true != $true) | ($true != (s__instance @ X1 @ s__ParkingLot))) )),
  inference(superposition,[status(thm)],[f46,f42])).
thf(f46,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (accreln1 @ s__Permission @ cw @ X0)) | ((s__instance @ X1 @ s__Parking) != $true) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true != (s__located @ X1 @ X2 @ X0)) | ($true != (s__instance @ X2 @ s__ParkingLot))) )),
  inference(cnf_transformation,[status(thm)],[f35])).
thf(f35,plain,(
  ! [X0 : w] : (! [X1 : $i,X2 : $i] : (((s__instance @ X1 @ s__Parking) != $true) | ($true != (s__located @ X1 @ X2 @ X0)) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X2 @ s__ParkingLot))) | ($true != (accreln1 @ s__Permission @ cw @ X0)))),
  inference(ennf_transformation,[status(thm)],[f24])).
thf(f24,plain,(
  ~? [X0 : w] : (? [X1 : $i,X2 : $i] : (((s__instance @ X1 @ s__Parking) = $true) & ($true = (s__located @ X1 @ X2 @ X0)) & ($true = (s__instance @ X2 @ s__ParkingLot)) & ($true = (s__agent @ X1 @ s__Bill @ X0))) & ($true = (accreln1 @ s__Permission @ cw @ X0)))),
  inference(fool_elimination,[status(thm)],[f23])).
thf(f23,plain,(
  ~? [X0 : w] : ((accreln1 @ s__Permission @ cw @ X0) & ? [X1 : $i,X2 : $i] : ((s__located @ X1 @ X2 @ X0) & (s__agent @ X1 @ s__Bill @ X0) & (s__instance @ X2 @ s__ParkingLot) & (s__instance @ X1 @ s__Parking)))),
  inference(rectify,[status(thm)],[f21])).
thf(f21,negated_conjecture,(
  ~? [X5 : w] : ((accreln1 @ s__Permission @ cw @ X5) & ? [X3 : $i,X4 : $i] : ((s__located @ X3 @ X4 @ X5) & (s__agent @ X3 @ s__Bill @ X5) & (s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking)))),
  inference(negated_conjecture,[status(cth)],[f20])).
thf(f20,conjecture,(
  ? [X5 : w] : ((accreln1 @ s__Permission @ cw @ X5) & ? [X3 : $i,X4 : $i] : ((s__located @ X3 @ X4 @ X5) & (s__agent @ X3 @ s__Bill @ X5) & (s__instance @ X4 @ s__ParkingLot) & (s__instance @ X3 @ s__Parking)))),
  file('/tmp/B7lmNlObdh/SOT_O6ja4t',conjecture_is_permitted)).
% SZS output end Proof for SOT_O6ja4t
% (2605253)------------------------------
% (2605253)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2605253)Termination reason: Refutation

% (2605253)Memory used [KB]: 5500
% (2605253)Time elapsed: 0.007 s
% (2605253)------------------------------
% (2605253)------------------------------
% (2605244)Success in time 0.018 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_O6ja4t - Vampire---5.0.1 says Theorem - CPU = 0.00 WC = 0.06 
% OUTPUT: SOT_O6ja4t - Vampire---5.0.1 says Refutation - CPU = 0.00 WC = 0.06 
