 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/3allq8YZdk/SOT_hmwlcs
% (2151297)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_hmwlcs for (1ds/183Mi)
% (2151302)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_hmwlcs for (1ds/275Mi)
% (2151298)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_hmwlcs for (1ds/4Mi)
% (2151303)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_hmwlcs for (1ds/18Mi)
% (2151301)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_hmwlcs for (1ds/2Mi)
% (2151299)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_hmwlcs for (1ds/27Mi)
% (2151300)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_hmwlcs for (1ds/2Mi)
% (2151301)Refutation not found, incomplete strategy
% (2151301)------------------------------
% (2151301)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2151301)Termination reason: Refutation not found, incomplete strategy


% (2151301)Memory used [KB]: 5628
% (2151301)Time elapsed: 0.006 s
% (2151301)------------------------------
% (2151301)------------------------------
% (2151304)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_hmwlcs for (1ds/3Mi)
% (2151300)First to succeed.
% (2151300)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_hmwlcs
% SZS output start Proof for SOT_hmwlcs
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, accreln1: m > w > w > $o).
thf(func_def_5, type, accreln2: m > $i > w > w > $o).
thf(func_def_6, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_8, type, s__instance: $i > $i > $o).
thf(func_def_9, type, s__agent: $i > $i > w > $o).
thf(func_def_10, type, s__patient: $i > $i > w > $o).
thf(func_def_11, type, s__destination: $i > $i > w > $o).
thf(func_def_12, type, s__located: $i > $i > w > $o).
thf(func_def_24, type, sK0: w > $i).
thf(func_def_25, type, sK1: w > $i).
thf(func_def_26, type, sK2: w > w).
thf(func_def_27, type, sK3: w > $i).
thf(func_def_28, type, sK4: w > $i).
thf(func_def_30, type, sK6: w > w).
thf(func_def_31, type, sK7: w > $i).
thf(func_def_32, type, sK8: w > $i).
thf(f933,plain,(
  $false),
  inference(avatar_sat_refutation,[status(thm)],[f354,f356,f358,f360,f928,f932])).
thf(f932,plain,(
  ~spl9_13),
  inference(avatar_contradiction_clause,[status(thm)],[f931])).
thf(f931,plain,(
  $false | ~spl9_13),
  inference(subsumption_resolution,[status(thm)],[f930,f305])).
thf(f305,plain,(
  ($true = (s__instance @ (sK3 @ (sK2 @ (sK2 @ cw))) @ s__OfficeBuilding))),
  inference(subsumption_resolution,[status(thm)],[f304,f239])).
thf(f239,plain,(
  ($true = (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f234])).
thf(f234,plain,(
  ($true != $true) | ($true = (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw)))),
  inference(superposition,[status(thm)],[f162,f102])).
thf(f102,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Obligation @ X0 @ (sK2 @ X0)))) )),
  inference(cnf_transformation,[status(thm)],[f78])).
thf(f78,plain,(
  ! [X0 : w] : ($true = (accreln1 @ s__Obligation @ X0 @ (sK2 @ X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f59,f77])).
thf(f77,plain,(
  ! [X0 : w] : (? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ($true = (accreln1 @ s__Obligation @ X0 @ (sK2 @ X0))))),
  introduced(definition,[],[choice_axiom])).
thf(f59,plain,(
  ! [X0 : w] : ? [X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) = $true)),
  inference(fool_elimination,[status(thm)],[f58])).
thf(f58,plain,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1)),
  file('/tmp/3allq8YZdk/SOT_hmwlcs',serial_obligation)).
thf(f162,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ((s__agent @ (sK1 @ X0) @ s__Bill @ X0) = $true)) )),
  inference(subsumption_resolution,[status(thm)],[f161,f124])).
thf(f124,plain,(
  ((s__instance @ sK5 @ s__PayCheck) = $true)),
  inference(cnf_transformation,[status(thm)],[f83])).
thf(f83,plain,(
  ((s__instance @ sK5 @ s__PayCheck) = $true) & ($true = (s__destination @ sK5 @ s__Bill @ cw))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f61,f82])).
thf(f82,plain,(
  ? [X0 : $i] : (($true = (s__instance @ X0 @ s__PayCheck)) & ($true = (s__destination @ X0 @ s__Bill @ cw))) => (((s__instance @ sK5 @ s__PayCheck) = $true) & ($true = (s__destination @ sK5 @ s__Bill @ cw)))),
  introduced(definition,[],[choice_axiom])).
thf(f61,plain,(
  ? [X0 : $i] : (($true = (s__instance @ X0 @ s__PayCheck)) & ($true = (s__destination @ X0 @ s__Bill @ cw)))),
  inference(fool_elimination,[status(thm)],[f60])).
thf(f60,plain,(
  ? [X0 : $i] : ((s__destination @ X0 @ s__Bill @ cw) & (s__instance @ X0 @ s__PayCheck))),
  inference(rectify,[status(thm)],[f40])).
thf(f40,axiom,(
  ? [X3 : $i] : ((s__destination @ X3 @ s__Bill @ cw) & (s__instance @ X3 @ s__PayCheck))),
  file('/tmp/3allq8YZdk/SOT_hmwlcs',optionA_bill_gets_paycheck)).
thf(f161,plain,(
  ( ! [X0 : w] : (((s__instance @ sK5 @ s__PayCheck) != $true) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ((s__agent @ (sK1 @ X0) @ s__Bill @ X0) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f160])).
thf(f160,plain,(
  ( ! [X0 : w] : (((s__agent @ (sK1 @ X0) @ s__Bill @ X0) = $true) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != $true) | ((s__instance @ sK5 @ s__PayCheck) != $true)) )),
  inference(superposition,[status(thm)],[f97,f123])).
thf(f123,plain,(
  ($true = (s__destination @ sK5 @ s__Bill @ cw))),
  inference(cnf_transformation,[status(thm)],[f83])).
thf(f97,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ($true = (s__agent @ (sK1 @ X2) @ s__Bill @ X2))) )),
  inference(cnf_transformation,[status(thm)],[f76])).
thf(f76,plain,(
  ! [X0 : w,X1 : $i,X2 : w] : (((s__instance @ X1 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | (($true = (s__agent @ (sK1 @ X2) @ s__Bill @ X2)) & ($true = (s__instance @ (sK0 @ X2) @ s__Tax)) & ($true = (s__patient @ (sK1 @ X2) @ (sK0 @ X2) @ X2)) & ($true = (s__instance @ (sK1 @ X2) @ s__Giving))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f74,f75])).
thf(f75,plain,(
  ! [X2 : w] : (? [X3 : $i,X4 : $i] : (((s__agent @ X4 @ s__Bill @ X2) = $true) & ($true = (s__instance @ X3 @ s__Tax)) & ($true = (s__patient @ X4 @ X3 @ X2)) & ((s__instance @ X4 @ s__Giving) = $true)) => (($true = (s__agent @ (sK1 @ X2) @ s__Bill @ X2)) & ($true = (s__instance @ (sK0 @ X2) @ s__Tax)) & ($true = (s__patient @ (sK1 @ X2) @ (sK0 @ X2) @ X2)) & ($true = (s__instance @ (sK1 @ X2) @ s__Giving))))),
  introduced(definition,[],[choice_axiom])).
thf(f74,plain,(
  ! [X0 : w,X1 : $i,X2 : w] : (((s__instance @ X1 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ? [X3 : $i,X4 : $i] : (((s__agent @ X4 @ s__Bill @ X2) = $true) & ($true = (s__instance @ X3 @ s__Tax)) & ($true = (s__patient @ X4 @ X3 @ X2)) & ((s__instance @ X4 @ s__Giving) = $true)))),
  inference(rectify,[status(thm)],[f71])).
thf(f71,plain,(
  ! [X2 : w,X0 : $i,X1 : w] : (($true != (s__instance @ X0 @ s__PayCheck)) | ($true != (accreln1 @ s__Obligation @ X2 @ X1)) | ($true != (s__destination @ X0 @ s__Bill @ X2)) | ? [X4 : $i,X3 : $i] : (($true = (s__agent @ X3 @ s__Bill @ X1)) & ((s__instance @ X4 @ s__Tax) = $true) & ($true = (s__patient @ X3 @ X4 @ X1)) & ($true = (s__instance @ X3 @ s__Giving))))),
  inference(flattening,[status(thm)],[f70])).
thf(f70,plain,(
  ! [X0 : $i,X2 : w,X1 : w] : ((? [X4 : $i,X3 : $i] : (($true = (s__agent @ X3 @ s__Bill @ X1)) & ((s__instance @ X4 @ s__Tax) = $true) & ($true = (s__patient @ X3 @ X4 @ X1)) & ($true = (s__instance @ X3 @ s__Giving))) | ($true != (accreln1 @ s__Obligation @ X2 @ X1))) | (($true != (s__destination @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X0 @ s__PayCheck))))),
  inference(ennf_transformation,[status(thm)],[f63])).
thf(f63,plain,(
  ! [X0 : $i,X2 : w,X1 : w] : ((($true = (s__destination @ X0 @ s__Bill @ X2)) & ($true = (s__instance @ X0 @ s__PayCheck))) => (($true = (accreln1 @ s__Obligation @ X2 @ X1)) => ? [X4 : $i,X3 : $i] : (($true = (s__agent @ X3 @ s__Bill @ X1)) & ((s__instance @ X4 @ s__Tax) = $true) & ($true = (s__patient @ X3 @ X4 @ X1)) & ($true = (s__instance @ X3 @ s__Giving)))))),
  inference(fool_elimination,[status(thm)],[f62])).
thf(f62,plain,(
  ! [X0 : $i,X1 : w,X2 : w] : (((s__instance @ X0 @ s__PayCheck) & (s__destination @ X0 @ s__Bill @ X2)) => ((accreln1 @ s__Obligation @ X2 @ X1) => ? [X3 : $i,X4 : $i] : ((s__instance @ X4 @ s__Tax) & (s__patient @ X3 @ X4 @ X1) & (s__agent @ X3 @ s__Bill @ X1) & (s__instance @ X3 @ s__Giving))))),
  inference(rectify,[status(thm)],[f37])).
thf(f37,axiom,(
  ! [X3 : $i,X1 : w,X0 : w] : (((s__instance @ X3 @ s__PayCheck) & (s__destination @ X3 @ s__Bill @ X0)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4 : $i,X5 : $i] : ((s__instance @ X5 @ s__Tax) & (s__patient @ X4 @ X5 @ X1) & (s__agent @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Giving))))),
  file('/tmp/3allq8YZdk/SOT_hmwlcs',obligation1_pay_taxes)).
thf(f304,plain,(
  ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true = (s__instance @ (sK3 @ (sK2 @ (sK2 @ cw))) @ s__OfficeBuilding))),
  inference(subsumption_resolution,[status(thm)],[f303,f233])).
thf(f233,plain,(
  ($true = (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax))),
  inference(trivial_inequality_removal,[status(thm)],[f226])).
thf(f226,plain,(
  ($true = (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ($true != $true)),
  inference(superposition,[status(thm)],[f159,f102])).
thf(f159,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__instance @ (sK0 @ X0) @ s__Tax))) )),
  inference(subsumption_resolution,[status(thm)],[f158,f124])).
thf(f158,plain,(
  ( ! [X0 : w] : (($true = (s__instance @ (sK0 @ X0) @ s__Tax)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ((s__instance @ sK5 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f157])).
thf(f157,plain,(
  ( ! [X0 : w] : (((s__instance @ sK5 @ s__PayCheck) != $true) | ($true != $true) | ($true = (s__instance @ (sK0 @ X0) @ s__Tax)) | ($true != (accreln1 @ s__Obligation @ cw @ X0))) )),
  inference(superposition,[status(thm)],[f96,f123])).
thf(f96,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true = (s__instance @ (sK0 @ X2) @ s__Tax)) | ((s__instance @ X1 @ s__PayCheck) != $true) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f76])).
thf(f303,plain,(
  ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ($true = (s__instance @ (sK3 @ (sK2 @ (sK2 @ cw))) @ s__OfficeBuilding)) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw)))),
  inference(subsumption_resolution,[status(thm)],[f294,f201])).
thf(f201,plain,(
  ($true = (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f194])).
thf(f194,plain,(
  ($true = (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving)) | ($true != $true)),
  inference(superposition,[status(thm)],[f150,f102])).
thf(f150,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__instance @ (sK1 @ X0) @ s__Giving))) )),
  inference(subsumption_resolution,[status(thm)],[f149,f124])).
thf(f149,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ((s__instance @ sK5 @ s__PayCheck) != $true) | ($true = (s__instance @ (sK1 @ X0) @ s__Giving))) )),
  inference(trivial_inequality_removal,[status(thm)],[f148])).
thf(f148,plain,(
  ( ! [X0 : w] : (((s__instance @ sK5 @ s__PayCheck) != $true) | ($true = (s__instance @ (sK1 @ X0) @ s__Giving)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f94,f123])).
thf(f94,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true = (s__instance @ (sK1 @ X2) @ s__Giving)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f76])).
thf(f294,plain,(
  ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving)) | ($true = (s__instance @ (sK3 @ (sK2 @ (sK2 @ cw))) @ s__OfficeBuilding)) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax))),
  inference(trivial_inequality_removal,[status(thm)],[f292])).
thf(f292,plain,(
  ($true = (s__instance @ (sK3 @ (sK2 @ (sK2 @ cw))) @ s__OfficeBuilding)) | ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving)) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ($true != $true)),
  inference(superposition,[status(thm)],[f190,f248])).
thf(f248,plain,(
  ((s__patient @ (sK1 @ (sK2 @ cw)) @ (sK0 @ (sK2 @ cw)) @ (sK2 @ cw)) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f242])).
thf(f242,plain,(
  ((s__patient @ (sK1 @ (sK2 @ cw)) @ (sK0 @ (sK2 @ cw)) @ (sK2 @ cw)) = $true) | ($true != $true)),
  inference(superposition,[status(thm)],[f171,f102])).
thf(f171,plain,(
  ( ! [X0 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true = (s__patient @ (sK1 @ X0) @ (sK0 @ X0) @ X0))) )),
  inference(subsumption_resolution,[status(thm)],[f170,f124])).
thf(f170,plain,(
  ( ! [X0 : w] : (($true = (s__patient @ (sK1 @ X0) @ (sK0 @ X0) @ X0)) | ((s__instance @ sK5 @ s__PayCheck) != $true) | ($true != (accreln1 @ s__Obligation @ cw @ X0))) )),
  inference(trivial_inequality_removal,[status(thm)],[f169])).
thf(f169,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true = (s__patient @ (sK1 @ X0) @ (sK0 @ X0) @ X0)) | ($true != (accreln1 @ s__Obligation @ cw @ X0)) | ((s__instance @ sK5 @ s__PayCheck) != $true)) )),
  inference(superposition,[status(thm)],[f95,f123])).
thf(f95,plain,(
  ( ! [X2 : w,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true = (s__patient @ (sK1 @ X2) @ (sK0 @ X2) @ X2)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ((s__instance @ X1 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f76])).
thf(f190,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X2 @ X0)) | ((s__instance @ (sK3 @ (sK2 @ X0)) @ s__OfficeBuilding) = $true) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X2 @ s__Tax)) | ((s__instance @ X1 @ s__Giving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f186])).
thf(f186,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__instance @ (sK3 @ (sK2 @ X0)) @ s__OfficeBuilding) = $true) | ($true != (s__instance @ X2 @ s__Tax)) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__Giving) != $true) | ($true != $true) | ($true != (s__patient @ X1 @ X2 @ X0))) )),
  inference(superposition,[status(thm)],[f121,f102])).
thf(f121,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ($true != (s__instance @ X2 @ s__Giving)) | ((s__agent @ X2 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X3 @ s__Tax)) | ((s__patient @ X2 @ X3 @ X0) != $true) | ($true = (s__instance @ (sK3 @ X1) @ s__OfficeBuilding))) )),
  inference(cnf_transformation,[status(thm)],[f81])).
thf(f81,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (($true != (s__instance @ X2 @ s__Giving)) | ((s__agent @ X2 @ s__Bill @ X0) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | (($true = (s__instance @ (sK3 @ X1) @ s__OfficeBuilding)) & ((s__destination @ (sK4 @ X1) @ (sK3 @ X1) @ X1) = $true) & ($true = (s__instance @ (sK4 @ X1) @ s__Driving)) & ($true = (s__agent @ (sK4 @ X1) @ s__Bill @ X1))) | ((s__patient @ X2 @ X3 @ X0) != $true) | ($true != (s__instance @ X3 @ s__Tax)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4])],[f79,f80])).
thf(f80,plain,(
  ! [X1 : w] : (? [X4 : $i,X5 : $i] : (($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X1)) & ((s__instance @ X5 @ s__Driving) = $true) & ((s__agent @ X5 @ s__Bill @ X1) = $true)) => (($true = (s__instance @ (sK3 @ X1) @ s__OfficeBuilding)) & ((s__destination @ (sK4 @ X1) @ (sK3 @ X1) @ X1) = $true) & ($true = (s__instance @ (sK4 @ X1) @ s__Driving)) & ($true = (s__agent @ (sK4 @ X1) @ s__Bill @ X1))))),
  introduced(definition,[],[choice_axiom])).
thf(f79,plain,(
  ! [X0 : w,X1 : w,X2 : $i,X3 : $i] : (($true != (s__instance @ X2 @ s__Giving)) | ((s__agent @ X2 @ s__Bill @ X0) != $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ? [X4 : $i,X5 : $i] : (($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X1)) & ((s__instance @ X5 @ s__Driving) = $true) & ((s__agent @ X5 @ s__Bill @ X1) = $true)) | ((s__patient @ X2 @ X3 @ X0) != $true) | ($true != (s__instance @ X3 @ s__Tax)))),
  inference(rectify,[status(thm)],[f66])).
thf(f66,plain,(
  ! [X0 : w,X2 : w,X3 : $i,X1 : $i] : (($true != (s__instance @ X3 @ s__Giving)) | ($true != (s__agent @ X3 @ s__Bill @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ? [X4 : $i,X5 : $i] : (($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X2)) & ((s__instance @ X5 @ s__Driving) = $true) & ($true = (s__agent @ X5 @ s__Bill @ X2))) | ($true != (s__patient @ X3 @ X1 @ X0)) | ($true != (s__instance @ X1 @ s__Tax)))),
  inference(flattening,[status(thm)],[f65])).
thf(f65,plain,(
  ! [X0 : w,X1 : $i,X3 : $i,X2 : w] : ((? [X4 : $i,X5 : $i] : (($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X2)) & ((s__instance @ X5 @ s__Driving) = $true) & ($true = (s__agent @ X5 @ s__Bill @ X2))) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true)) | (($true != (s__instance @ X1 @ s__Tax)) | ($true != (s__instance @ X3 @ s__Giving)) | ($true != (s__patient @ X3 @ X1 @ X0)) | ($true != (s__agent @ X3 @ s__Bill @ X0))))),
  inference(ennf_transformation,[status(thm)],[f45])).
thf(f45,plain,(
  ! [X0 : w,X1 : $i,X3 : $i,X2 : w] : ((($true = (s__instance @ X1 @ s__Tax)) & ($true = (s__instance @ X3 @ s__Giving)) & ($true = (s__patient @ X3 @ X1 @ X0)) & ($true = (s__agent @ X3 @ s__Bill @ X0))) => (((accreln1 @ s__Obligation @ X0 @ X2) = $true) => ? [X4 : $i,X5 : $i] : (($true = (s__instance @ X4 @ s__OfficeBuilding)) & ($true = (s__destination @ X5 @ X4 @ X2)) & ((s__instance @ X5 @ s__Driving) = $true) & ($true = (s__agent @ X5 @ s__Bill @ X2)))))),
  inference(fool_elimination,[status(thm)],[f44])).
thf(f44,plain,(
  ! [X0 : w,X1 : $i,X2 : w,X3 : $i] : (((s__instance @ X3 @ s__Giving) & (s__instance @ X1 @ s__Tax) & (s__patient @ X3 @ X1 @ X0) & (s__agent @ X3 @ s__Bill @ X0)) => ((accreln1 @ s__Obligation @ X0 @ X2) => ? [X4 : $i,X5 : $i] : ((s__instance @ X4 @ s__OfficeBuilding) & (s__destination @ X5 @ X4 @ X2) & (s__agent @ X5 @ s__Bill @ X2) & (s__instance @ X5 @ s__Driving))))),
  inference(rectify,[status(thm)],[f38])).
thf(f38,axiom,(
  ! [X0 : w,X5 : $i,X1 : w,X4 : $i] : (((s__instance @ X4 @ s__Giving) & (s__instance @ X5 @ s__Tax) & (s__patient @ X4 @ X5 @ X0) & (s__agent @ X4 @ s__Bill @ X0)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X7 : $i,X6 : $i] : ((s__instance @ X7 @ s__OfficeBuilding) & (s__destination @ X6 @ X7 @ X1) & (s__agent @ X6 @ s__Bill @ X1) & (s__instance @ X6 @ s__Driving))))),
  file('/tmp/3allq8YZdk/SOT_hmwlcs',obligation2_drive_to_office)).
thf(f930,plain,(
  ($true != (s__instance @ (sK3 @ (sK2 @ (sK2 @ cw))) @ s__OfficeBuilding)) | ~spl9_13),
  inference(trivial_inequality_removal,[status(thm)],[f929])).
thf(f929,plain,(
  ($true != $true) | ($true != (s__instance @ (sK3 @ (sK2 @ (sK2 @ cw))) @ s__OfficeBuilding)) | ~spl9_13),
  inference(superposition,[status(thm)],[f353,f310])).
thf(f310,plain,(
  ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ (sK3 @ (sK2 @ (sK2 @ cw))) @ (sK2 @ (sK2 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f309,f239])).
thf(f309,plain,(
  ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ (sK3 @ (sK2 @ (sK2 @ cw))) @ (sK2 @ (sK2 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f308,f201])).
thf(f308,plain,(
  ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ (sK3 @ (sK2 @ (sK2 @ cw))) @ (sK2 @ (sK2 @ cw))) = $true) | ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving)) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw)))),
  inference(subsumption_resolution,[status(thm)],[f307,f233])).
thf(f307,plain,(
  ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ (sK3 @ (sK2 @ (sK2 @ cw))) @ (sK2 @ (sK2 @ cw))) = $true) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f306])).
thf(f306,plain,(
  ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving)) | ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ (sK3 @ (sK2 @ (sK2 @ cw))) @ (sK2 @ (sK2 @ cw))) = $true) | ($true != $true) | ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax))),
  inference(superposition,[status(thm)],[f222,f248])).
thf(f222,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X2 @ X0)) | ($true = (s__destination @ (sK4 @ (sK2 @ X0)) @ (sK3 @ (sK2 @ X0)) @ (sK2 @ X0))) | ($true != (s__instance @ X2 @ s__Tax)) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__Giving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f218])).
thf(f218,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true = (s__destination @ (sK4 @ (sK2 @ X0)) @ (sK3 @ (sK2 @ X0)) @ (sK2 @ X0))) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true != (s__patient @ X1 @ X2 @ X0)) | ($true != $true) | ((s__instance @ X1 @ s__Giving) != $true) | ($true != (s__instance @ X2 @ s__Tax))) )),
  inference(superposition,[status(thm)],[f120,f102])).
thf(f120,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((s__agent @ X2 @ s__Bill @ X0) != $true) | ((s__patient @ X2 @ X3 @ X0) != $true) | ($true != (s__instance @ X2 @ s__Giving)) | ((s__destination @ (sK4 @ X1) @ (sK3 @ X1) @ X1) = $true) | ($true != (s__instance @ X3 @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f81])).
thf(f353,plain,(
  ( ! [X0 : $i] : (($true != (s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X0 @ (sK2 @ (sK2 @ cw)))) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) ) | ~spl9_13),
  inference(avatar_component_clause,[status(thm)],[f352])).
thf(f352,definition,(
  spl9_13 <=> ! [X0 : $i] : (((s__instance @ X0 @ s__OfficeBuilding) != $true) | ($true != (s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X0 @ (sK2 @ (sK2 @ cw)))))),
  introduced(definition,[new_symbols(naming,[spl9_13])],[avatar_definition])).
thf(f928,plain,(
  ~spl9_9 | ~spl9_10 | ~spl9_11 | ~spl9_12),
  inference(avatar_contradiction_clause,[status(thm)],[f927])).
thf(f927,plain,(
  $false | (~spl9_9 | ~spl9_10 | ~spl9_11 | ~spl9_12)),
  inference(subsumption_resolution,[status(thm)],[f926,f477])).
thf(f477,plain,(
  ($true = (s__instance @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Keeping)) | ~spl9_12),
  inference(trivial_inequality_removal,[status(thm)],[f470])).
thf(f470,plain,(
  ($true = (s__instance @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Keeping)) | ($true != $true) | ~spl9_12),
  inference(superposition,[status(thm)],[f340,f143])).
thf(f143,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Permission @ X0 @ (sK2 @ X0)))) )),
  inference(trivial_inequality_removal,[status(thm)],[f142])).
thf(f142,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true = (accreln1 @ s__Permission @ X0 @ (sK2 @ X0)))) )),
  inference(superposition,[status(thm)],[f127,f102])).
thf(f127,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ((accreln1 @ s__Permission @ X0 @ X1) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f64])).
thf(f64,plain,(
  ! [X0 : w,X1 : w] : (((accreln1 @ s__Permission @ X0 @ X1) = $true) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true))),
  inference(ennf_transformation,[status(thm)],[f47])).
thf(f47,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) => ((accreln1 @ s__Permission @ X0 @ X1) = $true))),
  inference(fool_elimination,[status(thm)],[f46])).
thf(f46,plain,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  inference(rectify,[status(thm)],[f4])).
thf(f4,axiom,(
  ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => (accreln1 @ s__Permission @ X0 @ X1))),
  file('/tmp/3allq8YZdk/SOT_hmwlcs',bridge_obligation_permission)).
thf(f340,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ((s__instance @ (sK8 @ X0) @ s__Keeping) = $true)) ) | ~spl9_12),
  inference(avatar_component_clause,[status(thm)],[f339])).
thf(f339,definition,(
  spl9_12 <=> ! [X0 : w] : (((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ((s__instance @ (sK8 @ X0) @ s__Keeping) = $true))),
  introduced(definition,[new_symbols(naming,[spl9_12])],[avatar_definition])).
thf(f926,plain,(
  ($true != (s__instance @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Keeping)) | (~spl9_9 | ~spl9_10 | ~spl9_11)),
  inference(subsumption_resolution,[status(thm)],[f925,f463])).
thf(f463,plain,(
  ($true = (s__instance @ (sK7 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__ParkingLot)) | ~spl9_9),
  inference(trivial_inequality_removal,[status(thm)],[f458])).
thf(f458,plain,(
  ($true != $true) | ($true = (s__instance @ (sK7 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__ParkingLot)) | ~spl9_9),
  inference(superposition,[status(thm)],[f328,f143])).
thf(f328,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ($true = (s__instance @ (sK7 @ X0) @ s__ParkingLot))) ) | ~spl9_9),
  inference(avatar_component_clause,[status(thm)],[f327])).
thf(f327,definition,(
  spl9_9 <=> ! [X0 : w] : (($true = (s__instance @ (sK7 @ X0) @ s__ParkingLot)) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true))),
  introduced(definition,[new_symbols(naming,[spl9_9])],[avatar_definition])).
thf(f925,plain,(
  ($true != (s__instance @ (sK7 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__ParkingLot)) | ($true != (s__instance @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Keeping)) | (~spl9_10 | ~spl9_11)),
  inference(subsumption_resolution,[status(thm)],[f924,f547])).
thf(f547,plain,(
  ($true = (s__agent @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Bill @ (sK2 @ (sK2 @ (sK2 @ cw))))) | ~spl9_11),
  inference(trivial_inequality_removal,[status(thm)],[f542])).
thf(f542,plain,(
  ($true = (s__agent @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Bill @ (sK2 @ (sK2 @ (sK2 @ cw))))) | ($true != $true) | ~spl9_11),
  inference(superposition,[status(thm)],[f336,f143])).
thf(f336,plain,(
  ( ! [X1 : w] : (((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true) | ($true = (s__agent @ (sK8 @ X1) @ s__Bill @ X1))) ) | ~spl9_11),
  inference(avatar_component_clause,[status(thm)],[f335])).
thf(f335,definition,(
  spl9_11 <=> ! [X1 : w] : (($true = (s__agent @ (sK8 @ X1) @ s__Bill @ X1)) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true))),
  introduced(definition,[new_symbols(naming,[spl9_11])],[avatar_definition])).
thf(f924,plain,(
  ($true != (s__agent @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Bill @ (sK2 @ (sK2 @ (sK2 @ cw))))) | ($true != (s__instance @ (sK7 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__ParkingLot)) | ($true != (s__instance @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Keeping)) | ~spl9_10),
  inference(trivial_inequality_removal,[status(thm)],[f923])).
thf(f923,plain,(
  ($true != (s__instance @ (sK7 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__ParkingLot)) | ($true != (s__instance @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Keeping)) | ($true != (s__agent @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ s__Bill @ (sK2 @ (sK2 @ (sK2 @ cw))))) | ($true != $true) | ~spl9_10),
  inference(superposition,[status(thm)],[f288,f592])).
thf(f592,plain,(
  ($true = (s__located @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ (sK7 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ (sK2 @ (sK2 @ (sK2 @ cw))))) | ~spl9_10),
  inference(trivial_inequality_removal,[status(thm)],[f585])).
thf(f585,plain,(
  ($true = (s__located @ (sK8 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ (sK7 @ (sK2 @ (sK2 @ (sK2 @ cw)))) @ (sK2 @ (sK2 @ (sK2 @ cw))))) | ($true != $true) | ~spl9_10),
  inference(superposition,[status(thm)],[f332,f143])).
thf(f332,plain,(
  ( ! [X1 : w] : (((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true) | ($true = (s__located @ (sK8 @ X1) @ (sK7 @ X1) @ X1))) ) | ~spl9_10),
  inference(avatar_component_clause,[status(thm)],[f331])).
thf(f331,definition,(
  spl9_10 <=> ! [X1 : w] : (($true = (s__located @ (sK8 @ X1) @ (sK7 @ X1) @ X1)) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true))),
  introduced(definition,[new_symbols(naming,[spl9_10])],[avatar_definition])).
thf(f288,plain,(
  ( ! [X0 : $i,X1 : $i] : (($true != (s__located @ X0 @ X1 @ (sK2 @ (sK2 @ (sK2 @ cw))))) | ((s__instance @ X1 @ s__ParkingLot) != $true) | ((s__agent @ X0 @ s__Bill @ (sK2 @ (sK2 @ (sK2 @ cw)))) != $true) | ($true != (s__instance @ X0 @ s__Keeping))) )),
  inference(trivial_inequality_removal,[status(thm)],[f283])).
thf(f283,plain,(
  ( ! [X0 : $i,X1 : $i] : (((s__instance @ X1 @ s__ParkingLot) != $true) | ($true != $true) | ($true != (s__instance @ X0 @ s__Keeping)) | ((s__agent @ X0 @ s__Bill @ (sK2 @ (sK2 @ (sK2 @ cw)))) != $true) | ($true != (s__located @ X0 @ X1 @ (sK2 @ (sK2 @ (sK2 @ cw)))))) )),
  inference(superposition,[status(thm)],[f207,f154])).
thf(f154,plain,(
  ( ! [X0 : w] : (((accreln1 @ s__Permission @ X0 @ (sK2 @ (sK2 @ X0))) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f152])).
thf(f152,plain,(
  ( ! [X0 : w] : (($true != $true) | ((accreln1 @ s__Permission @ X0 @ (sK2 @ (sK2 @ X0))) = $true)) )),
  inference(superposition,[status(thm)],[f127,f147])).
thf(f147,plain,(
  ( ! [X0 : w] : (($true = (accreln1 @ s__Obligation @ X0 @ (sK2 @ (sK2 @ X0))))) )),
  inference(trivial_inequality_removal,[status(thm)],[f146])).
thf(f146,plain,(
  ( ! [X0 : w] : (($true != $true) | ($true = (accreln1 @ s__Obligation @ X0 @ (sK2 @ (sK2 @ X0))))) )),
  inference(superposition,[status(thm)],[f145,f102])).
thf(f145,plain,(
  ( ! [X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ (sK2 @ X0) @ X1)) | ((accreln1 @ s__Obligation @ X0 @ X1) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f144])).
thf(f144,plain,(
  ( ! [X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) = $true) | ($true != $true) | ($true != (accreln1 @ s__Obligation @ (sK2 @ X0) @ X1))) )),
  inference(superposition,[status(thm)],[f141,f102])).
thf(f141,plain,(
  ( ! [X2 : w,X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ X1 @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true) | ((accreln1 @ s__Obligation @ X1 @ X2) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f90])).
thf(f90,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X1 @ X2) = $true) | ($true != (accreln1 @ s__Obligation @ X1 @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X2) != $true))),
  inference(rectify,[status(thm)],[f73])).
thf(f73,plain,(
  ! [X0 : w,X2 : w,X1 : w] : (($true = (accreln1 @ s__Obligation @ X2 @ X1)) | ($true != (accreln1 @ s__Obligation @ X2 @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true))),
  inference(flattening,[status(thm)],[f72])).
thf(f72,plain,(
  ! [X1 : w,X2 : w,X0 : w] : (($true = (accreln1 @ s__Obligation @ X2 @ X1)) | (($true != (accreln1 @ s__Obligation @ X2 @ X0)) | ((accreln1 @ s__Obligation @ X0 @ X1) != $true)))),
  inference(ennf_transformation,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X1 : w,X2 : w,X0 : w] : ((($true = (accreln1 @ s__Obligation @ X2 @ X0)) & ((accreln1 @ s__Obligation @ X0 @ X1) = $true)) => ($true = (accreln1 @ s__Obligation @ X2 @ X1)))),
  inference(fool_elimination,[status(thm)],[f48])).
thf(f48,plain,(
  ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) & (accreln1 @ s__Obligation @ X2 @ X0)) => (accreln1 @ s__Obligation @ X2 @ X1))),
  inference(rectify,[status(thm)],[f2])).
thf(f2,axiom,(
  ! [X1 : w,X2 : w,X0 : w] : (((accreln1 @ s__Obligation @ X1 @ X2) & (accreln1 @ s__Obligation @ X0 @ X1)) => (accreln1 @ s__Obligation @ X0 @ X2))),
  file('/tmp/3allq8YZdk/SOT_hmwlcs',transitive_obligation)).
thf(f207,plain,(
  ( ! [X2 : $i,X0 : $i,X1 : w] : (($true != (accreln1 @ s__Permission @ (sK2 @ cw) @ X1)) | ((s__instance @ X2 @ s__Keeping) != $true) | ((s__located @ X2 @ X0 @ X1) != $true) | ($true != (s__instance @ X0 @ s__ParkingLot)) | ($true != (s__agent @ X2 @ s__Bill @ X1))) )),
  inference(trivial_inequality_removal,[status(thm)],[f202])).
thf(f202,plain,(
  ( ! [X2 : $i,X0 : $i,X1 : w] : (((s__located @ X2 @ X0 @ X1) != $true) | ($true != (s__agent @ X2 @ s__Bill @ X1)) | ($true != (s__instance @ X0 @ s__ParkingLot)) | ($true != $true) | ($true != (accreln1 @ s__Permission @ (sK2 @ cw) @ X1)) | ((s__instance @ X2 @ s__Keeping) != $true)) )),
  inference(superposition,[status(thm)],[f130,f102])).
thf(f130,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ($true != (s__instance @ X3 @ s__ParkingLot)) | ((accreln1 @ s__Permission @ X0 @ X1) != $true) | ((s__instance @ X2 @ s__Keeping) != $true) | ($true != (s__agent @ X2 @ s__Bill @ X1)) | ($true != (s__located @ X2 @ X3 @ X1))) )),
  inference(cnf_transformation,[status(thm)],[f86])).
thf(f86,plain,(
  ! [X0 : w,X1 : w] : (($true != (accreln1 @ s__Obligation @ cw @ X0)) | ! [X2 : $i,X3 : $i] : (($true != (s__instance @ X3 @ s__ParkingLot)) | ((s__instance @ X2 @ s__Keeping) != $true) | ($true != (s__agent @ X2 @ s__Bill @ X1)) | ($true != (s__located @ X2 @ X3 @ X1))) | ((accreln1 @ s__Permission @ X0 @ X1) != $true))),
  inference(rectify,[status(thm)],[f67])).
thf(f67,plain,(
  ! [X1 : w,X0 : w] : (((accreln1 @ s__Obligation @ cw @ X1) != $true) | ! [X3 : $i,X2 : $i] : (($true != (s__instance @ X2 @ s__ParkingLot)) | ((s__instance @ X3 @ s__Keeping) != $true) | ($true != (s__agent @ X3 @ s__Bill @ X0)) | ($true != (s__located @ X3 @ X2 @ X0))) | ($true != (accreln1 @ s__Permission @ X1 @ X0)))),
  inference(ennf_transformation,[status(thm)],[f53])).
thf(f53,plain,(
  ~? [X0 : w,X1 : w] : (($true = (accreln1 @ s__Permission @ X1 @ X0)) & ((accreln1 @ s__Obligation @ cw @ X1) = $true) & ? [X3 : $i,X2 : $i] : (($true = (s__instance @ X2 @ s__ParkingLot)) & ($true = (s__located @ X3 @ X2 @ X0)) & ((s__instance @ X3 @ s__Keeping) = $true) & ($true = (s__agent @ X3 @ s__Bill @ X0))))),
  inference(fool_elimination,[status(thm)],[f52])).
thf(f52,plain,(
  ~? [X0 : w,X1 : w] : ((accreln1 @ s__Permission @ X1 @ X0) & (accreln1 @ s__Obligation @ cw @ X1) & ? [X2 : $i,X3 : $i] : ((s__instance @ X3 @ s__Keeping) & (s__located @ X3 @ X2 @ X0) & (s__agent @ X3 @ s__Bill @ X0) & (s__instance @ X2 @ s__ParkingLot)))),
  inference(rectify,[status(thm)],[f42])).
thf(f42,negated_conjecture,(
  ~? [X11 : w,X10 : w] : ((accreln1 @ s__Permission @ X10 @ X11) & (accreln1 @ s__Obligation @ cw @ X10) & ? [X9 : $i,X8 : $i] : ((s__instance @ X8 @ s__Keeping) & (s__located @ X8 @ X9 @ X11) & (s__agent @ X8 @ s__Bill @ X11) & (s__instance @ X9 @ s__ParkingLot)))),
  inference(negated_conjecture,[status(cth)],[f41])).
thf(f41,conjecture,(
  ? [X11 : w,X10 : w] : ((accreln1 @ s__Permission @ X10 @ X11) & (accreln1 @ s__Obligation @ cw @ X10) & ? [X9 : $i,X8 : $i] : ((s__instance @ X8 @ s__Keeping) & (s__located @ X8 @ X9 @ X11) & (s__agent @ X8 @ s__Bill @ X11) & (s__instance @ X9 @ s__ParkingLot)))),
  file('/tmp/3allq8YZdk/SOT_hmwlcs',conjecture_is_permitted)).
thf(f360,plain,(
  spl9_11 | spl9_13),
  inference(avatar_split_clause,[status(thm)],[f359,f352,f335])).
thf(f359,plain,(
  ( ! [X0 : $i,X1 : w] : (($true != (s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X0 @ (sK2 @ (sK2 @ cw)))) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true) | ($true = (s__agent @ (sK8 @ X1) @ s__Bill @ X1)) | ((s__instance @ X0 @ s__OfficeBuilding) != $true)) )),
  inference(subsumption_resolution,[status(thm)],[f346,f299])).
thf(f299,plain,(
  ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) = $true)),
  inference(subsumption_resolution,[status(thm)],[f298,f201])).
thf(f298,plain,(
  ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving)) | ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) = $true)),
  inference(subsumption_resolution,[status(thm)],[f297,f233])).
thf(f297,plain,(
  ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) = $true) | ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving))),
  inference(subsumption_resolution,[status(thm)],[f296,f239])).
thf(f296,plain,(
  ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) = $true) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f293])).
thf(f293,plain,(
  ($true != $true) | ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving)) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) = $true)),
  inference(superposition,[status(thm)],[f177,f248])).
thf(f177,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X2 @ X0)) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true = (s__instance @ (sK4 @ (sK2 @ X0)) @ s__Driving)) | ($true != (s__instance @ X2 @ s__Tax)) | ((s__instance @ X1 @ s__Giving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f172])).
thf(f172,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__agent @ X1 @ s__Bill @ X0)) | ($true = (s__instance @ (sK4 @ (sK2 @ X0)) @ s__Driving)) | ($true != (s__patient @ X1 @ X2 @ X0)) | ($true != $true) | ($true != (s__instance @ X2 @ s__Tax)) | ((s__instance @ X1 @ s__Giving) != $true)) )),
  inference(superposition,[status(thm)],[f119,f102])).
thf(f119,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ($true = (s__instance @ (sK4 @ X1) @ s__Driving)) | ((s__patient @ X2 @ X3 @ X0) != $true) | ((s__agent @ X2 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X3 @ s__Tax)) | ($true != (s__instance @ X2 @ s__Giving))) )),
  inference(cnf_transformation,[status(thm)],[f81])).
thf(f346,plain,(
  ( ! [X0 : $i,X1 : w] : (((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ($true != (s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X0 @ (sK2 @ (sK2 @ cw)))) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true) | ($true = (s__agent @ (sK8 @ X1) @ s__Bill @ X1))) )),
  inference(trivial_inequality_removal,[status(thm)],[f343])).
thf(f343,plain,(
  ( ! [X0 : $i,X1 : w] : (($true = (s__agent @ (sK8 @ X1) @ s__Bill @ X1)) | ($true != (s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X0 @ (sK2 @ (sK2 @ cw)))) | ($true != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) != $true) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true)) )),
  inference(superposition,[status(thm)],[f139,f302])).
thf(f302,plain,(
  ((s__agent @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Bill @ (sK2 @ (sK2 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f301,f239])).
thf(f301,plain,(
  ((s__agent @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Bill @ (sK2 @ (sK2 @ cw))) = $true) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw)))),
  inference(subsumption_resolution,[status(thm)],[f300,f233])).
thf(f300,plain,(
  ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ((s__agent @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Bill @ (sK2 @ (sK2 @ cw))) = $true)),
  inference(subsumption_resolution,[status(thm)],[f295,f201])).
thf(f295,plain,(
  ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving)) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ((s__agent @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Bill @ (sK2 @ (sK2 @ cw))) = $true)),
  inference(trivial_inequality_removal,[status(thm)],[f291])).
thf(f291,plain,(
  ($true != (s__instance @ (sK0 @ (sK2 @ cw)) @ s__Tax)) | ($true != $true) | ((s__agent @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Bill @ (sK2 @ (sK2 @ cw))) = $true) | ($true != (s__agent @ (sK1 @ (sK2 @ cw)) @ s__Bill @ (sK2 @ cw))) | ($true != (s__instance @ (sK1 @ (sK2 @ cw)) @ s__Giving))),
  inference(superposition,[status(thm)],[f216,f248])).
thf(f216,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__patient @ X1 @ X2 @ X0)) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__agent @ (sK4 @ (sK2 @ X0)) @ s__Bill @ (sK2 @ X0)) = $true) | ($true != (s__instance @ X2 @ s__Tax))) )),
  inference(trivial_inequality_removal,[status(thm)],[f210])).
thf(f210,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__instance @ X2 @ s__Tax)) | ($true != (s__agent @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__Giving) != $true) | ((s__agent @ (sK4 @ (sK2 @ X0)) @ s__Bill @ (sK2 @ X0)) = $true) | ($true != $true) | ($true != (s__patient @ X1 @ X2 @ X0))) )),
  inference(superposition,[status(thm)],[f118,f102])).
thf(f118,plain,(
  ( ! [X2 : $i,X3 : $i,X0 : w,X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1) != $true) | ($true != (s__instance @ X2 @ s__Giving)) | ($true = (s__agent @ (sK4 @ X1) @ s__Bill @ X1)) | ((s__patient @ X2 @ X3 @ X0) != $true) | ($true != (s__instance @ X3 @ s__Tax)) | ((s__agent @ X2 @ s__Bill @ X0) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f81])).
thf(f139,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (($true != (s__agent @ X0 @ s__Bill @ X1)) | ((s__destination @ X0 @ X2 @ X1) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ($true = (s__agent @ (sK8 @ X3) @ s__Bill @ X3)) | ((accreln1 @ s__Permission @ X1 @ X3) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f89])).
thf(f89,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((accreln1 @ s__Permission @ X1 @ X3) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | (($true = (s__agent @ (sK8 @ X3) @ s__Bill @ X3)) & ((s__instance @ (sK8 @ X3) @ s__Keeping) = $true) & ($true = (s__instance @ (sK7 @ X3) @ s__ParkingLot)) & ($true = (s__located @ (sK8 @ X3) @ (sK7 @ X3) @ X3))) | ((s__destination @ X0 @ X2 @ X1) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ($true != (s__agent @ X0 @ s__Bill @ X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f87,f88])).
thf(f88,plain,(
  ! [X3 : w] : (? [X4 : $i,X5 : $i] : (($true = (s__agent @ X5 @ s__Bill @ X3)) & ($true = (s__instance @ X5 @ s__Keeping)) & ($true = (s__instance @ X4 @ s__ParkingLot)) & ($true = (s__located @ X5 @ X4 @ X3))) => (($true = (s__agent @ (sK8 @ X3) @ s__Bill @ X3)) & ((s__instance @ (sK8 @ X3) @ s__Keeping) = $true) & ($true = (s__instance @ (sK7 @ X3) @ s__ParkingLot)) & ($true = (s__located @ (sK8 @ X3) @ (sK7 @ X3) @ X3))))),
  introduced(definition,[],[choice_axiom])).
thf(f87,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((accreln1 @ s__Permission @ X1 @ X3) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ? [X4 : $i,X5 : $i] : (($true = (s__agent @ X5 @ s__Bill @ X3)) & ($true = (s__instance @ X5 @ s__Keeping)) & ($true = (s__instance @ X4 @ s__ParkingLot)) & ($true = (s__located @ X5 @ X4 @ X3))) | ((s__destination @ X0 @ X2 @ X1) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ($true != (s__agent @ X0 @ s__Bill @ X1)))),
  inference(rectify,[status(thm)],[f69])).
thf(f69,plain,(
  ! [X0 : $i,X2 : w,X3 : $i,X1 : w] : (((accreln1 @ s__Permission @ X2 @ X1) != $true) | ($true != (s__instance @ X3 @ s__OfficeBuilding)) | ? [X4 : $i,X5 : $i] : (((s__agent @ X5 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X5 @ s__Keeping)) & ($true = (s__instance @ X4 @ s__ParkingLot)) & ($true = (s__located @ X5 @ X4 @ X1))) | ($true != (s__destination @ X0 @ X3 @ X2)) | ((s__instance @ X0 @ s__Driving) != $true) | ($true != (s__agent @ X0 @ s__Bill @ X2)))),
  inference(flattening,[status(thm)],[f68])).
thf(f68,plain,(
  ! [X3 : $i,X1 : w,X0 : $i,X2 : w] : ((? [X4 : $i,X5 : $i] : (((s__agent @ X5 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X5 @ s__Keeping)) & ($true = (s__instance @ X4 @ s__ParkingLot)) & ($true = (s__located @ X5 @ X4 @ X1))) | ((accreln1 @ s__Permission @ X2 @ X1) != $true)) | (((s__instance @ X0 @ s__Driving) != $true) | ($true != (s__destination @ X0 @ X3 @ X2)) | ($true != (s__agent @ X0 @ s__Bill @ X2)) | ($true != (s__instance @ X3 @ s__OfficeBuilding))))),
  inference(ennf_transformation,[status(thm)],[f57])).
thf(f57,plain,(
  ! [X3 : $i,X1 : w,X0 : $i,X2 : w] : ((((s__instance @ X0 @ s__Driving) = $true) & ($true = (s__destination @ X0 @ X3 @ X2)) & ($true = (s__agent @ X0 @ s__Bill @ X2)) & ($true = (s__instance @ X3 @ s__OfficeBuilding))) => (((accreln1 @ s__Permission @ X2 @ X1) = $true) => ? [X4 : $i,X5 : $i] : (((s__agent @ X5 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X5 @ s__Keeping)) & ($true = (s__instance @ X4 @ s__ParkingLot)) & ($true = (s__located @ X5 @ X4 @ X1)))))),
  inference(fool_elimination,[status(thm)],[f56])).
thf(f56,plain,(
  ! [X0 : $i,X1 : w,X2 : w,X3 : $i] : (((s__destination @ X0 @ X3 @ X2) & (s__instance @ X0 @ s__Driving) & (s__agent @ X0 @ s__Bill @ X2) & (s__instance @ X3 @ s__OfficeBuilding)) => ((accreln1 @ s__Permission @ X2 @ X1) => ? [X4 : $i,X5 : $i] : ((s__instance @ X5 @ s__Keeping) & (s__instance @ X4 @ s__ParkingLot) & (s__agent @ X5 @ s__Bill @ X1) & (s__located @ X5 @ X4 @ X1))))),
  inference(rectify,[status(thm)],[f39])).
thf(f39,axiom,(
  ! [X6 : $i,X1 : w,X0 : w,X7 : $i] : (((s__destination @ X6 @ X7 @ X0) & (s__instance @ X6 @ s__Driving) & (s__agent @ X6 @ s__Bill @ X0) & (s__instance @ X7 @ s__OfficeBuilding)) => ((accreln1 @ s__Permission @ X0 @ X1) => ? [X9 : $i,X8 : $i] : ((s__instance @ X8 @ s__Keeping) & (s__instance @ X9 @ s__ParkingLot) & (s__agent @ X8 @ s__Bill @ X1) & (s__located @ X8 @ X9 @ X1))))),
  file('/tmp/3allq8YZdk/SOT_hmwlcs',permission_park_in_lot)).
thf(f358,plain,(
  spl9_13 | spl9_9),
  inference(avatar_split_clause,[status(thm)],[f357,f327,f352])).
thf(f357,plain,(
  ( ! [X0 : w,X1 : $i] : (($true = (s__instance @ (sK7 @ X0) @ s__ParkingLot)) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ((s__instance @ X1 @ s__OfficeBuilding) != $true) | ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X1 @ (sK2 @ (sK2 @ cw))) != $true)) )),
  inference(subsumption_resolution,[status(thm)],[f347,f299])).
thf(f347,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X1 @ (sK2 @ (sK2 @ cw))) != $true) | ($true = (s__instance @ (sK7 @ X0) @ s__ParkingLot)) | ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f345])).
thf(f345,plain,(
  ( ! [X0 : w,X1 : $i] : (($true != $true) | ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X1 @ (sK2 @ (sK2 @ cw))) != $true) | ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) != $true) | ((s__instance @ X1 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ($true = (s__instance @ (sK7 @ X0) @ s__ParkingLot))) )),
  inference(superposition,[status(thm)],[f137,f302])).
thf(f137,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (($true != (s__agent @ X0 @ s__Bill @ X1)) | ($true = (s__instance @ (sK7 @ X3) @ s__ParkingLot)) | ((s__instance @ X0 @ s__Driving) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Permission @ X1 @ X3) != $true) | ((s__destination @ X0 @ X2 @ X1) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f89])).
thf(f356,plain,(
  spl9_13 | spl9_12),
  inference(avatar_split_clause,[status(thm)],[f355,f339,f352])).
thf(f355,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__instance @ X1 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X1 @ (sK2 @ (sK2 @ cw))) != $true) | ((s__instance @ (sK8 @ X0) @ s__Keeping) = $true)) )),
  inference(subsumption_resolution,[status(thm)],[f348,f299])).
thf(f348,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) != $true) | ((s__instance @ X1 @ s__OfficeBuilding) != $true) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X1 @ (sK2 @ (sK2 @ cw))) != $true) | ((s__instance @ (sK8 @ X0) @ s__Keeping) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f344])).
thf(f344,plain,(
  ( ! [X0 : w,X1 : $i] : (((s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X1 @ (sK2 @ (sK2 @ cw))) != $true) | ($true != $true) | ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) != $true) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X0) != $true) | ((s__instance @ X1 @ s__OfficeBuilding) != $true) | ((s__instance @ (sK8 @ X0) @ s__Keeping) = $true)) )),
  inference(superposition,[status(thm)],[f138,f302])).
thf(f138,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (($true != (s__agent @ X0 @ s__Bill @ X1)) | ((s__instance @ (sK8 @ X3) @ s__Keeping) = $true) | ((s__destination @ X0 @ X2 @ X1) != $true) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ((accreln1 @ s__Permission @ X1 @ X3) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f89])).
thf(f354,plain,(
  spl9_13 | spl9_10),
  inference(avatar_split_clause,[status(thm)],[f350,f331,f352])).
thf(f350,plain,(
  ( ! [X0 : $i,X1 : w] : (($true = (s__located @ (sK8 @ X1) @ (sK7 @ X1) @ X1)) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ($true != (s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X0 @ (sK2 @ (sK2 @ cw))))) )),
  inference(subsumption_resolution,[status(thm)],[f349,f299])).
thf(f349,plain,(
  ( ! [X0 : $i,X1 : w] : (((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ($true = (s__located @ (sK8 @ X1) @ (sK7 @ X1) @ X1)) | ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) != $true) | ($true != (s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X0 @ (sK2 @ (sK2 @ cw))))) )),
  inference(trivial_inequality_removal,[status(thm)],[f342])).
thf(f342,plain,(
  ( ! [X0 : $i,X1 : w] : (($true = (s__located @ (sK8 @ X1) @ (sK7 @ X1) @ X1)) | ((s__instance @ (sK4 @ (sK2 @ (sK2 @ cw))) @ s__Driving) != $true) | ((accreln1 @ s__Permission @ (sK2 @ (sK2 @ cw)) @ X1) != $true) | ($true != (s__destination @ (sK4 @ (sK2 @ (sK2 @ cw))) @ X0 @ (sK2 @ (sK2 @ cw)))) | ((s__instance @ X0 @ s__OfficeBuilding) != $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f136,f302])).
thf(f136,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (($true != (s__agent @ X0 @ s__Bill @ X1)) | ((s__instance @ X2 @ s__OfficeBuilding) != $true) | ((s__instance @ X0 @ s__Driving) != $true) | ($true = (s__located @ (sK8 @ X3) @ (sK7 @ X3) @ X3)) | ((s__destination @ X0 @ X2 @ X1) != $true) | ((accreln1 @ s__Permission @ X1 @ X3) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f89])).
% SZS output end Proof for SOT_hmwlcs
% (2151300)------------------------------
% (2151300)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (2151300)Termination reason: Refutation

% (2151300)Memory used [KB]: 6396
% (2151300)Time elapsed: 0.096 s
% (2151300)------------------------------
% (2151300)------------------------------
% (2151293)Success in time 0.111 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_hmwlcs - Vampire---5.0.1 says Theorem - CPU = 0.02 WC = 0.16 
% OUTPUT: SOT_hmwlcs - Vampire---5.0.1 says Refutation - CPU = 0.02 WC = 0.16 

