% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 1h_Obligation_UO1_univConj
% SZS output start Proof for 1h_Obligation_UO1_univConj
1. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1) [input(axiom)]
3. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X0 @ X2) & (accreln1 @ s__Obligation @ X0 @ X1)) => (accreln1 @ s__Obligation @ X1 @ X2)) [input(axiom)]
20. ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X2,X3] : ((s__destination @ X2 @ X3 @ X1) & (s__agent @ X2 @ s__Bill @ X1) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying))) [input(axiom)]
21. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [input(axiom)]
23. ! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [input(conjecture)]
24. ~! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [negated conjecture 23]
25. ! [X0 : w] : ? [X1 : w] : (accreln1 @ s__Obligation @ X0 @ X1) [rectify 1]
26. ! [X0 : w] : ? [X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) [fool elimination 25]
29. ! [X0 : w,X1 : w,X2 : w] : (((accreln1 @ s__Obligation @ X0 @ X2) & (accreln1 @ s__Obligation @ X0 @ X1)) => (accreln1 @ s__Obligation @ X1 @ X2)) [rectify 3]
30. ! [X0 : w,X1 : w,X2 : w] : (((((accreln1 @ s__Obligation @ X0 @ X2)) = $true) & (((accreln1 @ s__Obligation @ X0 @ X1)) = $true)) => (((accreln1 @ s__Obligation @ X1 @ X2)) = $true)) [fool elimination 29]
33. ! [X0 : w,X1 : w] : ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X2,X3] : ((s__destination @ X2 @ X3 @ X1) & (s__agent @ X2 @ s__Bill @ X1) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying))) [rectify 20]
34. ! [X0 : w,X1 : w] : ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X2,X3] : ((((s__destination @ X2 @ X3 @ X1)) = $true) & (((s__agent @ X2 @ s__Bill @ X1)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true))) [fool elimination 33]
35. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [rectify 21]
36. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [fool elimination 35]
39. ~! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [rectify 24]
40. ~! [X0 : w] : ((((accreln1 @ s__Obligation @ cw @ X0)) = $true) => ? [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) = $true) & (((s__patient @ X1 @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__PayCheck)) = $true) & (((s__instance @ X1 @ s__Giving)) = $true))) [fool elimination 39]
56. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [flattening 36]
59. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Obligation @ X1 @ X2)) = $true) | ((((accreln1 @ s__Obligation @ X0 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true))) [ennf transformation 30]
60. ! [X0 : w,X1 : w,X2 : w] : ((((accreln1 @ s__Obligation @ X1 @ X2)) = $true) | (((accreln1 @ s__Obligation @ X0 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) [flattening 59]
61. ! [X0 : w,X1 : w] : (? [X2,X3] : ((((s__destination @ X2 @ X3 @ X1)) = $true) & (((s__agent @ X2 @ s__Bill @ X1)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) [ennf transformation 34]
62. ! [X0 : w,X1 : w,X2,X3] : ((? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true))) [ennf transformation 56]
63. ! [X0 : w,X1 : w,X2,X3] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [flattening 62]
64. ? [X0 : w] : (! [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) [ennf transformation 40]
65. ! [X0 : w] : (? [X1 : w] : (((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0))))) [skolem symbol introduction]
66. ! [X0 : w] : ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)))) [skolemisation 26,65]
67. ! [X1 : w] : (? [X2,X3] : ((((s__destination @ X2 @ X3 @ X1)) = $true) & (((s__agent @ X2 @ s__Bill @ X1)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => (($true = ((s__destination @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__Moon))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Flying))))) [skolem symbol introduction]
68. ! [X0 : w,X1 : w] : ((($true = ((s__destination @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) & ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) & ($true = ((s__instance @ (sK2 @ X1) @ s__Moon))) & ($true = ((s__instance @ (sK1 @ X1) @ s__Flying)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) [skolemisation 61,67]
69. ! [X1 : w] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) => (($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) & ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))))) [skolem symbol introduction]
70. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) & ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK3 @ X1) @ s__Giving)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [skolemisation 63,69]
73. ? [X0 : w] : (! [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) => (! [X2,X1] : (($true != ((s__destination @ X1 @ s__Bill @ sK7))) | ($true != ((s__patient @ X1 @ X2 @ sK7))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & ($true = ((accreln1 @ s__Obligation @ cw @ sK7)))) [skolem symbol introduction]
74. ! [X1,X2] : (($true != ((s__destination @ X1 @ s__Bill @ sK7))) | ($true != ((s__patient @ X1 @ X2 @ sK7))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & ($true = ((accreln1 @ s__Obligation @ cw @ sK7))) [skolemisation 64,73]
75. ($true = ((accreln1 @ s__Obligation @ X0 @ (sK0 @ X0)))) [cnf transformation 66]
77. (((accreln1 @ s__Obligation @ X0 @ X2)) != $true) | (((accreln1 @ s__Obligation @ X1 @ X2)) = $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) [cnf transformation 60]
94. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK1 @ X1) @ s__Flying))) [cnf transformation 68]
95. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK2 @ X1) @ s__Moon))) [cnf transformation 68]
96. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__agent @ (sK1 @ X1) @ s__Bill @ X1))) [cnf transformation 68]
97. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK1 @ X1) @ (sK2 @ X1) @ X1))) [cnf transformation 68]
98. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
99. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
100. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
101. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 70]
106. ($true = ((accreln1 @ s__Obligation @ cw @ sK7))) [cnf transformation 74]
107. ($true != ((s__destination @ X1 @ s__Bill @ sK7))) | ($true != ((s__patient @ X1 @ X2 @ sK7))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) [cnf transformation 74]
109. ($true = X0) | ($false = X0) [fool exhaustiveness axiom]
113. ($true != $true) | ($true != ((s__patient @ X0 @ X1 @ sK7))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK7))) [superposition 107,109]
114. ($true != ((s__patient @ X0 @ X1 @ sK7))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK7))) [trivial inequality removal 113]
115. ($true != $true) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK7))) | ($false = ((s__patient @ X0 @ X1 @ sK7))) [superposition 114,109]
116. ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK7))) | ($false = ((s__patient @ X0 @ X1 @ sK7))) [trivial inequality removal 115]
117. ($true != $true) | (((s__instance @ X1 @ s__Giving)) != $true) | (((s__destination @ X1 @ s__Bill @ sK7)) = $false) | ($false = ((s__patient @ X1 @ X0 @ sK7))) | ($false = ((s__instance @ X0 @ s__PayCheck))) [superposition 116,109]
118. (((s__instance @ X1 @ s__Giving)) != $true) | (((s__destination @ X1 @ s__Bill @ sK7)) = $false) | ($false = ((s__patient @ X1 @ X0 @ sK7))) | ($false = ((s__instance @ X0 @ s__PayCheck))) [trivial inequality removal 117]
119. ($true != $true) | ($false = ((s__destination @ X0 @ s__Bill @ sK7))) | ($false = ((s__patient @ X0 @ X1 @ sK7))) | ($false = ((s__instance @ X1 @ s__PayCheck))) | ($false = ((s__instance @ X0 @ s__Giving))) [superposition 118,109]
120. ($false = ((s__patient @ X0 @ X1 @ sK7))) | ($false = ((s__destination @ X0 @ s__Bill @ sK7))) | ($false = ((s__instance @ X1 @ s__PayCheck))) | ($false = ((s__instance @ X0 @ s__Giving))) [trivial inequality removal 119]
553. ($true != $true) | ($true = ((s__instance @ (sK1 @ sK7) @ s__Flying))) [superposition 94,106]
554. ($true != $true) | ($true = ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [superposition 94,75]
570. ($true = ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) [trivial inequality removal 554]
571. ($true = ((s__instance @ (sK1 @ sK7) @ s__Flying))) [trivial inequality removal 553]
606. ($true != $true) | ($true = ((s__instance @ (sK2 @ sK7) @ s__Moon))) [superposition 95,106]
607. ($true != $true) | ($true = ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) [superposition 95,75]
623. ($true = ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) [trivial inequality removal 607]
624. ($true = ((s__instance @ (sK2 @ sK7) @ s__Moon))) [trivial inequality removal 606]
653. ($true != $true) | ($true = ((s__agent @ (sK1 @ sK7) @ s__Bill @ sK7))) [superposition 96,106]
654. ($true != $true) | ($true = ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) [superposition 96,75]
670. ($true = ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) [trivial inequality removal 654]
671. ($true = ((s__agent @ (sK1 @ sK7) @ s__Bill @ sK7))) [trivial inequality removal 653]
700. ($true != $true) | ($true = ((s__destination @ (sK1 @ sK7) @ (sK2 @ sK7) @ sK7))) [superposition 97,106]
701. ($true != $true) | ($true = ((s__destination @ (sK1 @ (sK0 @ X0)) @ (sK2 @ (sK0 @ X0)) @ (sK0 @ X0)))) [superposition 97,75]
717. ($true = ((s__destination @ (sK1 @ (sK0 @ X0)) @ (sK2 @ (sK0 @ X0)) @ (sK0 @ X0)))) [trivial inequality removal 701]
718. ($true = ((s__destination @ (sK1 @ sK7) @ (sK2 @ sK7) @ sK7))) [trivial inequality removal 700]
793. ($true != $true) | ($true = ((accreln1 @ s__Obligation @ X0 @ sK7))) | (((accreln1 @ s__Obligation @ cw @ X0)) != $true) [superposition 77,106]
811. (((accreln1 @ s__Obligation @ cw @ X0)) != $true) | ($true = ((accreln1 @ s__Obligation @ X0 @ sK7))) [trivial inequality removal 793]
818. ($true != $true) | ($true = ((accreln1 @ s__Obligation @ sK7 @ sK7))) [superposition 811,106]
819. ($true != $true) | ($true = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ sK7))) [superposition 811,75]
835. ($true = ((accreln1 @ s__Obligation @ (sK0 @ cw) @ sK7))) [trivial inequality removal 819]
836. ($true = ((accreln1 @ s__Obligation @ sK7 @ sK7))) [trivial inequality removal 818]
845. ($true != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [superposition 98,109]
859. (((s__destination @ X2 @ X3 @ X0)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [trivial inequality removal 845]
872. 5 <=> ($true = ((s__instance @ (sK3 @ sK7) @ s__Giving))) [avatar definition]
873. ($true != ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (~5) [avatar component clause 872]
874. ($true = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (5) [avatar component clause 872]
898. 7 <=> ($false = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) [avatar definition]
900. ($false = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) <- (7) [avatar component clause 898]
910. ($true != $true) | ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [superposition 99,109]
924. (((s__destination @ X2 @ X3 @ X0)) != $true) | ($true = ((s__instance @ (sK4 @ X1) @ s__PayCheck))) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [trivial inequality removal 910]
934. 8 <=> ($true = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) [avatar definition]
935. ($true != ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) <- (~8) [avatar component clause 934]
936. ($true = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) <- (8) [avatar component clause 934]
964. ($true != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [superposition 101,109]
978. (((s__destination @ X2 @ X3 @ X0)) != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [trivial inequality removal 964]
988. 9 <=> ($true = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) [avatar definition]
989. ($true != ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) <- (~9) [avatar component clause 988]
990. ($true = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) <- (9) [avatar component clause 988]
1038. ($true != $true) | ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [superposition 100,109]
1052. (((s__destination @ X2 @ X3 @ X0)) != $true) | ($true = ((s__patient @ (sK3 @ X1) @ (sK4 @ X1) @ X1))) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) | (((accreln1 @ s__Obligation @ X0 @ X1)) = $false) [trivial inequality removal 1038]
6719. ($true != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [superposition 859,717]
6772. ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [trivial inequality removal 6719]
6795. ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [forward subsumption resolution 6772,670]
6803. ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [forward subsumption resolution 6795,623]
6806. ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) | ($true = ((s__instance @ (sK3 @ X1) @ s__Giving))) [forward subsumption resolution 6803,570]
6854. ($true != $true) | ($true = ((s__instance @ (sK4 @ X0) @ s__PayCheck))) | ($true != ((s__agent @ (sK1 @ sK7) @ s__Bill @ sK7))) | ($true != ((s__instance @ (sK2 @ sK7) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ sK7) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [superposition 924,718]
6909. ($true = ((s__instance @ (sK4 @ X0) @ s__PayCheck))) | ($true != ((s__agent @ (sK1 @ sK7) @ s__Bill @ sK7))) | ($true != ((s__instance @ (sK2 @ sK7) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ sK7) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [trivial inequality removal 6854]
6927. 57 <=> ! [X0 : w] : (($true = ((s__instance @ (sK4 @ X0) @ s__PayCheck))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0)))) [avatar definition]
6928. ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) | ($true = ((s__instance @ (sK4 @ X0) @ s__PayCheck))) <- (57) [avatar component clause 6927]
6932. ($true = ((s__instance @ (sK4 @ X0) @ s__PayCheck))) | ($true != ((s__instance @ (sK2 @ sK7) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ sK7) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [forward subsumption resolution 6909,671]
6940. ($true = ((s__instance @ (sK4 @ X0) @ s__PayCheck))) | ($true != ((s__instance @ (sK1 @ sK7) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [forward subsumption resolution 6932,624]
6943. ($true = ((s__instance @ (sK4 @ X0) @ s__PayCheck))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [forward subsumption resolution 6940,571]
6944. 57 [avatar split clause 6943,6927]
7040. ($true != $true) | ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [superposition 978,717]
7093. ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | ($true != ((s__agent @ (sK1 @ (sK0 @ X0)) @ s__Bill @ (sK0 @ X0)))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [trivial inequality removal 7040]
7108. ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | ($true != ((s__instance @ (sK2 @ (sK0 @ X0)) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [forward subsumption resolution 7093,670]
7111. ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | ($true != ((s__instance @ (sK1 @ (sK0 @ X0)) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [forward subsumption resolution 7108,623]
7114. ($true = ((s__destination @ (sK3 @ X1) @ s__Bill @ X1))) | ($false = ((accreln1 @ s__Obligation @ (sK0 @ X0) @ X1))) [forward subsumption resolution 7111,570]
8444. ($true != $true) | ($true = ((s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0))) | ($true != ((s__agent @ (sK1 @ sK7) @ s__Bill @ sK7))) | ($true != ((s__instance @ (sK2 @ sK7) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ sK7) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [superposition 1052,718]
8510. ($true = ((s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0))) | ($true != ((s__agent @ (sK1 @ sK7) @ s__Bill @ sK7))) | ($true != ((s__instance @ (sK2 @ sK7) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ sK7) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [trivial inequality removal 8444]
8520. 62 <=> ! [X0 : w] : (($true = ((s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0)))) [avatar definition]
8521. ($true = ((s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) <- (62) [avatar component clause 8520]
8525. ($true = ((s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0))) | ($true != ((s__instance @ (sK2 @ sK7) @ s__Moon))) | ($true != ((s__instance @ (sK1 @ sK7) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [forward subsumption resolution 8510,671]
8528. ($true = ((s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0))) | ($true != ((s__instance @ (sK1 @ sK7) @ s__Flying))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [forward subsumption resolution 8525,624]
8531. ($true = ((s__patient @ (sK3 @ X0) @ (sK4 @ X0) @ X0))) | ($false = ((accreln1 @ s__Obligation @ sK7 @ X0))) [forward subsumption resolution 8528,571]
8532. 62 [avatar split clause 8531,8520]
13485. ($true != $true) | ($false = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) <- (~9) [superposition 989,109]
13564. ($false = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) <- (~9) [trivial inequality removal 13485]
13570. 7 | 9 [avatar split clause 13564,988,898]
14034. ($true = $false) | ($true = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) <- (57) [superposition 6928,836]
14142. ($true = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) <- (57) [trivial inequality removal 14034]
14173. $false <- (~8, 57) [forward subsumption resolution 14142,935]
14174. 8 | ~57 [avatar contradiction clause 14173]
15831. ($true = $false) | ($false = ((accreln1 @ s__Obligation @ sK7 @ sK7))) | ($false = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) | ($false = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) | ($false = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (62) [superposition 8521,120]
15923. ($false = ((accreln1 @ s__Obligation @ sK7 @ sK7))) | ($false = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) | ($false = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) | ($false = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (62) [trivial inequality removal 15831]
15956. ($true = $false) | ($false = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) | ($false = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) | ($false = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (62) [forward demodulation 15923,836]
15957. ($false = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) | ($false = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) | ($false = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (62) [trivial inequality removal 15956]
15976. ($true = $false) | ($false = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) | ($false = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (9, 62) [forward demodulation 15957,990]
15977. ($false = ((s__instance @ (sK4 @ sK7) @ s__PayCheck))) | ($false = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (9, 62) [trivial inequality removal 15976]
15978. ($true = $false) | ($false = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (8, 9, 62) [forward demodulation 15977,936]
15979. ($false = ((s__instance @ (sK3 @ sK7) @ s__Giving))) <- (8, 9, 62) [trivial inequality removal 15978]
15980. ($true = $false) <- (5, 8, 9, 62) [superposition 15979,874]
16066. $false <- (5, 8, 9, 62) [trivial inequality removal 15980]
16067. ~5 | ~8 | ~9 | ~62 [avatar contradiction clause 16066]
19947. ($true = $false) | ($true = ((s__instance @ (sK3 @ sK7) @ s__Giving))) [superposition 6806,835]
20079. ($true = ((s__instance @ (sK3 @ sK7) @ s__Giving))) [trivial inequality removal 19947]
24557. ($true = $false) | ($true = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) [superposition 7114,835]
24721. ($true = ((s__destination @ (sK3 @ sK7) @ s__Bill @ sK7))) [trivial inequality removal 24557]
24785. ($true = $false) <- (7) [forward demodulation 24721,900]
24786. $false <- (7) [trivial inequality removal 24785]
24787. ~7 [avatar contradiction clause 24786]
24931. $false <- (~5) [forward subsumption resolution 20079,873]
24932. 5 [avatar contradiction clause 24931]
s160. 57 [sat_conversion 6944]
s168. 62 [sat_conversion 8532]
s329. 7 | 9 [sat_conversion 13570]
s345. ~57 | 8 [sat_conversion 14174]
s401. ~62 | ~9 | ~8 | ~5 [sat_conversion 16067]
s748. ~7 [sat_conversion 24787]
s784. 5 [sat_conversion 24932]
s813. ~62 | ~9 | ~8 [rat s401,s784]
s824. 9 [rat s329,s748]
s838. ~8 [rat s813,s824,s168]
s839. ~57 [rat s345,s838]
s840. # [rat s160,s839]
24937. $false [avatar sat refutation s840]
% SZS output end Proof for 1h_Obligation_UO1_univConj
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 3.328 s
% Peak memory usage: 21 MB
% Instructions burned: 23881 (million)
% ------------------------------
% ------------------------------
