% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for 1h_Obligation_UO1_univConj
% SZS output start Proof for 1h_Obligation_UO1_univConj
18. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [input(axiom)]
19. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [input(axiom)]
20. ! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [input(conjecture)]
21. ~! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [negated conjecture 20]
26. ! [X0 : w,X1 : w] : ! [X2,X3] : (((s__destination @ X2 @ X3 @ X0) & (s__agent @ X2 @ s__Bill @ X0) & (s__instance @ X3 @ s__Moon) & (s__instance @ X2 @ s__Flying)) => ((accreln1 @ s__Obligation @ X0 @ X1) => ? [X4,X5] : ((s__destination @ X4 @ s__Bill @ X1) & (s__patient @ X4 @ X5 @ X1) & (s__instance @ X5 @ s__PayCheck) & (s__instance @ X4 @ s__Giving)))) [rectify 18]
27. ! [X0 : w,X1 : w] : ! [X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [fool elimination 26]
28. ? [X0,X1] : ((s__destination @ X0 @ X1 @ cw) & (s__agent @ X0 @ s__Bill @ cw) & (s__instance @ X1 @ s__Moon) & (s__instance @ X0 @ s__Flying)) [rectify 19]
29. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) [fool elimination 28]
30. ~! [X0 : w] : ((accreln1 @ s__Obligation @ cw @ X0) => ? [X1,X2] : ((s__destination @ X1 @ s__Bill @ X0) & (s__patient @ X1 @ X2 @ X0) & (s__instance @ X2 @ s__PayCheck) & (s__instance @ X1 @ s__Giving))) [rectify 21]
31. ~! [X0 : w] : ((((accreln1 @ s__Obligation @ cw @ X0)) = $true) => ? [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) = $true) & (((s__patient @ X1 @ X2 @ X0)) = $true) & (((s__instance @ X2 @ s__PayCheck)) = $true) & (((s__instance @ X1 @ s__Giving)) = $true))) [fool elimination 30]
47. ! [X0 : w,X1 : w,X2,X3] : (((((s__destination @ X2 @ X3 @ X0)) = $true) & (((s__agent @ X2 @ s__Bill @ X0)) = $true) & (((s__instance @ X3 @ s__Moon)) = $true) & (((s__instance @ X2 @ s__Flying)) = $true)) => ((((accreln1 @ s__Obligation @ X0 @ X1)) = $true) => ? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)))) [flattening 27]
49. ! [X0 : w,X1 : w,X2,X3] : ((? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true)) | ((((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true))) [ennf transformation 47]
50. ! [X0 : w,X1 : w,X2,X3] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [flattening 49]
51. ? [X0 : w] : (! [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) [ennf transformation 31]
54. ! [X1 : w] : (? [X4,X5] : ((((s__destination @ X4 @ s__Bill @ X1)) = $true) & (((s__patient @ X4 @ X5 @ X1)) = $true) & (((s__instance @ X5 @ s__PayCheck)) = $true) & (((s__instance @ X4 @ s__Giving)) = $true)) => (($true = ((s__destination @ (sK2 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK2 @ X1) @ (sK3 @ X1) @ X1))) & ($true = ((s__instance @ (sK3 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK2 @ X1) @ s__Giving))))) [skolem symbol introduction]
55. ! [X0 : w,X1 : w,X2,X3] : ((($true = ((s__destination @ (sK2 @ X1) @ s__Bill @ X1))) & ($true = ((s__patient @ (sK2 @ X1) @ (sK3 @ X1) @ X1))) & ($true = ((s__instance @ (sK3 @ X1) @ s__PayCheck))) & ($true = ((s__instance @ (sK2 @ X1) @ s__Giving)))) | (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true)) [skolemisation 50,54]
56. ? [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) = $true) & (((s__agent @ X0 @ s__Bill @ cw)) = $true) & (((s__instance @ X1 @ s__Moon)) = $true) & (((s__instance @ X0 @ s__Flying)) = $true)) => (($true = ((s__destination @ sK4 @ sK5 @ cw))) & ($true = ((s__agent @ sK4 @ s__Bill @ cw))) & ($true = ((s__instance @ sK5 @ s__Moon))) & ($true = ((s__instance @ sK4 @ s__Flying)))) [skolem symbol introduction]
57. ($true = ((s__destination @ sK4 @ sK5 @ cw))) & ($true = ((s__agent @ sK4 @ s__Bill @ cw))) & ($true = ((s__instance @ sK5 @ s__Moon))) & ($true = ((s__instance @ sK4 @ s__Flying))) [skolemisation 29,56]
58. ? [X0 : w] : (! [X1,X2] : ((((s__destination @ X1 @ s__Bill @ X0)) != $true) | (((s__patient @ X1 @ X2 @ X0)) != $true) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & (((accreln1 @ s__Obligation @ cw @ X0)) = $true)) => (! [X2,X1] : (($true != ((s__destination @ X1 @ s__Bill @ sK6))) | ($true != ((s__patient @ X1 @ X2 @ sK6))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & ($true = ((accreln1 @ s__Obligation @ cw @ sK6)))) [skolem symbol introduction]
59. ! [X1,X2] : (($true != ((s__destination @ X1 @ s__Bill @ sK6))) | ($true != ((s__patient @ X1 @ X2 @ sK6))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true)) & ($true = ((accreln1 @ s__Obligation @ cw @ sK6))) [skolemisation 51,58]
80. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK2 @ X1) @ s__Giving))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 55]
81. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__instance @ (sK3 @ X1) @ s__PayCheck))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 55]
82. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__patient @ (sK2 @ X1) @ (sK3 @ X1) @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 55]
83. (((accreln1 @ s__Obligation @ X0 @ X1)) != $true) | ($true = ((s__destination @ (sK2 @ X1) @ s__Bill @ X1))) | (((s__destination @ X2 @ X3 @ X0)) != $true) | (((s__agent @ X2 @ s__Bill @ X0)) != $true) | (((s__instance @ X3 @ s__Moon)) != $true) | (((s__instance @ X2 @ s__Flying)) != $true) [cnf transformation 55]
84. ($true = ((s__instance @ sK4 @ s__Flying))) [cnf transformation 57]
85. ($true = ((s__instance @ sK5 @ s__Moon))) [cnf transformation 57]
86. ($true = ((s__agent @ sK4 @ s__Bill @ cw))) [cnf transformation 57]
87. ($true = ((s__destination @ sK4 @ sK5 @ cw))) [cnf transformation 57]
88. ($true = ((accreln1 @ s__Obligation @ cw @ sK6))) [cnf transformation 59]
89. ($true != ((s__destination @ X1 @ s__Bill @ sK6))) | ($true != ((s__patient @ X1 @ X2 @ sK6))) | (((s__instance @ X2 @ s__PayCheck)) != $true) | (((s__instance @ X1 @ s__Giving)) != $true) [cnf transformation 59]
91. ($true = X0) | ($false = X0) [fool exhaustiveness axiom]
95. ($true != $true) | ($true != ((s__patient @ X0 @ X1 @ sK6))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK6))) [superposition 89,91]
96. ($true != ((s__patient @ X0 @ X1 @ sK6))) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK6))) [trivial inequality removal 95]
97. ($true != $true) | ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK6))) | ($false = ((s__patient @ X0 @ X1 @ sK6))) [superposition 96,91]
98. ($true != ((s__instance @ X1 @ s__PayCheck))) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK6))) | ($false = ((s__patient @ X0 @ X1 @ sK6))) [trivial inequality removal 97]
269. ($true != $true) | ($true = ((s__instance @ (sK2 @ sK6) @ s__Giving))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 80,88]
276. ($true = ((s__instance @ (sK2 @ sK6) @ s__Giving))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 269]
278. 1 <=> ! [X0,X1] : ((((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true)) [avatar definition]
279. (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) <- (1) [avatar component clause 278]
281. 2 <=> ($true = ((s__instance @ (sK2 @ sK6) @ s__Giving))) [avatar definition]
283. ($true = ((s__instance @ (sK2 @ sK6) @ s__Giving))) <- (2) [avatar component clause 281]
284. 1 | 2 [avatar split clause 276,281,278]
297. 4 <=> ($false = ((s__destination @ (sK2 @ sK6) @ s__Bill @ sK6))) [avatar definition]
299. ($false = ((s__destination @ (sK2 @ sK6) @ s__Bill @ sK6))) <- (4) [avatar component clause 297]
303. ($true != $true) | ($true = ((s__instance @ (sK3 @ sK6) @ s__PayCheck))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 81,88]
310. ($true = ((s__instance @ (sK3 @ sK6) @ s__PayCheck))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 303]
312. 5 <=> ($true = ((s__instance @ (sK3 @ sK6) @ s__PayCheck))) [avatar definition]
314. ($true = ((s__instance @ (sK3 @ sK6) @ s__PayCheck))) <- (5) [avatar component clause 312]
315. 1 | 5 [avatar split clause 310,312,278]
320. ($true != $true) | ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK6))) | ($false = ((s__patient @ X0 @ (sK3 @ sK6) @ sK6))) <- (5) [superposition 98,314]
321. ($true != ((s__instance @ X0 @ s__Giving))) | ($false = ((s__destination @ X0 @ s__Bill @ sK6))) | ($false = ((s__patient @ X0 @ (sK3 @ sK6) @ sK6))) <- (5) [trivial inequality removal 320]
326. ($true != $true) | ($true = ((s__destination @ (sK2 @ sK6) @ s__Bill @ sK6))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 83,88]
333. ($true = ((s__destination @ (sK2 @ sK6) @ s__Bill @ sK6))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 326]
335. 6 <=> ($true = ((s__destination @ (sK2 @ sK6) @ s__Bill @ sK6))) [avatar definition]
337. ($true = ((s__destination @ (sK2 @ sK6) @ s__Bill @ sK6))) <- (6) [avatar component clause 335]
338. 1 | 6 [avatar split clause 333,335,278]
346. ($true != $true) | ($true = ((s__patient @ (sK2 @ sK6) @ (sK3 @ sK6) @ sK6))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [superposition 82,88]
353. ($true = ((s__patient @ (sK2 @ sK6) @ (sK3 @ sK6) @ sK6))) | (((s__destination @ X0 @ X1 @ cw)) != $true) | (((s__agent @ X0 @ s__Bill @ cw)) != $true) | (((s__instance @ X1 @ s__Moon)) != $true) | (((s__instance @ X0 @ s__Flying)) != $true) [trivial inequality removal 346]
354. ($true != $true) | ($false = ((s__destination @ (sK2 @ sK6) @ s__Bill @ sK6))) | ($false = ((s__patient @ (sK2 @ sK6) @ (sK3 @ sK6) @ sK6))) <- (2, 5) [superposition 321,283]
361. ($false = ((s__destination @ (sK2 @ sK6) @ s__Bill @ sK6))) | ($false = ((s__patient @ (sK2 @ sK6) @ (sK3 @ sK6) @ sK6))) <- (2, 5) [trivial inequality removal 354]
363. 7 <=> ($false = ((s__patient @ (sK2 @ sK6) @ (sK3 @ sK6) @ sK6))) [avatar definition]
365. ($false = ((s__patient @ (sK2 @ sK6) @ (sK3 @ sK6) @ sK6))) <- (7) [avatar component clause 363]
366. 7 | 4 | ~2 | ~5 [avatar split clause 361,312,281,297,363]
381. ($true != $true) | ($true != ((s__instance @ sK4 @ s__Flying))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) <- (1) [superposition 279,87]
388. ($true != ((s__instance @ sK4 @ s__Flying))) | ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) <- (1) [trivial inequality removal 381]
389. ($true != ((s__instance @ sK5 @ s__Moon))) | ($true != ((s__agent @ sK4 @ s__Bill @ cw))) <- (1) [forward subsumption resolution 388,84]
390. ($true != ((s__agent @ sK4 @ s__Bill @ cw))) <- (1) [forward subsumption resolution 389,85]
391. $false <- (1) [forward subsumption resolution 390,86]
392. ~1 [avatar contradiction clause 391]
393. ($true = $false) <- (4, 6) [forward demodulation 337,299]
394. $false <- (4, 6) [trivial inequality removal 393]
395. ~4 | ~6 [avatar contradiction clause 394]
397. 8 <=> ($true = ((s__patient @ (sK2 @ sK6) @ (sK3 @ sK6) @ sK6))) [avatar definition]
399. ($true = ((s__patient @ (sK2 @ sK6) @ (sK3 @ sK6) @ sK6))) <- (8) [avatar component clause 397]
400. 1 | 8 [avatar split clause 353,397,278]
401. ($true = $false) <- (7, 8) [forward demodulation 399,365]
402. $false <- (7, 8) [trivial inequality removal 401]
403. ~7 | ~8 [avatar contradiction clause 402]
s1. 1 | 2 [sat_conversion 284]
s3. 1 | 5 [sat_conversion 315]
s4. 1 | 6 [sat_conversion 338]
s5. ~5 | ~2 | 4 | 7 [sat_conversion 366]
s6. ~1 [sat_conversion 392]
s7. ~6 | ~4 [sat_conversion 395]
s8. 1 | 8 [sat_conversion 400]
s9. ~8 | ~7 [sat_conversion 403]
s10. 8 [rat s8,s6]
s11. ~7 [rat s9,s10]
s12. ~5 | ~2 | 4 [rat s5,s11]
s13. 6 [rat s4,s6]
s14. ~4 [rat s7,s13]
s15. 5 [rat s3,s6]
s16. ~2 [rat s12,s14,s15]
s17. # [rat s1,s16,s6]
404. $false [avatar sat refutation s17]
% SZS output end Proof for 1h_Obligation_UO1_univConj
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit ec9595e79 on 2026-01-27 11:46:19 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 z3-4.8.4-9178-g3c47fd96c
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.151 s
% Peak memory usage: 12 MB
% Instructions burned: 31 (million)
% ------------------------------
% ------------------------------
