 % Checking upload ... % Checker ran ... % No errors ...

% START OF SYSTEM OUTPUT
Running higher-order theorem proving
Running: /home/tptp/Systems/Vampire---5.0.1/vampire-ho --input_syntax tptp --proof tptp --output_axiom_names on --print_def_sorts on --mode portfolio --schedule snake_tptp_hol --cores 7 -m 16384 -t 60 /tmp/P_eXfui2V5/SOT_1oz6Aa
% (1248141)dis+1010_1:1_au=on:cbe=off:chr=on:fsr=off:hfsq=on:nm=64:sos=theory:sp=weighted_frequency:i=27:si=on:rtra=on_0 on SOT_1oz6Aa for (1ds/27Mi)
% (1248143)lrs+1002_1:128_aac=none:au=on:cnfonf=lazy_not_gen_be_off:sos=all:i=2:si=on:rtra=on_0 on SOT_1oz6Aa for (1ds/2Mi)
% (1248145)lrs+1004_1:128_cond=on:e2e=on:sp=weighted_frequency:i=18:si=on:rtra=on_0 on SOT_1oz6Aa for (1ds/18Mi)
% (1248140)lrs+10_1:1_c=on:cnfonf=conj_eager:fd=off:fe=off:kws=frequency:spb=intro:i=4:si=on:rtra=on_0 on SOT_1oz6Aa for (1ds/4Mi)
% (1248139)lrs+1002_1:8_bd=off:fd=off:hud=10:tnu=1:i=183:si=on:rtra=on_0 on SOT_1oz6Aa for (1ds/183Mi)
% (1248144)lrs+1002_1:1_au=on:bd=off:e2e=on:sd=2:sos=on:ss=axioms:i=275:si=on:rtra=on_0 on SOT_1oz6Aa for (1ds/275Mi)
% (1248142)lrs+10_1:1_au=on:inj=on:i=2:si=on:rtra=on_0 on SOT_1oz6Aa for (1ds/2Mi)
% (1248143)Refutation not found, incomplete strategy
% (1248143)------------------------------
% (1248143)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (1248143)Termination reason: Refutation not found, incomplete strategy


% (1248143)Memory used [KB]: 5628
% (1248143)Time elapsed: 0.005 s
% (1248143)------------------------------
% (1248143)------------------------------
% (1248146)lrs+10_1:1_bet=on:cnfonf=off:fd=off:hud=5:inj=on:i=3:si=on:rtra=on_0 on SOT_1oz6Aa for (1ds/3Mi)
% (1248141)First to succeed.
% (1248141)Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_1oz6Aa
% SZS output start Proof for SOT_1oz6Aa
thf(type_def_5, type, w: $tType).
thf(type_def_6, type, m: $tType).
thf(func_def_0, type, cw: w).
thf(func_def_1, type, s__Obligation: m).
thf(func_def_2, type, s__Permission: m).
thf(func_def_3, type, s__Prohibition: m).
thf(func_def_4, type, s__confersObligation: m).
thf(func_def_5, type, accreln1: m > w > w > $o).
thf(func_def_6, type, accreln2: m > $i > w > w > $o).
thf(func_def_7, type, accreln3: m > $i > $i > w > w > $o).
thf(func_def_9, type, s__instance: $i > $i > $o).
thf(func_def_10, type, s__agent: $i > $i > w > $o).
thf(func_def_11, type, s__patient: $i > $i > w > $o).
thf(func_def_12, type, s__origin: $i > $i > w > $o).
thf(func_def_13, type, s__destination: $i > $i > w > $o).
thf(func_def_24, type, vEPSILON: !>[X0: $tType]:((X0 > $o) > X0)).
thf(func_def_29, type, sK2: w > $i).
thf(func_def_30, type, sK3: w > $i).
thf(func_def_31, type, sK4: $i > $i > w > w).
thf(func_def_32, type, sK5: w > $i).
thf(f235,plain,(
  $false),
  inference(avatar_sat_refutation,[status(thm)],[f184,f198,f231,f234])).
thf(f234,plain,(
  ~spl6_6 | ~spl6_9),
  inference(avatar_split_clause,[status(thm)],[f233,f195,f178])).
thf(f178,definition,(
  spl6_6 <=> ($true = (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Tax))),
  introduced(definition,[new_symbols(naming,[spl6_6])],[avatar_definition])).
thf(f195,definition,(
  spl6_9 <=> ($true = (s__origin @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))))),
  introduced(definition,[new_symbols(naming,[spl6_9])],[avatar_definition])).
thf(f233,plain,(
  ($true != (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Tax)) | ~spl6_9),
  inference(trivial_inequality_removal,[status(thm)],[f232])).
thf(f232,plain,(
  ($true != (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Tax)) | ($true != $true) | ~spl6_9),
  inference(superposition,[status(thm)],[f226,f197])).
thf(f197,plain,(
  ($true = (s__origin @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))) | ~spl6_9),
  inference(avatar_component_clause,[status(thm)],[f195])).
thf(f226,plain,(
  ( ! [X0 : $i] : (($true != (s__origin @ X0 @ s__Bill @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))) | ($true != (s__instance @ X0 @ s__Tax))) )),
  inference(trivial_inequality_removal,[status(thm)],[f216])).
thf(f216,plain,(
  ( ! [X0 : $i] : (($true != $true) | ($true != (s__origin @ X0 @ s__Bill @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))) | ($true != (s__instance @ X0 @ s__Tax))) )),
  inference(superposition,[status(thm)],[f96,f134])).
thf(f134,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (((accreln3 @ s__confersObligation @ X0 @ X1 @ X2 @ (sK4 @ X0 @ X1 @ (sK4 @ X0 @ X1 @ X2))) = $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f133])).
thf(f133,plain,(
  ( ! [X2 : w,X0 : $i,X1 : $i] : (((accreln3 @ s__confersObligation @ X0 @ X1 @ X2 @ (sK4 @ X0 @ X1 @ (sK4 @ X0 @ X1 @ X2))) = $true) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f120,f92])).
thf(f92,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ (sK4 @ X2 @ X1 @ X0)) = $true)) )),
  inference(cnf_transformation,[status(thm)],[f63])).
thf(f63,plain,(
  ! [X0 : w,X1 : $i,X2 : $i] : ((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ (sK4 @ X2 @ X1 @ X0)) = $true)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f47,f62])).
thf(f62,plain,(
  ! [X0 : w,X1 : $i,X2 : $i] : (? [X3 : w] : ($true = (accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ X3)) => ((accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ (sK4 @ X2 @ X1 @ X0)) = $true))),
  introduced(definition,[],[choice_axiom])).
thf(f47,plain,(
  ! [X0 : w,X1 : $i,X2 : $i] : ? [X3 : w] : ($true = (accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ X3))),
  inference(fool_elimination,[status(thm)],[f46])).
thf(f46,plain,(
  ! [X0 : w,X1 : $i,X2 : $i] : ? [X3 : w] : (accreln3 @ s__confersObligation @ X2 @ X1 @ X0 @ X3)),
  inference(rectify,[status(thm)],[f1])).
thf(f1,axiom,(
  ! [X0 : w,X2 : $i,X1 : $i] : ? [X3 : w] : (accreln3 @ s__confersObligation @ X1 @ X2 @ X0 @ X3)),
  file('/tmp/P_eXfui2V5/SOT_1oz6Aa',serial_obligation)).
thf(f120,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (($true != (accreln3 @ s__confersObligation @ X0 @ X1 @ (sK4 @ X0 @ X1 @ X2) @ X3)) | ($true = (accreln3 @ s__confersObligation @ X0 @ X1 @ X2 @ X3))) )),
  inference(trivial_inequality_removal,[status(thm)],[f119])).
thf(f119,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (($true != $true) | ($true = (accreln3 @ s__confersObligation @ X0 @ X1 @ X2 @ X3)) | ($true != (accreln3 @ s__confersObligation @ X0 @ X1 @ (sK4 @ X0 @ X1 @ X2) @ X3))) )),
  inference(superposition,[status(thm)],[f89,f92])).
thf(f89,plain,(
  ( ! [X2 : w,X3 : $i,X0 : $i,X1 : w,X4 : w] : (($true != (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X1)) | ($true = (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X4)) | ($true != (accreln3 @ s__confersObligation @ X0 @ X3 @ X1 @ X4))) )),
  inference(cnf_transformation,[status(thm)],[f55])).
thf(f55,plain,(
  ! [X0 : $i,X1 : w,X2 : w,X3 : $i,X4 : w] : (($true = (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X4)) | ($true != (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X1)) | ($true != (accreln3 @ s__confersObligation @ X0 @ X3 @ X1 @ X4)))),
  inference(flattening,[status(thm)],[f54])).
thf(f54,plain,(
  ! [X1 : w,X0 : $i,X3 : $i,X4 : w,X2 : w] : (($true = (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X4)) | (($true != (accreln3 @ s__confersObligation @ X0 @ X3 @ X1 @ X4)) | ($true != (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X1))))),
  inference(ennf_transformation,[status(thm)],[f39])).
thf(f39,plain,(
  ! [X1 : w,X0 : $i,X3 : $i,X4 : w,X2 : w] : ((($true = (accreln3 @ s__confersObligation @ X0 @ X3 @ X1 @ X4)) & ($true = (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X1))) => ($true = (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X4)))),
  inference(fool_elimination,[status(thm)],[f38])).
thf(f38,plain,(
  ! [X0 : $i,X1 : w,X2 : w,X3 : $i,X4 : w] : (((accreln3 @ s__confersObligation @ X0 @ X3 @ X1 @ X4) & (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X1)) => (accreln3 @ s__confersObligation @ X0 @ X3 @ X2 @ X4))),
  inference(rectify,[status(thm)],[f2])).
thf(f2,axiom,(
  ! [X1 : $i,X3 : w,X0 : w,X2 : $i,X4 : w] : (((accreln3 @ s__confersObligation @ X1 @ X2 @ X3 @ X4) & (accreln3 @ s__confersObligation @ X1 @ X2 @ X0 @ X3)) => (accreln3 @ s__confersObligation @ X1 @ X2 @ X0 @ X4))),
  file('/tmp/P_eXfui2V5/SOT_1oz6Aa',transitive_obligation)).
thf(f96,plain,(
  ( ! [X0 : w,X1 : $i] : (((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0) != $true) | ($true != (s__origin @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X1 @ s__Tax))) )),
  inference(cnf_transformation,[status(thm)],[f51])).
thf(f51,plain,(
  ! [X0 : w] : (! [X1 : $i] : (($true != (s__origin @ X1 @ s__Bill @ X0)) | ($true != (s__instance @ X1 @ s__Tax))) | ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0) != $true))),
  inference(ennf_transformation,[status(thm)],[f48])).
thf(f48,plain,(
  ~? [X0 : w] : (? [X1 : $i] : (($true = (s__origin @ X1 @ s__Bill @ X0)) & ($true = (s__instance @ X1 @ s__Tax))) & ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0) = $true))),
  inference(rectify,[status(thm)],[f33])).
thf(f33,plain,(
  ~? [X0 : w] : (((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0) = $true) & ? [X2 : $i] : (($true = (s__instance @ X2 @ s__Tax)) & ($true = (s__origin @ X2 @ s__Bill @ X0))))),
  inference(fool_elimination,[status(thm)],[f32])).
thf(f32,plain,(
  ~? [X0 : w] : ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X0) & ? [X2 : $i] : ((s__instance @ X2 @ s__Tax) & (s__origin @ X2 @ s__Bill @ X0)))),
  inference(rectify,[status(thm)],[f30])).
thf(f30,negated_conjecture,(
  ~? [X12 : w,X13 : $i] : ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X12) & ? [X7 : $i] : ((s__instance @ X7 @ s__Tax) & (s__origin @ X7 @ s__Bill @ X12)))),
  inference(negated_conjecture,[status(cth)],[f29])).
thf(f29,conjecture,(
  ? [X12 : w,X13 : $i] : ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ cw @ X12) & ? [X7 : $i] : ((s__instance @ X7 @ s__Tax) & (s__origin @ X7 @ s__Bill @ X12)))),
  file('/tmp/P_eXfui2V5/SOT_1oz6Aa',conjecture_he_get_taxed)).
thf(f231,plain,(
  ~spl6_7),
  inference(avatar_contradiction_clause,[status(thm)],[f230])).
thf(f230,plain,(
  $false | ~spl6_7),
  inference(trivial_inequality_removal,[status(thm)],[f229])).
thf(f229,plain,(
  ($true != $true) | ~spl6_7),
  inference(superposition,[status(thm)],[f228,f128])).
thf(f128,plain,(
  ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck))),
  inference(trivial_inequality_removal,[status(thm)],[f127])).
thf(f127,plain,(
  ($true != $true) | ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck))),
  inference(forward_demodulation,[status(thm)],[f126,f78])).
thf(f78,plain,(
  ($true = (s__instance @ sK0 @ s__Flying))),
  inference(cnf_transformation,[status(thm)],[f58])).
thf(f58,plain,(
  ($true = (s__destination @ sK0 @ sK1 @ cw)) & ($true = (s__instance @ sK1 @ s__Moon)) & ($true = (s__agent @ sK0 @ s__Bill @ cw)) & ($true = (s__instance @ sK0 @ s__Flying))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f56,f57])).
thf(f57,plain,(
  ? [X0 : $i,X1 : $i] : (($true = (s__destination @ X0 @ X1 @ cw)) & ($true = (s__instance @ X1 @ s__Moon)) & ((s__agent @ X0 @ s__Bill @ cw) = $true) & ($true = (s__instance @ X0 @ s__Flying))) => (($true = (s__destination @ sK0 @ sK1 @ cw)) & ($true = (s__instance @ sK1 @ s__Moon)) & ($true = (s__agent @ sK0 @ s__Bill @ cw)) & ($true = (s__instance @ sK0 @ s__Flying)))),
  introduced(definition,[],[choice_axiom])).
thf(f56,plain,(
  ? [X0 : $i,X1 : $i] : (($true = (s__destination @ X0 @ X1 @ cw)) & ($true = (s__instance @ X1 @ s__Moon)) & ((s__agent @ X0 @ s__Bill @ cw) = $true) & ($true = (s__instance @ X0 @ s__Flying)))),
  inference(rectify,[status(thm)],[f45])).
thf(f45,plain,(
  ? [X1 : $i,X0 : $i] : (($true = (s__destination @ X1 @ X0 @ cw)) & ((s__instance @ X0 @ s__Moon) = $true) & ($true = (s__agent @ X1 @ s__Bill @ cw)) & ($true = (s__instance @ X1 @ s__Flying)))),
  inference(fool_elimination,[status(thm)],[f44])).
thf(f44,plain,(
  ? [X0 : $i,X1 : $i] : ((s__instance @ X0 @ s__Moon) & (s__instance @ X1 @ s__Flying) & (s__agent @ X1 @ s__Bill @ cw) & (s__destination @ X1 @ X0 @ cw))),
  inference(rectify,[status(thm)],[f28])).
thf(f28,axiom,(
  ? [X9 : $i,X8 : $i] : ((s__instance @ X9 @ s__Moon) & (s__instance @ X8 @ s__Flying) & (s__agent @ X8 @ s__Bill @ cw) & (s__destination @ X8 @ X9 @ cw))),
  file('/tmp/P_eXfui2V5/SOT_1oz6Aa',fact_flies_to_moon)).
thf(f126,plain,(
  ($true != (s__instance @ sK0 @ s__Flying)) | ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck))),
  inference(trivial_inequality_removal,[status(thm)],[f125])).
thf(f125,plain,(
  ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck)) | ($true != (s__instance @ sK0 @ s__Flying)) | ($true != $true)),
  inference(forward_demodulation,[status(thm)],[f124,f79])).
thf(f79,plain,(
  ($true = (s__agent @ sK0 @ s__Bill @ cw))),
  inference(cnf_transformation,[status(thm)],[f58])).
thf(f124,plain,(
  ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck)) | ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true != (s__instance @ sK0 @ s__Flying))),
  inference(trivial_inequality_removal,[status(thm)],[f123])).
thf(f123,plain,(
  ($true != (s__instance @ sK0 @ s__Flying)) | ($true != $true) | ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck))),
  inference(forward_demodulation,[status(thm)],[f122,f80])).
thf(f80,plain,(
  ($true = (s__instance @ sK1 @ s__Moon))),
  inference(cnf_transformation,[status(thm)],[f58])).
thf(f122,plain,(
  ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck)) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true != (s__instance @ sK0 @ s__Flying))),
  inference(trivial_inequality_removal,[status(thm)],[f121])).
thf(f121,plain,(
  ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck)) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true != (s__instance @ sK0 @ s__Flying)) | ($true != $true) | ($true != (s__agent @ sK0 @ s__Bill @ cw))),
  inference(superposition,[status(thm)],[f108,f81])).
thf(f81,plain,(
  ($true = (s__destination @ sK0 @ sK1 @ cw))),
  inference(cnf_transformation,[status(thm)],[f58])).
thf(f108,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X2 @ X1 @ X0)) | ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__PayCheck)) | ($true != (s__instance @ X2 @ s__Flying)) | ((s__agent @ X2 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X1 @ s__Moon))) )),
  inference(trivial_inequality_removal,[status(thm)],[f107])).
thf(f107,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__agent @ X2 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X1 @ s__Moon)) | ($true = (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__PayCheck)) | ($true != (s__destination @ X2 @ X1 @ X0)) | ($true != (s__instance @ X2 @ s__Flying)) | ($true != $true)) )),
  inference(superposition,[status(thm)],[f85,f92])).
thf(f85,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3)) | ((s__instance @ X0 @ s__Moon) != $true) | ((s__instance @ (sK2 @ X3) @ s__PayCheck) = $true) | ((s__destination @ X1 @ X0 @ X2) != $true) | ((s__agent @ X1 @ s__Bill @ X2) != $true) | ($true != (s__instance @ X1 @ s__Flying))) )),
  inference(cnf_transformation,[status(thm)],[f61])).
thf(f61,plain,(
  ! [X0 : $i,X1 : $i,X2 : w,X3 : w] : (((s__destination @ X1 @ X0 @ X2) != $true) | (($true = (s__patient @ (sK3 @ X3) @ (sK2 @ X3) @ X3)) & ($true = (s__destination @ (sK3 @ X3) @ s__Bill @ X3)) & ((s__instance @ (sK2 @ X3) @ s__PayCheck) = $true) & ($true = (s__instance @ (sK3 @ X3) @ s__Giving))) | ((s__agent @ X1 @ s__Bill @ X2) != $true) | ((s__instance @ X0 @ s__Moon) != $true) | ($true != (s__instance @ X1 @ s__Flying)) | ($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f59,f60])).
thf(f60,plain,(
  ! [X3 : w] : (? [X4 : $i,X5 : $i] : (($true = (s__patient @ X5 @ X4 @ X3)) & ($true = (s__destination @ X5 @ s__Bill @ X3)) & ($true = (s__instance @ X4 @ s__PayCheck)) & ((s__instance @ X5 @ s__Giving) = $true)) => (($true = (s__patient @ (sK3 @ X3) @ (sK2 @ X3) @ X3)) & ($true = (s__destination @ (sK3 @ X3) @ s__Bill @ X3)) & ((s__instance @ (sK2 @ X3) @ s__PayCheck) = $true) & ($true = (s__instance @ (sK3 @ X3) @ s__Giving))))),
  introduced(definition,[],[choice_axiom])).
thf(f59,plain,(
  ! [X0 : $i,X1 : $i,X2 : w,X3 : w] : (((s__destination @ X1 @ X0 @ X2) != $true) | ? [X4 : $i,X5 : $i] : (($true = (s__patient @ X5 @ X4 @ X3)) & ($true = (s__destination @ X5 @ s__Bill @ X3)) & ($true = (s__instance @ X4 @ s__PayCheck)) & ((s__instance @ X5 @ s__Giving) = $true)) | ((s__agent @ X1 @ s__Bill @ X2) != $true) | ((s__instance @ X0 @ s__Moon) != $true) | ($true != (s__instance @ X1 @ s__Flying)) | ($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3)))),
  inference(rectify,[status(thm)],[f53])).
thf(f53,plain,(
  ! [X2 : $i,X0 : $i,X1 : w,X3 : w] : (((s__destination @ X0 @ X2 @ X1) != $true) | ? [X4 : $i,X5 : $i] : (($true = (s__patient @ X5 @ X4 @ X3)) & ($true = (s__destination @ X5 @ s__Bill @ X3)) & ($true = (s__instance @ X4 @ s__PayCheck)) & ((s__instance @ X5 @ s__Giving) = $true)) | ((s__agent @ X0 @ s__Bill @ X1) != $true) | ($true != (s__instance @ X2 @ s__Moon)) | ($true != (s__instance @ X0 @ s__Flying)) | ($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X1 @ X3)))),
  inference(flattening,[status(thm)],[f52])).
thf(f52,plain,(
  ! [X0 : $i,X3 : w,X1 : w,X2 : $i] : ((? [X4 : $i,X5 : $i] : (($true = (s__patient @ X5 @ X4 @ X3)) & ($true = (s__destination @ X5 @ s__Bill @ X3)) & ($true = (s__instance @ X4 @ s__PayCheck)) & ((s__instance @ X5 @ s__Giving) = $true)) | ($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X1 @ X3))) | (($true != (s__instance @ X2 @ s__Moon)) | ((s__destination @ X0 @ X2 @ X1) != $true) | ((s__agent @ X0 @ s__Bill @ X1) != $true) | ($true != (s__instance @ X0 @ s__Flying))))),
  inference(ennf_transformation,[status(thm)],[f35])).
thf(f35,plain,(
  ! [X0 : $i,X3 : w,X1 : w,X2 : $i] : ((($true = (s__instance @ X2 @ s__Moon)) & ((s__destination @ X0 @ X2 @ X1) = $true) & ((s__agent @ X0 @ s__Bill @ X1) = $true) & ($true = (s__instance @ X0 @ s__Flying))) => (($true = (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X1 @ X3)) => ? [X4 : $i,X5 : $i] : (($true = (s__patient @ X5 @ X4 @ X3)) & ($true = (s__destination @ X5 @ s__Bill @ X3)) & ($true = (s__instance @ X4 @ s__PayCheck)) & ((s__instance @ X5 @ s__Giving) = $true))))),
  inference(fool_elimination,[status(thm)],[f34])).
thf(f34,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__agent @ X0 @ s__Bill @ X1) & (s__instance @ X0 @ s__Flying) & (s__instance @ X2 @ s__Moon) & (s__destination @ X0 @ X2 @ X1)) => ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X1 @ X3) => ? [X4 : $i,X5 : $i] : ((s__patient @ X5 @ X4 @ X3) & (s__instance @ X4 @ s__PayCheck) & (s__instance @ X5 @ s__Giving) & (s__destination @ X5 @ s__Bill @ X3))))),
  inference(rectify,[status(thm)],[f27])).
thf(f27,axiom,(
  ! [X8 : $i,X0 : w,X9 : $i,X3 : w] : (((s__agent @ X8 @ s__Bill @ X0) & (s__instance @ X8 @ s__Flying) & (s__instance @ X9 @ s__Moon) & (s__destination @ X8 @ X9 @ X0)) => ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X0 @ X3) => ? [X11 : $i,X10 : $i] : ((s__patient @ X10 @ X11 @ X3) & (s__instance @ X11 @ s__PayCheck) & (s__instance @ X10 @ s__Giving) & (s__destination @ X10 @ s__Bill @ X3))))),
  file('/tmp/P_eXfui2V5/SOT_1oz6Aa',obligation2_must_be_paid)).
thf(f228,plain,(
  ($true != (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck)) | ~spl6_7),
  inference(trivial_inequality_removal,[status(thm)],[f227])).
thf(f227,plain,(
  ($true != $true) | ($true != (s__instance @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__PayCheck)) | ~spl6_7),
  inference(superposition,[status(thm)],[f183,f208])).
thf(f208,plain,(
  ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK4 @ s__USG @ s__Bill @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f207])).
thf(f207,plain,(
  ($true != $true) | ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK4 @ s__USG @ s__Bill @ cw)))),
  inference(forward_demodulation,[status(thm)],[f206,f79])).
thf(f206,plain,(
  ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK4 @ s__USG @ s__Bill @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f205])).
thf(f205,plain,(
  ($true != $true) | ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != (s__agent @ sK0 @ s__Bill @ cw))),
  inference(forward_demodulation,[status(thm)],[f204,f78])).
thf(f204,plain,(
  ($true != (s__instance @ sK0 @ s__Flying)) | ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != (s__agent @ sK0 @ s__Bill @ cw))),
  inference(trivial_inequality_removal,[status(thm)],[f203])).
thf(f203,plain,(
  ($true != (s__instance @ sK0 @ s__Flying)) | ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != $true)),
  inference(forward_demodulation,[status(thm)],[f202,f80])).
thf(f202,plain,(
  ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true != (s__instance @ sK0 @ s__Flying))),
  inference(trivial_inequality_removal,[status(thm)],[f199])).
thf(f199,plain,(
  ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ cw)) @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true != $true) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true != (s__instance @ sK0 @ s__Flying))),
  inference(superposition,[status(thm)],[f136,f81])).
thf(f136,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ X2 @ X0)) | ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ X0)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ X0)) @ (sK4 @ s__USG @ s__Bill @ X0))) | ($true != (s__instance @ X2 @ s__Moon)) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X1 @ s__Flying))) )),
  inference(trivial_inequality_removal,[status(thm)],[f135])).
thf(f135,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__instance @ X2 @ s__Moon)) | ($true != $true) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ($true = (s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ X0)) @ (sK2 @ (sK4 @ s__USG @ s__Bill @ X0)) @ (sK4 @ s__USG @ s__Bill @ X0))) | ($true != (s__destination @ X1 @ X2 @ X0)) | ($true != (s__instance @ X1 @ s__Flying))) )),
  inference(superposition,[status(thm)],[f87,f92])).
thf(f87,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3)) | ($true = (s__patient @ (sK3 @ X3) @ (sK2 @ X3) @ X3)) | ((s__agent @ X1 @ s__Bill @ X2) != $true) | ($true != (s__instance @ X1 @ s__Flying)) | ((s__destination @ X1 @ X0 @ X2) != $true) | ((s__instance @ X0 @ s__Moon) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f61])).
thf(f183,plain,(
  ( ! [X0 : $i] : (((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true)) ) | ~spl6_7),
  inference(avatar_component_clause,[status(thm)],[f182])).
thf(f182,definition,(
  spl6_7 <=> ! [X0 : $i] : (((s__instance @ X0 @ s__PayCheck) != $true) | ((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true))),
  introduced(definition,[new_symbols(naming,[spl6_7])],[avatar_definition])).
thf(f198,plain,(
  spl6_7 | spl6_9),
  inference(avatar_split_clause,[status(thm)],[f193,f195,f182])).
thf(f193,plain,(
  ( ! [X0 : $i] : (((s__instance @ X0 @ s__PayCheck) != $true) | ($true = (s__origin @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))) | ((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f192])).
thf(f192,plain,(
  ( ! [X0 : $i] : (($true = (s__origin @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))) | ((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true) | ($true != $true) | ((s__instance @ X0 @ s__PayCheck) != $true)) )),
  inference(forward_demodulation,[status(thm)],[f191,f116])).
thf(f116,plain,(
  ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f115])).
thf(f115,plain,(
  ($true != $true) | ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving))),
  inference(forward_demodulation,[status(thm)],[f114,f80])).
thf(f114,plain,(
  ($true != (s__instance @ sK1 @ s__Moon)) | ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f113])).
thf(f113,plain,(
  ($true != (s__instance @ sK1 @ s__Moon)) | ($true != $true) | ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving))),
  inference(forward_demodulation,[status(thm)],[f112,f78])).
thf(f112,plain,(
  ($true != (s__instance @ sK0 @ s__Flying)) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f111])).
thf(f111,plain,(
  ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving)) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true != (s__instance @ sK0 @ s__Flying)) | ($true != $true)),
  inference(forward_demodulation,[status(thm)],[f110,f79])).
thf(f110,plain,(
  ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true != (s__instance @ sK0 @ s__Flying)) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving))),
  inference(trivial_inequality_removal,[status(thm)],[f109])).
thf(f109,plain,(
  ($true != $true) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving)) | ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true != (s__instance @ sK0 @ s__Flying))),
  inference(superposition,[status(thm)],[f106,f81])).
thf(f106,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ X2 @ X0)) | ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__Giving)) | ($true != (s__instance @ X1 @ s__Flying)) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X2 @ s__Moon))) )),
  inference(trivial_inequality_removal,[status(thm)],[f105])).
thf(f105,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__instance @ X1 @ s__Flying)) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ($true = (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__Giving)) | ($true != (s__instance @ X2 @ s__Moon)) | ($true != $true) | ($true != (s__destination @ X1 @ X2 @ X0))) )),
  inference(superposition,[status(thm)],[f84,f92])).
thf(f84,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3)) | ($true != (s__instance @ X1 @ s__Flying)) | ((s__instance @ X0 @ s__Moon) != $true) | ((s__destination @ X1 @ X0 @ X2) != $true) | ($true = (s__instance @ (sK3 @ X3) @ s__Giving)) | ((s__agent @ X1 @ s__Bill @ X2) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f61])).
thf(f191,plain,(
  ( ! [X0 : $i] : (($true = (s__origin @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))) | ($true != (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving)) | ((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true) | ((s__instance @ X0 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f190])).
thf(f190,plain,(
  ( ! [X0 : $i] : (($true != (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving)) | ((s__instance @ X0 @ s__PayCheck) != $true) | ($true != $true) | ((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true) | ($true = (s__origin @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))))) )),
  inference(superposition,[status(thm)],[f132,f144])).
thf(f144,plain,(
  ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f143])).
thf(f143,plain,(
  ($true != $true) | ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))),
  inference(forward_demodulation,[status(thm)],[f142,f78])).
thf(f142,plain,(
  ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != (s__instance @ sK0 @ s__Flying))),
  inference(trivial_inequality_removal,[status(thm)],[f141])).
thf(f141,plain,(
  ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != (s__instance @ sK0 @ s__Flying)) | ($true != $true)),
  inference(forward_demodulation,[status(thm)],[f140,f79])).
thf(f140,plain,(
  ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true != (s__instance @ sK0 @ s__Flying)) | ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f139])).
thf(f139,plain,(
  ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != (s__instance @ sK0 @ s__Flying)) | ($true != $true)),
  inference(forward_demodulation,[status(thm)],[f138,f80])).
thf(f138,plain,(
  ($true != (s__instance @ sK1 @ s__Moon)) | ($true != (s__instance @ sK0 @ s__Flying)) | ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw)))),
  inference(trivial_inequality_removal,[status(thm)],[f137])).
thf(f137,plain,(
  ($true != (s__agent @ sK0 @ s__Bill @ cw)) | ($true != (s__instance @ sK1 @ s__Moon)) | ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) | ($true != $true) | ($true != (s__instance @ sK0 @ s__Flying))),
  inference(superposition,[status(thm)],[f130,f81])).
thf(f130,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ X2 @ X0)) | ((s__agent @ X1 @ s__Bill @ X0) != $true) | ($true != (s__instance @ X1 @ s__Flying)) | ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ X0))) | ($true != (s__instance @ X2 @ s__Moon))) )),
  inference(trivial_inequality_removal,[status(thm)],[f129])).
thf(f129,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ X2 @ X0)) | ($true != (s__instance @ X1 @ s__Flying)) | ($true != (s__instance @ X2 @ s__Moon)) | ($true != $true) | ($true = (s__destination @ (sK3 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ X0))) | ((s__agent @ X1 @ s__Bill @ X0) != $true)) )),
  inference(superposition,[status(thm)],[f86,f92])).
thf(f86,plain,(
  ( ! [X2 : w,X3 : w,X0 : $i,X1 : $i] : (($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X3)) | ($true = (s__destination @ (sK3 @ X3) @ s__Bill @ X3)) | ((s__destination @ X1 @ X0 @ X2) != $true) | ((s__agent @ X1 @ s__Bill @ X2) != $true) | ((s__instance @ X0 @ s__Moon) != $true) | ($true != (s__instance @ X1 @ s__Flying))) )),
  inference(cnf_transformation,[status(thm)],[f61])).
thf(f132,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ X2 @ s__PayCheck) != $true) | ($true = (s__origin @ (sK5 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ X0))) | ($true != (s__patient @ X1 @ X2 @ X0)) | ((s__instance @ X1 @ s__Giving) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f131])).
thf(f131,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true = (s__origin @ (sK5 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__Bill @ (sK4 @ s__USG @ s__Bill @ X0))) | ($true != (s__patient @ X1 @ X2 @ X0)) | ($true != $true) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ X2 @ s__PayCheck) != $true) | ((s__instance @ X1 @ s__Giving) != $true)) )),
  inference(superposition,[status(thm)],[f101,f92])).
thf(f101,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X1 @ X3)) | ($true != (s__destination @ X2 @ s__Bill @ X1)) | ($true = (s__origin @ (sK5 @ X3) @ s__Bill @ X3)) | ($true != (s__patient @ X2 @ X0 @ X1)) | ($true != (s__instance @ X2 @ s__Giving)) | ((s__instance @ X0 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f66])).
thf(f66,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__instance @ X0 @ s__PayCheck) != $true) | ($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X1 @ X3)) | (($true = (s__instance @ (sK5 @ X3) @ s__Tax)) & ($true = (s__origin @ (sK5 @ X3) @ s__Bill @ X3))) | ($true != (s__destination @ X2 @ s__Bill @ X1)) | ($true != (s__patient @ X2 @ X0 @ X1)) | ($true != (s__instance @ X2 @ s__Giving)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f64,f65])).
thf(f65,plain,(
  ! [X3 : w] : (? [X4 : $i] : (($true = (s__instance @ X4 @ s__Tax)) & ($true = (s__origin @ X4 @ s__Bill @ X3))) => (($true = (s__instance @ (sK5 @ X3) @ s__Tax)) & ($true = (s__origin @ (sK5 @ X3) @ s__Bill @ X3))))),
  introduced(definition,[],[choice_axiom])).
thf(f64,plain,(
  ! [X0 : $i,X1 : w,X2 : $i,X3 : w] : (((s__instance @ X0 @ s__PayCheck) != $true) | ($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X1 @ X3)) | ? [X4 : $i] : (($true = (s__instance @ X4 @ s__Tax)) & ($true = (s__origin @ X4 @ s__Bill @ X3))) | ($true != (s__destination @ X2 @ s__Bill @ X1)) | ($true != (s__patient @ X2 @ X0 @ X1)) | ($true != (s__instance @ X2 @ s__Giving)))),
  inference(rectify,[status(thm)],[f50])).
thf(f50,plain,(
  ! [X0 : $i,X2 : w,X3 : $i,X1 : w] : (((s__instance @ X0 @ s__PayCheck) != $true) | ($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X1)) | ? [X4 : $i] : (($true = (s__instance @ X4 @ s__Tax)) & ((s__origin @ X4 @ s__Bill @ X1) = $true)) | ($true != (s__destination @ X3 @ s__Bill @ X2)) | ($true != (s__patient @ X3 @ X0 @ X2)) | ($true != (s__instance @ X3 @ s__Giving)))),
  inference(flattening,[status(thm)],[f49])).
thf(f49,plain,(
  ! [X0 : $i,X3 : $i,X2 : w,X1 : w] : ((? [X4 : $i] : (($true = (s__instance @ X4 @ s__Tax)) & ((s__origin @ X4 @ s__Bill @ X1) = $true)) | ($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X1))) | (($true != (s__instance @ X3 @ s__Giving)) | ((s__instance @ X0 @ s__PayCheck) != $true) | ($true != (s__destination @ X3 @ s__Bill @ X2)) | ($true != (s__patient @ X3 @ X0 @ X2))))),
  inference(ennf_transformation,[status(thm)],[f41])).
thf(f41,plain,(
  ! [X0 : $i,X3 : $i,X2 : w,X1 : w] : ((($true = (s__instance @ X3 @ s__Giving)) & ((s__instance @ X0 @ s__PayCheck) = $true) & ($true = (s__destination @ X3 @ s__Bill @ X2)) & ($true = (s__patient @ X3 @ X0 @ X2))) => (($true = (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X1)) => ? [X4 : $i] : (($true = (s__instance @ X4 @ s__Tax)) & ((s__origin @ X4 @ s__Bill @ X1) = $true))))),
  inference(fool_elimination,[status(thm)],[f40])).
thf(f40,plain,(
  ! [X0 : $i,X1 : w,X2 : w,X3 : $i] : (((s__instance @ X3 @ s__Giving) & (s__patient @ X3 @ X0 @ X2) & (s__destination @ X3 @ s__Bill @ X2) & (s__instance @ X0 @ s__PayCheck)) => ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X2 @ X1) => ? [X4 : $i] : ((s__origin @ X4 @ s__Bill @ X1) & (s__instance @ X4 @ s__Tax))))),
  inference(rectify,[status(thm)],[f26])).
thf(f26,axiom,(
  ! [X6 : $i,X3 : w,X0 : w,X5 : $i] : (((s__instance @ X5 @ s__Giving) & (s__patient @ X5 @ X6 @ X0) & (s__destination @ X5 @ s__Bill @ X0) & (s__instance @ X6 @ s__PayCheck)) => ((accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X0 @ X3) => ? [X7 : $i] : ((s__origin @ X7 @ s__Bill @ X3) & (s__instance @ X7 @ s__Tax))))),
  file('/tmp/P_eXfui2V5/SOT_1oz6Aa',obligation1_bill_taxed)).
thf(f184,plain,(
  spl6_6 | spl6_7),
  inference(avatar_split_clause,[status(thm)],[f176,f182,f178])).
thf(f176,plain,(
  ( ! [X0 : $i] : (((s__instance @ X0 @ s__PayCheck) != $true) | ((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true) | ($true = (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Tax))) )),
  inference(trivial_inequality_removal,[status(thm)],[f175])).
thf(f175,plain,(
  ( ! [X0 : $i] : (((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true) | ($true != $true) | ((s__instance @ X0 @ s__PayCheck) != $true) | ($true = (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Tax))) )),
  inference(forward_demodulation,[status(thm)],[f150,f116])).
thf(f150,plain,(
  ( ! [X0 : $i] : (((s__instance @ X0 @ s__PayCheck) != $true) | ((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true) | ($true = (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Tax)) | ($true != (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving))) )),
  inference(trivial_inequality_removal,[status(thm)],[f145])).
thf(f145,plain,(
  ( ! [X0 : $i] : (($true = (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ (sK4 @ s__USG @ s__Bill @ cw))) @ s__Tax)) | ((s__instance @ X0 @ s__PayCheck) != $true) | ($true != (s__instance @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ s__Giving)) | ($true != $true) | ((s__patient @ (sK3 @ (sK4 @ s__USG @ s__Bill @ cw)) @ X0 @ (sK4 @ s__USG @ s__Bill @ cw)) != $true)) )),
  inference(superposition,[status(thm)],[f118,f144])).
thf(f118,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (($true != (s__destination @ X1 @ s__Bill @ X0)) | ($true = (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__Tax)) | ((s__instance @ X1 @ s__Giving) != $true) | ($true != (s__patient @ X1 @ X2 @ X0)) | ((s__instance @ X2 @ s__PayCheck) != $true)) )),
  inference(trivial_inequality_removal,[status(thm)],[f117])).
thf(f117,plain,(
  ( ! [X2 : $i,X0 : w,X1 : $i] : (((s__instance @ X2 @ s__PayCheck) != $true) | ($true != $true) | ($true = (s__instance @ (sK5 @ (sK4 @ s__USG @ s__Bill @ X0)) @ s__Tax)) | ($true != (s__destination @ X1 @ s__Bill @ X0)) | ((s__instance @ X1 @ s__Giving) != $true) | ($true != (s__patient @ X1 @ X2 @ X0))) )),
  inference(superposition,[status(thm)],[f102,f92])).
thf(f102,plain,(
  ( ! [X2 : $i,X3 : w,X0 : $i,X1 : w] : (($true != (accreln3 @ s__confersObligation @ s__USG @ s__Bill @ X1 @ X3)) | ($true != (s__destination @ X2 @ s__Bill @ X1)) | ($true = (s__instance @ (sK5 @ X3) @ s__Tax)) | ($true != (s__instance @ X2 @ s__Giving)) | ($true != (s__patient @ X2 @ X0 @ X1)) | ((s__instance @ X0 @ s__PayCheck) != $true)) )),
  inference(cnf_transformation,[status(thm)],[f66])).
% SZS output end Proof for SOT_1oz6Aa
% (1248141)------------------------------
% (1248141)Version: Vampire 4.8 (commit b486e5b56 on 2026-03-23 16:09:11 +0100)
% (1248141)Termination reason: Refutation

% (1248141)Memory used [KB]: 5756
% (1248141)Time elapsed: 0.021 s
% (1248141)------------------------------
% (1248141)------------------------------
% (1248135)Success in time 0.033 s
% Vampire exiting

% END OF SYSTEM OUTPUT
% RESULT: SOT_1oz6Aa - Vampire---5.0.1 says Theorem - CPU = 0.02 WC = 0.08 
% OUTPUT: SOT_1oz6Aa - Vampire---5.0.1 says Refutation - CPU = 0.02 WC = 0.08 

